<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2.4.6.4 Originally 5.3.3 Odd and even CAP W1 Q10 (used) edited</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Give an example of an odd function. \(f_1(x)=\)
            </span>
            <span class="multilang" lang="no">
                Gi et eksempel av en odde funksjon. \(f_1(x)=\)
            </span>
            [[input:ans1]]. [[validation:ans1]] [[feedback:odd]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Give an example of an even function. \(f_2(x)=\)
            </span>
            <span class="multilang" lang="no">
                Gi et eksempel p&#229; en jevn funksjon. \(f_2(x)=\)
            </span>
            [[input:ans2]]. [[validation:ans2]] [[feedback:even]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Give an example of a function which is odd and even. \(f_3(x)=\)
            </span>
            <span class="multilang" lang="no">
                Gi et eksempel p&#229; en funksjon som er b&#229;de jevn og odde. \(f_3(x)=\)
            </span>
            [[input:ans3]]. [[validation:ans3]] [[feedback:oddeven]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                True or false: the answer to 3. is unique.
            </span>
            <span class="multilang" lang="no">
                Sant eller usant: svaret til 3. er unik.
            </span>
            [[input:ans4]] [[validation:ans4]] [[feedback:unique]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        A function \(y=f(x)\) is odd if \(f(-x) = -f(x)\).&#160; A function \(y = f(x)\) is even if \(f(-x) = f(x)\).
    </span>
    <span class="multilang" lang="no">
        En funksjon \(y = f(x) \) er odde hvis \(f(-x) = -f(x) \). &amp;Nbsp; En funksjon \(y = f(x)\) er jevn om \(f(-x) = f(x)\).
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                For example, \(f_1(x) = x^3\) is an odd function, since
            </span>
            <span class="multilang" lang="no">
                For eksempel er \(f_1(x) = x^3\) en oddefunksjon, siden
            </span>
        </p>
        \[f_1(-x) = (-x)^3 = -x^3 = -f_1(x).\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For example, \(f_2(x) = x^4\) is an even function, since
            </span>
            <span class="multilang" lang="no">
                For eksempel, \(f_2(x) = x^4\) er en jevn funksjon, since
            </span>
        </p>
        \[f_2(-x) = (-x)^4 = x^4 = f_2(x).\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                A function that is both even and odd is \(f_3(x) = 0\), since
            </span>
            <span class="multilang" lang="no">
                En funksjon som er b&#229;de jevn og odde er \(f_3(x) = 0\), since
            </span>
        </p>
        \[f_3(-x) = 0 = f_3(x) = -f_3(x).\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                The function \(f_3(x) = 0\) is indeed the unique function that is both even and odd.
            </span>
            <span class="multilang" lang="no">
                Funksjonen \(f_3 (x) = 0 \) er den unike funksjonen som er b&#229;de jevn og oddetall.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                To see this, let \(f\) be a function that is even and odd and \(x\) be any point on the real line. Then since \(f\) is even we have \(f(-x) = f(x)\), and since \(f\) is odd, \(f(-x) = -f(x)\). Therefore, we must have \(f(x) = -f(x)\), which forces \(f(x) = 0\).
            </span>
            <span class="multilang" lang="no">
                For &#229; se dette, la \(f \) v&#230;re en funksjon som er jevn og odde og \(x \) v&#230;re hvilket som helst punkt p&#229; den reelle linjen. Siden \(f\) er jevn, har vi \(f(-x) = f(x)\), og siden \(f\) er odde, \(f(-x) = -f(x)\). Derfor m&#229; vi ha \(f (x) = -f (x) \), som tvinger \(f (x) = 0 \).
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>5.0000000</defaultgrade>
        <penalty>0.3333333</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text/>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text/>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>x^3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>x^4</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>0</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>boolean</type>
            <tans>true</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>even</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>sa:ans2-subst(x=-x,ans2);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>sa</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>even-0-T </trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>even-0-F </falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is not an even function:\[ f(x)-f(-x)={@sa@} \neq 0.\]
    </span>
    <span class="multilang" lang="no">
        Svaret ditt er ikke en jevn funksjon:\[ f(x)-f(-x)={@sa@} \neq 0.\]
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>odd</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>sa:subst(x=-x,ans1)+ans1;</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>sa</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>odd-0-T </trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>odd-0-F </falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is not an odd function: \[ f(x)+f(-x)={@sa@} \neq 0.\]
    </span>
    <span class="multilang" lang="no">
        Svaret ditt er ikke en odde funksjon: \[ f(x)+f(-x)={@sa@} \neq 0.\]
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>oddeven</name>
            <value>2.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>sa1:subst(x=-x,ans3)+ans3;
 sa2:ans3-subst(x=-x,ans3);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>sa1</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>ODD</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>oddeven-0-F </falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is not an odd function:\[ f(x)+f(-x)={@sa1@} \neq 0.\]
    </span>
    <span class="multilang" lang="no">
        Svaret ditt er ikke en odde funksjon:\[ f(x)+f(-x)={@sa1@} \neq 0.\]
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>sa2</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>EVEN</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>+</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>oddeven-1-F </falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is not an even function:\[ f(x)-f(-x)={@sa2@} \neq 0.\]
    </span>
    <span class="multilang" lang="no">
        Svaret ditt er ikke en jevn funksjon:\[ f(x)-f(-x)={@sa2@} \neq 0.\]
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>unique</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>unique-0-T </trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>unique-0-F </falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>x^3</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>cos(x)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>0</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>true</value>
            </testinput>
            <expected>
                <name>even</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>even-0-T</expectedanswernote>
            </expected>
            <expected>
                <name>odd</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>odd-0-T</expectedanswernote>
            </expected>
            <expected>
                <name>oddeven</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>EVEN</expectedanswernote>
            </expected>
            <expected>
                <name>unique</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>unique-0-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>x^2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>x^3</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>x^3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>false</value>
            </testinput>
            <expected>
                <name>even</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.3333333</expectedpenalty>
                <expectedanswernote>even-0-F</expectedanswernote>
            </expected>
            <expected>
                <name>odd</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.3333333</expectedpenalty>
                <expectedanswernote>odd-0-F</expectedanswernote>
            </expected>
            <expected>
                <name>oddeven</name>
                <expectedscore>0.5000000</expectedscore>
                <expectedpenalty>0.3333333</expectedpenalty>
                <expectedanswernote>oddeven-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>unique</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.3333333</expectedpenalty>
                <expectedanswernote>unique-0-F</expectedanswernote>
            </expected>
        </qtest>
    </question>


</quiz>