<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2.4.4.4 Frequency from period</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        A periodic function has period \({\displaystyle{@float(pp)@}}\) seconds. How many times will the pattern in the graph repeat over an interval of {@tt@} seconds?
    </span>
    <span class="multilang" lang="no">
        En periodisk funksjon har periode \({\displaystyle{@float(pp)@}} \) sekunder. Hvor mange ganger vil m&#248;nsteret i grafen gjenta seg over et intervall p&#229; {@tt@} sekunder?
    </span>
</p>
<p>
    \( N = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The periodic function has period \({\displaystyle{@float(pp)@}}\) seconds.
    </span>
    <span class="multilang" lang="no">
        Den periodiske funksjonen har perioden \({\displaystyle{@float(pp)@}}\) sekunder.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The frequency (the number of repetitions per time unit) is \( f = \frac{1}{T} = {@1/pp@}\, \textrm{s}^{-1} = {@1/pp@}\, \textrm{Hz}.\)
    </span>
    <span class="multilang" lang="no">
        Frekvensen (antall repetisjoner per tidsenhet) er \(f = \frac{1}{T} = {@1/pp@} \, \textrm{s}^{-1} = {@1/pp @} \, \textrm{Hz}. \)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        So in the time interval of \( {@tt@} \,\textrm{s}\) the pattern repeats \( N = {@1/pp@}\, \textrm{s}^{-1} \times {@tt@} \, \textrm{s} = {@Ans@}\) times.
    </span>
    <span class="multilang" lang="no">
        S&#229; i tidsintervallet for \({@tt@} \, \textrm{s} \) gjentas m&#248;nsteret \(N = {@1/pp@} \, \textrm{s}^{-1} \times {@tt@} \, \textrm{s} = {@Ans@}\) ganger.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>pp: 10^-(1+rand(3));
tt: rand([10,20,30,40,50,50,60,80]);

Ans: tt/pp;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\(T ={@pp@}, t={@tt@}, N = {@Ans@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>Ans</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans1</sans>
                <tans>Ans</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>474454178</deployedseed>
        <deployedseed>2103682428</deployedseed>
        <deployedseed>26316965</deployedseed>
        <deployedseed>771985559</deployedseed>
        <deployedseed>748981128</deployedseed>
        <deployedseed>745917705</deployedseed>
        <deployedseed>398645102</deployedseed>
        <deployedseed>1867571400</deployedseed>
        <deployedseed>68311927</deployedseed>
        <deployedseed>1038670467</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>Ans</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>pp*tt</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>