<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>4.3.4.1 Projectile</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        A projectile is fired from the ground with an initial speed \(u\ \mathrm{ms}^{-1}\) at an angle of elevation \(\alpha^\circ\). If air resistance is neglected, the vertical height, \(y\ \mathrm{m}\), is related to the horizontal distance, \(x\ \mathrm{m}\), by the equation
    </span>
    <span class="multilang" lang="no">
        Et prosjektil skytes fra bakken med en starthastighet \(u\ \mathrm{ms}^{-1}\) og vinkel p&#229; \(\alpha^\circ\). Hvis luftmotstanden neglisjeres, er den vertikale h&#248;yden, \(y\ \mathrm{m}\), relatert til den horisontale avstanden, \(x\ \mathrm{m}\), gitt ved ligningen
    </span>
</p>
\[y= x\tan\alpha -\dfrac{gx^2{\rm sec}^2\alpha}{2u^2},\]
<p>
    <span class="multilang" lang="en">
        where \(g\ \mathrm{ms^{-2}}\) is the gravitational constant.
    </span>
    <span class="multilang" lang="no">
        hvor \(g\ \mathrm{ms^{-2}}\) er gravitasjonskonstanten.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Find the height of the projectile when it has not yet travelled any horizontal distance.
            </span>
            <span class="multilang" lang="no">
                Finn h&#248;yden p&#229; prosjektilet n&#229;r det enn&#229; ikke har tilbakelagt noen horisontal avstand.
            </span>
        </p>
        <p>
            \(y(x=0) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For what value of \(x\) other than \(x=0\) do we have \(y=0\)?
            </span>
            <span class="multilang" lang="no">
                For hvilken verdi av \(x\) annet enn \(x=0\) har vi \(y=0\)?
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Simplify your answer to get a function of \(\sin(2\alpha)\).
            </span>
            <span class="multilang" lang="no">
                Forenkle svaret ditt for &#229; f&#229; en funksjon av \(\sin(2\alpha)\).
            </span>
        </p>
        <p>
            \(x=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Find the value \(x_\text{max}\) for which the value of \(y\) would be a maximum.
            </span>
            <span class="multilang" lang="no">
                Finn verdien \(x_\text{maks}\) som gir st&#248;rst \(y\)-verdi.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Simplify your answer to get a function of \(\sin(2\alpha)\).
            </span>
            <span class="multilang" lang="no">
                Forenkle svaret ditt for &#229; f&#229; en funksjon av \(\sin(2\alpha)\).
            </span>
        </p>
        <p>
            \(x_\text{max}=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Simplify the expression for the maximum height to obtain a function of \(\sin \alpha\).
            </span>
            <span class="multilang" lang="no">
                Forenkle uttrykket for maksimal h&#248;yde for &#229; f&#229; en funksjon av \(\sin \alpha\).
            </span>
        </p>
        <p>
            \(y_\text{max} =\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Assuming \(u = {@u1@}\ \mathrm{ms}^{-1}\), \(\alpha = {@a1@}^{\circ}\) and \(g = 10\ \mathrm{ms}^{-2}\), find the two values \(x_1 \le x_2\) of \(x\) such that \(y = \displaystyle{{@y1@}}\).
            </span>
            <span class="multilang" lang="no">
                Forutsatt \(u = {@u1@}\ \mathrm{ms}^{-1}\), \(\alpha = {@a1@}^{\circ}\) og \(g = 10\ \mathrm {ms}^{-2}\), finn de to verdiene \(x_1 \le x_2\) av \(x\) slik at \(y = \displaystyle{{@y1@}}\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Give your answer with three significant figures.
            </span>
            <span class="multilang" lang="no">
                Gi svaret ditt med tre signifikante sifre.
            </span>
        </p>
        <p>
            \(x_1=\) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]
            <span class="multilang" lang="en">
                and
            </span>
            <span class="multilang" lang="no">
                og
            </span>
            \(x_2=\) [[input:ans6]] [[validation:ans6]] [[feedback:prt6]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Since \(x\) appears in both of the terms on the right-hand side, when \(x=0\), the right-hand side is zero. Hence, \(y(x = 0) = 0\).
            </span>
            <span class="multilang" lang="no">
                Siden \(x\) er tilstede i begge leddene p&#229; h&#248;yre side er h&#248;yre side null n&#229;r \(x=0\). Derfor er \(y(x = 0) = 0\).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                When \(y=0\), the equation can be written as
            </span>
            <span class="multilang" lang="no">
                N&#229;r \(y=0\), kan ligningen skrives som
            </span>
        </p>
        \[\dfrac{gx^2{\sec}^2\alpha}{2u^2}-x\tan\alpha=0.\]
        <p>
            <span class="multilang" lang="en">
                If \(x=0\) is excluded from consideration, we can divide through by \(x\) and rearrange to give
            </span>
            <span class="multilang" lang="no">
                Hvis vi antar at \(x\neq 0\), kan vi dele med \(x\) og omorganisere for &#229; f&#229;
            </span>
        </p>
        \[\dfrac{gx{\sec}^2\alpha}{2u^2}=\tan\alpha.\]
        <p>
            <span class="multilang" lang="en">
                To solve for \(x\) we need to multiply both sides by \(2u^2/g{\sec}^2\alpha\). Given that
            </span>
            <span class="multilang" lang="no">
                For &#229; l&#248;se for \(x\) m&#229; vi multiplisere begge sider med \(2u^2/g{\sec}^2\alpha\). Gitt at
            </span>
        </p>
        \[\begin{aligned} \frac{1}{\sec^2\alpha} &amp;\equiv \cos^2 \alpha, \\ \tan \alpha &amp;\equiv \sin\alpha \cos \alpha, \\ \sin 2 \alpha &amp;\equiv 2\sin\alpha \cos \alpha, \end{aligned}\]
        <p>
            <span class="multilang" lang="en">
                this results in
            </span>
            <span class="multilang" lang="no">
                dette resulterer med
            </span>
        </p>
        \[x=\dfrac{2u^2\sin\alpha\cos\alpha}{g}=\dfrac{u^2\sin 2\alpha}{g}\]
        <p>
            <span class="multilang" lang="en">
                This represents the maximum range.
            </span>
            <span class="multilang" lang="no">
                Dette representerer maksimal rekkevidde.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                As air resistance is neglected, the path of the projectile is a parabola, which is symmetrical about the vertical line through its highest point. Therefore, the maximum height will occur at half the maximum range, i.e. where
            </span>
            <span class="multilang" lang="no">
                Ettersom luftmotstanden blir neglisjert, er banen til prosjektilet en parabel, som er symmetrisk om den vertikale linjen gjennom det h&#248;yeste punktet. Derfor vil den maksimale h&#248;yden inntreffe ved halve maksimale rekkevidden, dvs. hvor
            </span>
        </p>
        \[x=\dfrac{u^2\sin 2\alpha}{2g}.\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Substituting this expression for \(x\) in the equation for \(y\) gives
            </span>
            <span class="multilang" lang="no">
                &#197; erstatte dette uttrykket med \(x\) i ligningen for \(y\) gir
            </span>
        </p>
        \[y=\left(\dfrac{u^2\sin2\alpha}{2g}\right)\tan\alpha-\left(\dfrac{u^2\sin2\alpha}{2g}\right)^2\dfrac{g{\sec}^2\alpha}{2u^2}.\]
        <p>
            <span class="multilang" lang="en">
                Using the same trigonometric identities as before,
            </span>
            <span class="multilang" lang="no">
                Ved &#229; bruke de samme trigonometriske identitetene som f&#248;r f&#229;r vi,
            </span>
        </p>
        \[\qquad y=\dfrac{u^2\sin^2\alpha}{g}-\dfrac{u^2\sin^2\alpha}{2g}=\dfrac{u^2\sin^2\alpha}{2g}.\]
        <p>
            <span class="multilang" lang="en">
                This represents the maximum height.
            </span>
            <span class="multilang" lang="no">
                Dette representerer maksimal h&#248;yde.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Substitution of \(u={@u1@}\), \(\alpha={@a1@}\), \(g=10\) and \(y={@y1@}\) in the original equation gives a quadratic:
            </span>
            <span class="multilang" lang="no">
                Substitusjon av \(u={@u1@}\), \(\alpha={@a1@}\), \(g=10\) og \(y={@y1@}\) i den opprinnelige ligningen gir et kvadrat:
            </span>
        </p>
        \[{@eq@} = 0\]
        <p>
            <span class="multilang" lang="en">
                Solving this quadratic yields \(x_1 = {@ta5@}\) and \(x_2 = {@ta6@}\) as the two horizontal ranges at which \(y = {@y1@}\).
            </span>
            <span class="multilang" lang="no">
                L&#248;sning av dette kvadratet gir \(x_1 = {@ta5@}\) og \(x_2 = {@ta6@}\) som de to horisontale omr&#229;dene der \(y = {@y1@}\).
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>fpprintprec : 3;
ta1 : 0;
ta2 : u^2*sin(2*alpha)/g;
ta2_dd : u^2*dd/g;
ta3 : u^2*sin(2*alpha)/(2*g);
ta3_dd : u^2*dd/(2*g);
ta4 : u^2*(sin(alpha))^2/(2*g);
ta4_ee : u^2*ee^2/(2*g);
u1: 5*(4+rand(4));
a1: 10*(6+rand(3));
ar : pi*a1/180;
y1 : u1/2;
g1 : 10;
aa : (g1*sec(ar)^2)/(2*u1^2);
bb : -tan(ar);
cc : y1;
eq: aa*x^2 + bb*x + cc;
D1 : bb^2 - 4*aa*cc;
ta5 : float((-bb - sqrt(D1))/(2*aa));
ta6 : float((-bb + sqrt(D1))/(2*aa)); </text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( u = {@u1@}, \alpha = {@a1@} \implies x_1 = {@ta5@}, x_2 = {@ta6@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>numerical</type>
            <tans>ta5</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans6</name>
            <type>numerical</type>
            <tans>ta6</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ans2_dd:subst(sin(2*alpha)=dd,ans2)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2_dd</sans>
                <tans>ta2_dd</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer can be simplified.
    </span>
    <span class="multilang" lang="no">
        Svaret ditt kan forenkles.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ans3_dd:subst(sin(2*alpha)=dd,ans3)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3_dd</sans>
                <tans>ta3_dd</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer can be simplified.
    </span>
    <span class="multilang" lang="no">
        Svaret ditt kan forenkles.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ans4_ee:subst(sin(alpha)=ee,ans4)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4_ee</sans>
                <tans>ta4_ee</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Your answer can be simplified.
</span>
<span class="multilang" lang="no">
    Svaret ditt kan forenkles.
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumSigFigs</answertest>
                <sans>ans5</sans>
                <tans>ta5</tans>
                <testoptions>3</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumSigFigs</answertest>
                <sans>ans6</sans>
                <tans>ta6</tans>
                <testoptions>3</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1631184873</deployedseed>
        <deployedseed>1457071667</deployedseed>
        <deployedseed>1470256475</deployedseed>
        <deployedseed>2084921921</deployedseed>
        <deployedseed>1621182874</deployedseed>
        <deployedseed>83065210</deployedseed>
        <deployedseed>1781346533</deployedseed>
        <deployedseed>1390297486</deployedseed>
        <deployedseed>1155164518</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta5</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ta6</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt6-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>(2*cos(alpha)*sin(alpha)*u^2)/g</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>(cos(alpha)*sin(alpha)*u^2)/g</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4+1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta5+1</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ta6+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt6-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>