<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>9.3.2.3  The scalar product</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If \(\underline{ r}={@r1@}\underline{ i}+{@r2@}\underline{ j}+{@r3@}\underline{ k}\),&#160;find \(\underline{ r}\cdot\underline{ r}\)
    </span>
    <span class="multilang" lang="no">
        Hvis \(\underline{ r}={@r1@}\underline{ i}+{@r2@}\underline{ j}+{@r3@}\underline{ k}\),&#160;bestem \(\underline{r}\cdot\underline{r}\)
    </span>
    .
</p>
<p>
    [[input:ans1]] [[validation:ans1]]
</p>
<p>
    [[feedback:prt1]]
</p>
<p>
    <span class="multilang" lang="en">
        Now compute \(|\underline{ r}|^2\).
    </span>
    <span class="multilang" lang="no">
        Beregn n&#229; \(|\underline{ r}|^2\).
    </span>
</p>
<p>
    [[input:ans2]] [[validation:ans2]]
</p>
<p>
    [[feedback:prt2]]
    <br/>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We use the formula given in a previous Key point:
    </span>
    <span class="multilang" lang="no">
        Vi bruker formelen gitt i et tidligere n&#248;kkelpunkt:
    </span>
</p>
<p>
    \({\underline r}\cdot {\underline r}=({@r1@}{\underline i}+{@r2@}{\underline j}+{@r3@}{\underline k})\cdot({@r1@}{\underline i}+{@r2@}{\underline j}+{@r3@}{\underline k})=({@r1@})({@r1@})+({@r2@})({@r2@})+({@r3@})({@r3@})={@ta@}\).
</p>
<p>
    <span class="multilang" lang="en">
        Then:
    </span>
    <span class="multilang" lang="no">
        Videre:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(|\underline{r}|=\sqrt{{@r1sq@}+{@r2sq@}+{@r3sq@}}=\sqrt{{@ta@}}\), hence&#160;\(|\underline{r}|^2= {@ta@}\).
    </span>
    <span class="multilang" lang="no">
        \(|\underline{r}|=\sqrt{{@r1sq@}+{@r2sq@}+{@r3sq@}}=\sqrt{{@ta@}}\), s&#229;&#160;\(|\underline{r}|^2= {@ta@}\).
    </span>
</p>
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Thus, illustrating that:&#160;\(|\underline{r}|^2=\underline{r}\cdot\underline{r}\).
    </span>
    <span class="multilang" lang="no">
        Dermed illustrerer at:&#160;\(|\underline{r}|^2=\underline{r}\cdot\underline{r}\).
    </span>
</p>
<p>
</p>
<p>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022060100</text>
        </stackversion>
        <questionvariables>
            <text>r1: 2+rand(3);
r1sq: r1^2;
r2: 2+rand(3);
r2sq: r2^2;
r3: 2+rand(3);
r3sq: r3^2;
ta: r1*r1 + r2*r2 + r3*r3;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[ {@r1@}\underline{ i}+{@r2@}\underline{ j}+{@r3@}\underline{ k} \cdot {@r1@}\underline{ i}+{@r2@}\underline{ j}+{@r3@}\underline{ k} = {@ta@}  \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1773321182</deployedseed>
        <deployedseed>273952110</deployedseed>
        <deployedseed>1332435204</deployedseed>
        <deployedseed>681954021</deployedseed>
        <deployedseed>1982775636</deployedseed>
        <deployedseed>1273883291</deployedseed>
        <deployedseed>1556040258</deployedseed>
        <deployedseed>1852471241</deployedseed>
        <deployedseed>2001340123</deployedseed>
        <deployedseed>685102625</deployedseed>
        <deployedseed>101716277</deployedseed>
        <deployedseed>637076689</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta-1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>