<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>9.3.4.7 Work done by a force</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The work done by a force \(\underline{F}\) in moving a body through a displacement \(\underline{r}\) is given by \(\underline{F}\cdot\underline{r}\). Find the work done by the force \(\underline{F}=3\underline{i}+{@f2@}\underline{k}\) if it causes a body to move from the point with coordinates \((1,1,2)\) to the point \(({@p2x@},3,5)\).
    </span>
    <span class="multilang" lang="no">
        Arbeidet som utf&#248;res av en kraft \(\underline{F}\) for &#229; bevege et legeme \(\underline{r}\) er gitt ved \(\underline{F}\cdot\underline{r}\) . Finn arbeidet utf&#248;rt av kraften \(\underline{F}=3\underline{i}+{@f2@}\underline{k}\) hvis den f&#229;r et legeme til &#229; bevege seg fra punktet med koordinatene \((1) ,1,2)\) til punktet \(({@p2x@},3,5)\).
    </span>
</p>
<p>
    [[input:ans1]]
    <span class="multilang" lang="en">
        units
    </span>
    <span class="multilang" lang="no">
        enheter
    </span>
    [[validation:ans1]]
</p>
<p>
    [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        First, we proceed to compute the displacement vector \(\underline{r}\) subtracting the final point \(({@p2x@},3,5)\) minus the initial point&#160;\((1,1,2)\):
    </span>
    <span class="multilang" lang="no">
        F&#248;rst beregner vi forskyvningsvektoren \(\underline{r}\) ved &#229; trekke fra sluttpunktet \(({@p2x@},3,5)\) minus startpunktet&#160;\((1,1,2) )\):
    </span>
</p>
<p>
    \(\underline{r} =&#160;({@p2x@},3,5) -&#160;(1,1,2)&#160;= {@p2x -1@}\underline{i} + 2\underline{j} + 3\underline{k}\)
</p>
<p>
    <span class="multilang" lang="en">
        Then, we proceed to compute the work done by the force \(\underline{F}=3\underline{i}+{@f2@}\underline{k}\) using the formula given in the statement:
    </span>
    <span class="multilang" lang="no">
        Deretter fortsetter vi ved &#229; beregne arbeidet utf&#248;rt av kraften \(\underline{F}=3\underline{i}+{@f2@}\underline{k}\) ved &#229; bruke formelen gitt i setningen:
    </span>
</p>
<p>
    \(W = \underline{F}\cdot\underline{r} = (3)({@p2x -1@})+(0)(2)+({@f2@})(3) = {@ta@} \) units.
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022060100</text>
        </stackversion>
        <questionvariables>
            <text>f2: 3 + rand(4)
p2x: 3 + rand(4)
ta: 3*(p2x -1) + f2*3
</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\(W = (3\underline{i}+{@f2@}\underline{k}) \cdot (({@p2x@}-1)\underline{i} + 2 \underline{j} + 3 \underline{k}) = {@ta@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1596035697</deployedseed>
        <deployedseed>428490799</deployedseed>
        <deployedseed>24406443</deployedseed>
        <deployedseed>1578315134</deployedseed>
        <deployedseed>1013805574</deployedseed>
        <deployedseed>108770404</deployedseed>
        <deployedseed>1460721703</deployedseed>
        <deployedseed>223450315</deployedseed>
        <deployedseed>526907617</deployedseed>
        <deployedseed>492674192</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>20</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>