<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>7.4.4.2 Inverse of a 3 x 3 matrix &#9472; determinant method</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Use the determinant method to find the inverse of the
                following matrices
    </span>
    <span class="multilang" lang="no">
        Bruk determinantmetoden for &#229; finne inversen til f&#248;lgende matriser.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \[A = {@A@},\] \[B = {@B@}.\]
    </span>
    <span class="multilang" lang="no">
        \[A = {@A@},\] \[B = {@B@}.\]
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <div>
            \(A^{-1} = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </div>
    </li>
    <li>
        <div>
            \(B^{-1} = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </div>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        To find the inverse of \(A = {@A@}\) using the determinant method will require four steps.
    </span>
    <span class="multilang" lang="no">
        For &#229; finne inversen av \(A = {@A@}\) ved hjelp av determinantmetoden vil det kreve fire trinn.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        First we find the determinant of \(A\).
    </span>
    <span class="multilang" lang="no">
        F&#248;rst finner vi determinanten til \(A\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \[ \det(A) = {@detA@}. \]
    </span>
    <span class="multilang" lang="no">
        \[ \det(A) = {@detA@}. \]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Since \(\det(A) \neq 0\) the inverse of \(A\) exists.
    </span>
    <span class="multilang" lang="no">
        Siden \(\det(A) \neq 0\) finnes inversen til \(A\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Now, we replace each element of \(A\) by its cofactor.
    </span>
    <span class="multilang" lang="no">
        N&#229; erstatter vi hvert element i \(A\) med dets kofaktor.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The cofactor of \({A}_{ij}\) is the determinant of the matrix that remains when row \(i\) and column \(j\) are
                removed, multiplied by \((-1)^{i+j}\). So we have,
    </span>
    <span class="multilang" lang="no">
        Kofaktoren til \({A}_{ij}\) er determinanten til matrisen som gjenst&#229;r n&#229;r rad \(i\) og kolonne \(j\) fjernes, multiplisert med \((-1)^{i+j}\). S&#229; vi har,
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        {@cofactorsA@}.
    </span>
    <span class="multilang" lang="no">
        {@cofactorsA@}.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Then we transpose the result to obtain the adjoint matrix:
    </span>
    <span class="multilang" lang="no">
        S&#229; transponerer vi resultatet for &#229; f&#229; den tilh&#248;rende adjungerende matrisen:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(\mathrm{adj}(A) = {@cofactorsA@}^{T} = {@adjA@}\).
    </span>
    <span class="multilang" lang="no">
        \(\mathrm{adj}(A) = {@cofactorsA@}^{T} = {@adjA@}\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Finally we obtain \(A^{-1}\) using the following formula:
    </span>
    <span class="multilang" lang="no">
        Til slutt f&#229;r vi \(A^{-1}\) ved &#229; bruke f&#248;lgende formel:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(A^{-1} = \frac{1}{\det(A)} \mathrm{adj}(A) = \frac{1}{@detA@} {@adjA@} = {@ta1@}\).
    </span>
    <span class="multilang" lang="no">
        \(A^{-1} = \frac{1}{\det(A)} \mathrm{adj}(A) = \frac{1}{@detA@} {@adjA@} = {@ta1@}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The method can be used to determine \(B^{-1}\) similarly, which gives \( B^{-1} = {@ta2@}.\)
    </span>
    <span class="multilang" lang="no">
        Metoden kan brukes p&#229; samme m&#229;te for &#229; bestemme \(B^{-1}\), som gir \(B^{-1} = {@ta2@}.\)
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>A:matrix([rand(9)+2,rand(9)+1,0],[1,0,0],[rand(9)+1,rand(9)+1,rand(9)+1]);
ta1:invert(A);

B:matrix([rand_with_prohib(-1,1,[0]),rand_with_prohib(-1,1,[0]),rand_with_prohib(-1,1,[0])],[0,rand_with_prohib(-1,1,[0]),rand_with_prohib(-1,1,[0])],[0,0,rand_with_prohib(-1,1,[0])]);
ta2:invert(B);

/* Determinant method*/
detA:determinant(A);
cofactorsA:transpose(adjoint(A));
adjA:adjoint(A)

upper: matrix([1,n,n],[0,1,n],[0,0,n]);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[\bf A = {@A@}, B = {@B@}\]
\[A^{-1} = {@ta1@}, B^{-1} = {@ta2@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text>Correct answer, well done.</text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text>Your answer is partially correct.</text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text>Incorrect answer.</text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>ta1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>matrix</type>
            <tans>ta2</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>wronganswer1:ev(-A,simp);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>wronganswer1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        It appears you have given the matrix \(\bf -A\)! This is not the inverse matrix. Please try again!
        <br/>
    </span>
    <span class="multilang" lang="no">
        Beklager, jeg mente &#229; si at jeg ga deg matrisen \(\bf A\), men ikke den inverse matrisen. La meg pr&#248;ve igjen:
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>wronganswer2:ev(-B,simp);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>wronganswer2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        It appears you have given the matrix \(\bf -B\)! This is not the inverse matrix. Please try again!
        <br/>
    </span>
    <span class="multilang" lang="no">
        Det virker som om du har gitt matrisen \(\bf -B\)! Dette er ikke den inverse matrisen. Vennligst pr&#248;v igjen!
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>361301438</deployedseed>
        <deployedseed>1469324625</deployedseed>
        <deployedseed>171644071</deployedseed>
        <deployedseed>358966885</deployedseed>
        <deployedseed>818821470</deployedseed>
        <deployedseed>1536609802</deployedseed>
        <deployedseed>793837122</deployedseed>
        <deployedseed>254525453</deployedseed>
        <deployedseed>1140548974</deployedseed>
        <deployedseed>1437129438</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(-A,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(-B,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>