<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>3.3.9.4 Factorise quartic with repeated roots x=+-a</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the equation {@eqn@}:
    </span>
    <span class="multilang" lang="no">
        L&#248;s likningen {@eqn@}:
    </span>
</p>
\( x = \)&#160;[[input:ans1]] [[validation:ans1]]
<p>
</p>
<p>
    <b>
        <span class="multilang" lang="en">
            Give your answer as a set of precise numbers. Only provide the
            <u>
                distinct
            </u>
            solutions.
        </span>
        <span class="multilang" lang="no">
            Gi svaret ditt som et sett av presise tall. Gi bare de
            <u>
                distinkte
            </u>
            l&#248;sningene.
        </span>
    </b>
</p>
<p>
    [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Try viewing this as a quadratic function of \(x^2.\) We solve for \(x^2\), from which we can obtain \(x\).
    </span>
    <span class="multilang" lang="no">
        Pr&#248;v &#229; se dette som en kvadratisk funksjon av \(x^2.\) Vi l&#248;ser for \(x^2\), hvorfra vi kan f&#229; \(x\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We are given the equation {@eqn@}. Let \(y = x^2\). Then the equation becomes \[y^2 - {@2*aa^2*y@} + {@aa^4@} = 0.\] Recognise that this is equivalent to \((y - {@aa^2@})^2 = 0.\)
    </span>
    <span class="multilang" lang="no">
        Vi f&#229;r ligningen {@eqn@}. La \(y = x^2\). Deretter blir ligningen \[y^2 - {@2*aa^2*y@} + {@aa^4@} = 0.\] Gjenkjenne at dette tilsvarer \((y - {@aa^2@})^2 = 0.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore \(y = {@aa^2@}\) which implies \(x = \pm{@aa@}.\)
    </span>
    <span class="multilang" lang="no">
        Derfor er \(y = {@aa^2@}\) som inneb&#230;rer at \(x = \pm{@aa@}.\)
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand(3)+1;
ta: setify([aa,-aa]);

simp:true;
factor: x + aa;
%eqn: (x + aa)^2*(x - aa)^2;
eqn: x^4 - 2*aa^2*x^2 + aa^4 = 0;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\({@eqn@}, \quad sol: {@ta@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>{a,b,c,d}</syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0.1000000</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        A useful property to notice is that this quadratic can be viewed as a quadratic polynomial in \(x^2\), so we can solve {@eqn@} for \(x^2\) (by factorising or using the quadratic formula) which will give us two expressions for \(x^2\), and this expression gives us the values of \(x\) as desired.
    </span>
    <span class="multilang" lang="no">
        En nyttig egenskap &#229; legge merke til er at denne andregradslikningen kan sees p&#229; som et kvadratisk polynom i \(x^2\), s&#229; vi kan l&#248;se {@eqn@} for \(x^2\) (ved &#229; faktorisere eller bruke den kvadratiske formelen ) som vil gi oss to uttrykk for \(x^2\), og dette uttrykket gir oss verdiene til \(x\) som &#248;nsket.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>967783658</deployedseed>
        <deployedseed>2059533060</deployedseed>
        <deployedseed>778014411</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>{aa}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>{-aa}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>