<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>3.3.7.1 Solve cubic given a root I</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Verify that the given value \(x={@x1@}\)&#160;is a solution of the equation \({@eqn@} = 0\) and hence find all solutions:
    </span>
    <span class="multilang" lang="no">
        Bekreft at den gitte verdien \(x={@x1@}\)&#160;er en l&#248;sning av ligningen \({@eqn@} = 0\) og finn videre alle l&#248;sninger:
    </span>
</p>
<p>
    \( x = \)&#160;[[input:ans1]] [[validation:ans1]]&#160;&#160;[[feedback:prt1]]
</p>
<p>
    <b>
        <span class="multilang" lang="en">
            Give your answer as a set of precise numbers.
        </span>
        <span class="multilang" lang="no">
            Gi svaret ditt som et sett av presise tall.
        </span>
    </b>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Try taking the factor of \( ({@x-x1@}) \) out of \({@eqn@} = 0\) in order to be left with a quadratic function \( {@ratsimp((x-x2)*(x-x3))@} = 0 \) which you can solve using the quadratic formula to obtain all solutions {@ta@}.
    </span>
    <span class="multilang" lang="no">
        Pr&#248;v &#229; ta faktoren \( ({@x-x1@}) \) ut av \({@eqn@} = 0\) for &#229; sitte igjen med en kvadratisk funksjon \( {@ratsimp((x-x2)*(x-x3))@} = 0 \) som du kan l&#248;se ved &#229; bruke abc-formelen og f&#229; alle l&#248;sningene {@ta@}.
    </span>
    <br/>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand(3)+1;
bb: rand(3)+1;
cc: rand(4)+5;

x1: -aa;
x2: -bb+sqrt(cc);
x3: -bb-sqrt(cc);
ta: setify([x1,x2,x3]);

eqn1: x^3 + (2*bb - aa)*x^2 - (2*aa*bb - bb^2 + cc)*x - aa*bb^2 + aa*cc;
eqn: ratsimp((x-x1)*(x-x2)*(x-x3));</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\({@eqn@}=0,\)  sol: {{@x1@},{@x2@},{@x3@}}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>{a,b,c}</syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1946564725</deployedseed>
        <deployedseed>1731070056</deployedseed>
        <deployedseed>324123121</deployedseed>
        <deployedseed>1521231394</deployedseed>
        <deployedseed>339670746</deployedseed>
        <deployedseed>1633534430</deployedseed>
        <deployedseed>1585238934</deployedseed>
        <deployedseed>239446220</deployedseed>
        <deployedseed>1937162256</deployedseed>
        <deployedseed>1990880341</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>{x1,x2,x3}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>{x1,x2}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>{ev(x1+1,simp),x2,x3}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>