<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>8.2.2.3 Exercise 2x2 linear system no solutions</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Consider the following linear system of equations: \[\begin{matrix}
                \displaystyle{{@eqq1@}} \\
                {@eqq2@} \\
                \end{matrix}\]
    </span>
    <span class="multilang" lang="no">
        Betrakt f&#248;lgende line&#230;re likningssystem: \[\begin{matrix}
                \displaystyle{{@eqq1@}} \\
                {@eqq2@} \\
                \end{matrix}\]
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Consider the matrix form \(AX = B\) of the system and identify \(A\)
            </span>
            <span class="multilang" lang="no">
                Betrakt matriseformen \(AX = B\) av systemet og identifiser \(A\)
            </span>
        </p>
    </li>
    <div>
        \(A=\) [[input:ans1]] [[validation:ans1]]
    </div>
    <div>
        <br/>
    </div>
    <li>
        <p>
            <span class="multilang" lang="en">
                Find the determinant of \(A\):
            </span>
            <span class="multilang" lang="no">
                Finn determinanten til \(A\):
            </span>
        </p>
        <p>
            \(|A| =\) [[input:ans2]] [[validation:ans2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Determine if the system has no solution, a unique solution, or infinitely many solutions.
            </span>
            <span class="multilang" lang="no">
                Bestem om systemet ikke har noen l&#248;sning, en unik l&#248;sning eller uendelig mange l&#248;sninger.
            </span>
        </p>
    </li>
    <p>
        [[input:ans3]] [[validation:ans3]]
    </p>
    <li>
        <p>
            <span class="multilang" lang="en">
                If it has a unique solution or infinitely many solutions, these are:
            </span>
            <span class="multilang" lang="no">
                Hvis det har en unik l&#248;sning eller uendelig mange l&#248;sninger, s&#229; er disse:
            </span>
        </p>
    </li>
    <p>
        \(x =\) [[input:ans4]] [[validation:ans4]]
    </p>
    <p>
        \(y =\) [[input:ans5]] [[validation:ans5]]
    </p>
    <p>
        <span class="multilang" lang="en">
            Give all the solutions if there is more than one. If there are no solutions, fill the gaps for \(x \) and \(y \) with the word
            <b>
                null
            </b>
            .
        </span>
        <span class="multilang" lang="no">
            Gi alle l&#248;sningene hvis det er mer enn en. Hvis det ikke er noen l&#248;sninger, fyll plassene for \(x \) og \(y \) med ordet
            <b>
                null
            </b>
            .
        </span>
    </p>
</ol>
[[feedback:prt1]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                In matrix form \(AX = B\), the matrix \(A \) is \[ A = {@AA@}.\]
            </span>
            <span class="multilang" lang="no">
                I matriseform \(AX = B\), er matrisen \(A \) \[ A = {@AA@}.\]
            </span>
        </p>
    </li>
    <li>
        <p>
            \( \det A = {@determinant(AA)@} \).
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Since the determinant of \(A \) is \( 0 \), we already know that the system has either no solution or infinitely many solutions. Note that if we multiply the first equation by \({@kk@} \) we obtain \[ {@kk@} \cdot ({@lhs(eqq1)@}) = {@kk@} \cdot ({@rhs(eqq1)@})
                                                \Leftrightarrow {@expand(kk * lhs(eqq1))@} = {@kk * rhs(eqq1)@} \] which is in conflict with the second equation. Therefore, the system is incompatible, i.e., there are no solutions.
            </span>
            <span class="multilang" lang="no">
                Siden determinanten til \(A \) er \( 0 \), vet vi allerede at systemet enten ikke har noen l&#248;sning eller uendelig mange l&#248;sninger. Legg merke til at hvis vi multipliserer den f&#248;rste ligningen med \({@kk@} \) f&#229;r vi \[ {@kk@} \cdot ({@lhs(eqq1)@}) = {@kk@} \cdot ({@rhs(eqq1)@})
                                                \Leftrightarrow {@expand(kk * lhs(eqq1))@} = {@kk * rhs(eqq1)@} \] som er i konflikt med den andre ligningen. Derfor er systemet uforenlig, dvs. det er ingen l&#248;sninger.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We have already found that there are no solutions to this system. As indicated in the question, you should enter the word
                <b>
                    null
                </b>
                in both \(x \) and \(y \) as an answer to this part.
            </span>
            <span class="multilang" lang="no">
                Vi har allerede funnet at det ikke er noen l&#248;sninger til dette systemet. Som angitt i sp&#248;rsm&#229;let, b&#248;r du skrive inn ordet
                <b>
                    null
                </b>
                i b&#229;de \(x \) og \(y \) som svar p&#229; denne delen.
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/* Set an order in the variables */
ordergreat(x, y);
/* Generate non-zero random numbers. They will be the coefficients of the first equation */
aa : rand(6) * 2 - 9/2 ;
bb : rand(6) * 2 - 5  ;
cc : rand(6) * 2 - 5  ;
/* Generate a random number different than 0, 1. It will be the opposite coefficient by which we multiply the second equation*/
kk : rand(4) * 3 - 4;
/* Define equations */
eqq1 : aa * x + bb * y = cc;
eqq2: kk * aa * x + kk * bb * y = kk * cc + 1;

AA : matrix([aa, bb], [kk * aa, kk * bb]);
BB : matrix([cc], [kk * cc]);
detAA : determinant(AA);

/* Variables for dropdown list */
ta3:[["No solution",true],["One unique solution",false],["Infinitely many solutions",false]];
ta3:random_permutation(ta3);
ta3c: "No solution";

/* Determine solution to the system */
xx1 : null
xx2 : null;
XX : matrix([xx1], [xx2]);

nullMat : matrix([null], [null]);
]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text> \[\begin{matrix} {@eqq1@} \\ {@eqq2@} \\ \end{matrix}\]
\( A = {@AA@}, \quad |A| = {@detAA@},\) {@ta3c@}, \( x={@xx1@}, y = {@xx2@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>AA</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>detAA</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>dropdown</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>xx1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>xx2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/* Create the vector containing the answers */
vecAns : matrix([ans4], [ans5]);
</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>AA</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (a) Matrix \(A\) correctly identified.
    </span>
    <span class="multilang" lang="no">
        (a) Matrise \(A\) korrekt identifisert.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (a) Be careful, matrix \(A\) was not introduced correctly. Unfortunately, this error changed the difficulty of the problem and no more points can be awarded.
    </span>
    <span class="multilang" lang="no">
        (a) V&#230;r forsiktig, matrise \(A\) ble ikke introdusert riktig. Dessverre endret denne feilen problemets vanskelighetsgrad og flere poeng kan ikke tildeles.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>detAA</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (b) Determinant of \(A\) correctly computed.
    </span>
    <span class="multilang" lang="no">
        (b) Determinant til \(A\)  er korrekt beregnet.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (b) Be careful, the determinant of \(A\) was not computed correctly. Unfortunately, this error changed the difficulty of the problem and no more points can be awarded.
    </span>
    <span class="multilang" lang="no">
        (b) V&#230;r forsiktig, determinanten til \(A\) ble ikke beregnet riktig. Dessverre endret denne feilen problemets vanskelighetsgrad og flere poeng kan ikke tildeles.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3c</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.4000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (c) The dimension of the solution space was identified correctly.
    </span>
    <span class="multilang" lang="no">
        (c) Dimensjonen til l&#248;sningsrommet ble korrekt identifisert.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (c) The dimension of the solution space was not identified correctly. What happens if we multiply the first equation by \( {@kk@} \)? Unfortunately, this error changed the difficulty of the problem and no more points can be awarded.
    </span>
    <span class="multilang" lang="no">
        (c) Dimensjonen til l&#248;sningsrommet ble ikke identifisert riktig. Hva skjer hvis vi multipliserer den f&#248;rste ligningen med \( {@kk@} \)? Dessverre endret denne feilen problemets vanskelighetsgrad og flere poeng kan ikke tildeles.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <answertest>AlgEquiv</answertest>
                <sans>vecAns</sans>
                <tans>nullMat</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (d) Careful, as indicated in the question, if you think there is no solution, please fill in the entries for \( x \) and \( y \) with the word
        <b>
            null
        </b>
    </span>
    <span class="multilang" lang="no">
        (d) V&#230;r forsiktig, som angitt i sp&#248;rsm&#229;let, hvis du tror det ikke er noen l&#248;sning, fyll inn plassene for \( x \) og \( y \) med ordet
        <b>
            null
        </b>
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1672849393</deployedseed>
        <deployedseed>298980100</deployedseed>
        <deployedseed>1151982337</deployedseed>
        <deployedseed>1074868519</deployedseed>
        <deployedseed>646882530</deployedseed>
        <deployedseed>1354503646</deployedseed>
        <deployedseed>161684088</deployedseed>
        <deployedseed>839340090</deployedseed>
        <deployedseed>449968944</deployedseed>
        <deployedseed>1494058370</deployedseed>
        <deployedseed>1429100905</deployedseed>
        <deployedseed>1789009877</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>AA</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>detAA</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3c</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>xx1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>xx2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-4-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(2*AA,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(detAA+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["One unique solution"]]></value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>AA</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>detAA</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["One unique solution"]]></value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.4000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-3-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>