<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>10.3.2.4 Finding the exponential form</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<span class="multilang" lang="en">
    Express the complex number \(z = \displaystyle{@zed@}\)&#160; in exponential form. The argument of \(z\) should be given precisely in terms of \(\pi\), i.e. do not give your answer in decimal form:
</span>
<span class="multilang" lang="no">
    Uttrykk det komplekse tallet \(z = \displaystyle{@zed@}\)&#160; i eksponentiell form. Argumentet til \(z\) skal gis presist i form av \(\pi\), dvs. ikke gi svaret ditt i desimal form:
</span>
<p>
    \(z=\quad\)[[input:ans1]][[validation:ans1]]&#160;[[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We want to express \(z\) in the form \(re^{i\theta}\).
    </span>
    <span class="multilang" lang="no">
        Vi &#248;nsker &#229; uttrykke \(z\) i formen \(re^{i\theta}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We calculate \(r = \sqrt{({@re@})^2 + ({@im@})^2} = {@arr@}\) and \(\theta = \tan^{-1}\left({@im@}/{@re@}\right) = {@th@}\).
    </span>
    <span class="multilang" lang="no">
        Vi beregner \(r = \sqrt{({@re@})^2 + ({@im@})^2} = {@arr@}\) og \(\theta = \tan^{-1}\left({@im@}/{@re@}\right) = {@th@}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore: \[z = {@arr@}e^{{@th@}i}.\]
    </span>
    <span class="multilang" lang="no">
        Derfor: \[z = {@arr@}e^{{@th@}i}.\]
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>th: rand([%pi/3,%pi/4,-%pi/3,-%pi/4])
arr: 2 + rand(10);
re: arr*cos(th);
im: arr*sin(th);
simp:false;
ta: arr*e^(th*%i);
zed: re + im*%i;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[ {@ta@} = {@zed@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span data-darkreader-inline-color="" style="font-size: 1.5em; color: green; --darkreader-inline-color:#7fff7f;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span data-darkreader-inline-color="" style="font-size: 1.5em; color: orange; --darkreader-inline-color:#ffb733;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span data-darkreader-inline-color="" style="font-size: 1.5em; color: red; --darkreader-inline-color:#ff3333;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>13</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1901606859</deployedseed>
        <deployedseed>602818603</deployedseed>
        <deployedseed>894327509</deployedseed>
        <deployedseed>1479722686</deployedseed>
        <deployedseed>1994926504</deployedseed>
        <deployedseed>1023542860</deployedseed>
        <deployedseed>603625680</deployedseed>
        <deployedseed>1866674831</deployedseed>
        <deployedseed>1786551608</deployedseed>
        <deployedseed>1910411624</deployedseed>
        <deployedseed>1013862141</deployedseed>
        <deployedseed>583681724</deployedseed>
        <deployedseed>572819719</deployedseed>
        <deployedseed>686289728</deployedseed>
        <deployedseed>1394133457</deployedseed>
        <deployedseed>1380486129</deployedseed>
        <deployedseed>1750255028</deployedseed>
        <deployedseed>1893038075</deployedseed>
        <deployedseed>1857467541</deployedseed>
        <deployedseed>1886071763</deployedseed>
        <deployedseed>963949412</deployedseed>
        <deployedseed>1523582033</deployedseed>
        <deployedseed>1901542382</deployedseed>
        <deployedseed>1322232985</deployedseed>
        <deployedseed>1511949970</deployedseed>
        <deployedseed>233523471</deployedseed>
        <deployedseed>1198274324</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>0 + 2*%i</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>