<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>6.2.3.3 Compute sech given tanh</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Given \(\tanh{@arg@} = {@tanh@} \). Compute \( \rm{sech}\) \( {@arg@} \). Round up and enter 4 decimal places. If the answer has less than 4 decimal places, then enter zero(s) at the end.
    </span>
    <span class="multilang" lang="no">
        Gitt \(\tanh{@arg@} = {@tanh@} \). Beregn \( \rm{sech}\) \( {@arg@} \). Rund opp til 4 desimaler. Hvis svaret har f&#230;rre enn 4 desimaler, m&#229; du sette inn null(er) p&#229; slutten.
    </span>
</p>
<p>
    \( \rm{sech} \) \({@arg@} = \) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Use the identity \( \rm{sech}^2\)\({@arg@} = 1 - \tanh^2{@arg@}\) to calculate \(\rm{sech}^2\)\({@arg@}\), then take the (positive) square root to obtain \(\rm{sech}\) \({@arg@} \), noting that \(\rm{sech}\) \( {@arg@} = \frac{1}{\cosh{@arg@}} &gt;0\) for all values of \({@arg@}\).
    </span>
    <span class="multilang" lang="no">
        Bruk identiteten \( \rm{sech}^2\)\({@arg@} = 1 - \tanh^2{@arg@}\) for &#229; kalkulere \(\rm{sech}^2\)\({@arg@}\), og ta (den positive) kvadratroten for &#229; f&#229; \(\rm{sech}\) \({@arg@} \). Merk at \(\rm{sech}\) \( {@arg@} = \frac{1}{\cosh{@arg@}} &gt;0\) for alle verdier av \({@arg@}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(\rm{sech}\) \({@arg@} = \sqrt{1 - \tanh^2{@arg@}} = \sqrt{1 - ({@tanh@})^2} ={@ta@}\) to four significant figures (or four decimal places in this case).
    </span>
    <span class="multilang" lang="no">
        \(\rm{sech}\) \({@arg@} = \sqrt{1 - \tanh^2{@arg@}} = \sqrt{1 - ({@tanh@})^2} ={@ta@}\) (eller fire desimaler i dette tilfellet).
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022082900</text>
        </stackversion>
        <questionvariables>
            <text>arg: rand([x,y,z]);
xx: rand([23,24,25,33,34,35])/100;
tanh: decimalplaces(float(tanh(xx)),4);
sech2: 1-tanh*tanh;
sech: sqrt(sech2);
ta: decimalplaces(float(sech),4);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>For \(\tanh{@arg@} = {@tanh@} \) the \( \rm{sech} \) \({@arg@} = {@ta@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>sinh,cosh,tanh,coth,sech,cosech,[[BASIC-ALGEBRA]]</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>2</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>4</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>726280345</deployedseed>
        <deployedseed>340323282</deployedseed>
        <deployedseed>1488508266</deployedseed>
        <deployedseed>670076169</deployedseed>
        <deployedseed>1506111712</deployedseed>
        <deployedseed>2070103918</deployedseed>
        <deployedseed>348438463</deployedseed>
        <deployedseed>672256274</deployedseed>
        <deployedseed>973619248</deployedseed>
        <deployedseed>45045465</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>0.94</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>