<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>9.5.1.2 Direction ratios and cosines, Exercise</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    Exercise
</p>
<ol>
    <p>
        In all of the following questions give your answers exactly. Recall that \(\sqrt{2}\) can be entered by typing sqrt(2).
    </p>
    <li>
        <p>
            Points\(A\) and \(B\) have position vectors \(\underline{ a} ={@A1@}\), and \(\underline{ b} = {@B1@}\) respectively. Find
        </p>
        <p>
            (a) \(\overrightarrow{AB}=\) [[input:AB1]] [[validation:AB1]]
        </p>
        <p>
            [[feedback:diff]]
        </p>
        <p>
            (b) \(|\overrightarrow{AB}|=\) [[input:magAB1]] [[validation:magAB1]]
        </p>
        <p>
            [[feedback:mag]]
        </p>
        <p>
            (c) The direction ratios of \(\overrightarrow{AB}\): [[input:dr1]] [[validation:dr1]]:[[input:dr2]][[validation:dr2]]
        </p>
        <p>
            [[feedback:DR]]
        </p>
        <p>
            (d) The direction cosines of \(\overrightarrow{AB}\), \((\ell,m)=\) [[input:dc]] [[validation:dc]]
        </p>
        <p>
            [[feedback:DC]]
        </p>
        <p>
            (e) \(\ell^2+m^2=\) [[input:magCos]] [[validation:magCos]]
        </p>
        <p>
            [[feedback:sumDC]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <strong>
        FEEDBACK
    </strong>
</p>
<li>
    <p>
        We have that
    </p>
    <p>
        (a) \(\overrightarrow{AB}={@R1@}\)
    </p>
    <p>
        (b) \(|\overrightarrow{AB}|={@magr1@}\)
    </p>
    <p>
        (c) The direction ratios of \(\overrightarrow{AB}\): \({@r1[1,1]@}:&#160;{@r1[2,1]@}\)
    </p>
    <p>
        (d) The direction cosines of \(\overrightarrow{AB}\), \((\ell,m)=\){@tadc@}
    </p>
    <p>
        (e) \(\ell^2+m^2=1\).
    </p>
</li>
]]></text>
        </generalfeedback>
        <defaultgrade>2.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020070100</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[stackvectortex(ex):= block(sconcat("{ \\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k);
a1: matrix([5],[3]) - rand(matrix([10],[3]));
b1: matrix([5],[8]) - rand(matrix([10],[2]));
r1:  b1-a1;
magr1: sqrt(dotproduct(r1,r1));
u1: matrix([i],[j]);
A1: texboldatoms(dotproduct(a1,u1));
B1: texboldatoms(dotproduct(b1,u1));
R1:  texboldatoms(dotproduct(r1,u1));
l1: (r1[1,1])/(magr1);
m1: (r1[2,1])/(magr1);
tadc: matrix([l1,m1]);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>The direction ratios for the vector \(\mathbf{r} = {@R1@}\) are \({@r1[1,1]@}\,:\,{@r1[1,2]@}\), and its direction cosines are \(\ell = {@l1@}\) and \(m = {@m1@}\).</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<p>
    Correct answer, well done.
</p>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<p>
    Your answer is partially correct.
</p>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<p>
    Incorrect answer.
</p>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>symi</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>AB1</name>
            <type>matrix</type>
            <tans>r1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>dc</name>
            <type>matrix</type>
            <tans>tadc</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>dr1</name>
            <type>algebraic</type>
            <tans>r1[1,1]</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>dr2</name>
            <type>algebraic</type>
            <tans>r1[2,1]</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>magAB1</name>
            <type>algebraic</type>
            <tans>magr1</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>magCos</name>
            <type>algebraic</type>
            <tans>1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>DC</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>dc1 = dc[1,1];
dc2 = dc[1,2];</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>dc</sans>
                <tans>tadc</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0.0000000</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>VEC-DC_COR</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<br/>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>VEC-DC_INCOR</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <br/>
</p>
<p>
    Remember the formula for calculating the direction cosine (draw a diagram if this helps!), and try again.
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>diff</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>AB1;
</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>AB1</sans>
                <tans>r1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>diff-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>diff-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
Your expression for the difference seems to contain an error.
<br/>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>DR</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>student_r: [dr1*x=dr2*y];
taSys: [r1[1,1]*x=r1[2,1]*y];</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>SysEquiv</answertest>
                <sans>student_r</sans>
                <tans>taSys</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>VEC-DR-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>VEC-DR-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>mag</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>magAB1</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>magAB1</sans>
                <tans>magr1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>mag-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>mag-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    Your expression for the magnitude appears to be incorrect.
    <br/>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>sumDC</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>magCos</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>magCos</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>sumDC-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    Well done. Do you expect this formula to hold in general?
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>sumDC-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    This answer is incorrect. Please check your calculations.
    <br/>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>521325657</deployedseed>
        <deployedseed>453182377</deployedseed>
        <deployedseed>1422894494</deployedseed>
        <deployedseed>568066611</deployedseed>
        <deployedseed>538458330</deployedseed>
        <deployedseed>2070177701</deployedseed>
        <deployedseed>1009751502</deployedseed>
        <deployedseed>1293742219</deployedseed>
        <deployedseed>1724383337</deployedseed>
        <deployedseed>629072897</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>AB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dc</name>
                <value>tadc</value>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value>r1[1,1]</value>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value>r1[2,1]</value>
            </testinput>
            <testinput>
                <name>magAB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>magCos</name>
                <value/>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DC_COR</expectedanswernote>
            </expected>
            <expected>
                <name>diff</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>mag</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>sumDC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>AB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dc</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value>r1[1,1]</value>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value>r1[2,1]</value>
            </testinput>
            <testinput>
                <name>magAB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>magCos</name>
                <value/>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>diff</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>mag</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>sumDC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>AB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dc</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value>r1[1,1]</value>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value>r1[2,1]</value>
            </testinput>
            <testinput>
                <name>magAB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>magCos</name>
                <value/>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>diff</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>mag</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>sumDC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>AB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dc</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value>r1[1,1]/2</value>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value>r1[2,1]/2</value>
            </testinput>
            <testinput>
                <name>magAB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>magCos</name>
                <value/>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>diff</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty/>
                <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>mag</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>sumDC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>5</testcase>
            <testinput>
                <name>AB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dc</name>
                <value>matrix([2,2])</value>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value>10</value>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>magAB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>magCos</name>
                <value/>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DC_INCOR</expectedanswernote>
            </expected>
            <expected>
                <name>diff</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DR-FALSE</expectedanswernote>
            </expected>
            <expected>
                <name>mag</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>sumDC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>6</testcase>
            <testinput>
                <name>AB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dc</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value/>
            </testinput>
            <testinput>
                <name>magAB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>magCos</name>
                <value>1.1</value>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>diff</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>mag</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>sumDC</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty/>
                <expectedanswernote>sumDC-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>7</testcase>
            <testinput>
                <name>AB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dc</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value/>
            </testinput>
            <testinput>
                <name>magAB1</name>
                <value>magr1</value>
            </testinput>
            <testinput>
                <name>magCos</name>
                <value/>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>diff</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>mag</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty/>
                <expectedanswernote>mag-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>sumDC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>8</testcase>
            <testinput>
                <name>AB1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dc</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value/>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value/>
            </testinput>
            <testinput>
                <name>magAB1</name>
                <value>101</value>
            </testinput>
            <testinput>
                <name>magCos</name>
                <value/>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>diff</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>mag</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty/>
                <expectedanswernote>mag-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>sumDC</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>