<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>9.5.3.3 Vector equation of a line through two points in 2D Exercise</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Write down the vector equation of the line which passes through the points with position vectors \({@stackvector(a)@}={@A1@}\) and \({@stackvector(b)@}=
        {@B1@}\).
    </span>
    <span class="multilang" lang="no">
        Skriv ned vektorligningen til linjen som g&#229;r gjennom punktene med posisjonsvektorene \({@stackvector(a)@}={@A1@}\) og \({@stackvector(b)@}=
                                {@B1@}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Use the variable \(t\) in your answer.
    </span>
    <span class="multilang" lang="no">
        Bruk variabelen \(t \) i svaret ditt.
    </span>
</p>
\({@stackvector(r)@}=\) [[input:ans1i]] [[validation:ans1i]] \({@stackvector(i)@}+{}\) [[input:ans1j]] [[validation:ans1j]] \({@stackvector(j)@}\) [[feedback:prt1]]
<p>
    <span class="multilang" lang="en">
        Also express the equation in column vector form.
    </span>
    <span class="multilang" lang="no">
        Uttrykk ogs&#229; ligningen i kolonnevektorform.
    </span>
</p>
<div>
    \({@stackvector(r)@}=\) [[input:ans2b]] [[validation:ans2b]] \({}+t\ \) [[input:ans2d]] [[validation:ans2d]] [[feedback:prt2]]
</div>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The vector \(\overrightarrow{AB}\) is given by
    </span>
    <span class="multilang" lang="no">
        Vektoren \(\overrightarrow{AB}\) er gitt av
    </span>
</p>
\[ {@stackvector(b)@}-{@stackvector(a)@} = {@B1@} - ({@A1@}) = {@R1@}.\]
<p>
    <span class="multilang" lang="en">
        The equation of the line is then
    </span>
    <span class="multilang" lang="no">
        Linjens ligning er da
    </span>
</p>
\[{@stackvector(r)@} = {@stackvector(a)@} + t({@stackvector(b)@}-{@stackvector(a)@}) = ({@A1@}) + t({@R1@}).\]
<p>
    <span class="multilang" lang="en">
        Using column vector notation we could write
    </span>
    <span class="multilang" lang="no">
        Ved hjelp av kolonnevektornotasjon kan vi skrive
    </span>
</p>
\[{@stackvector(r)@} = {@a1@}+t\,{@r1@}.\]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022082900</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[stackvectortex(ex):= block(sconcat("{ \\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k,l);
/*two distinct vectors*/
a1: matrix([3],[5]) - rand(matrix([3],[10]));
b1: matrix([8],[5]) - rand(matrix([2],[10]));
r1:  b1-a1;
negr1:  -r1;
u1: matrix([i],[j]);
A1: texboldatoms(dotproduct(a1,u1));
B1: texboldatoms(dotproduct(b1,u1));
R1:  texboldatoms(dotproduct(r1,u1));
r2: matrix([1],[1],[1],[1]);
u2: matrix([i],[j],[l],[k]);
R2:  texboldatoms(dotproduct(r2,u2));]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[\underline{ a}={@A1@} \quad  \text{and} \quad \underline{b}= {@B1@}\]
\[\underline{r}= ({@a1[1,1]+t*r1[1,1]@}) i + ({@a1[2,1]+t*r1[2,1]@}) j\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>symi</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1i</name>
            <type>algebraic</type>
            <tans>a1[1,1]+t*r1[1,1]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans1j</name>
            <type>algebraic</type>
            <tans>a1[2,1]+t*r1[2,1]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2b</name>
            <type>matrix</type>
            <tans>a1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2d</name>
            <type>matrix</type>
            <tans>r1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ans1:matrix([ans1i], [ans1j]);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>a1+t*r1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1i</sans>
                <tans>a1[1,1]+t*r1[1,1]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    It seems that only one of your equation entries is correct, please check your answer.
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1j</sans>
                <tans>a1[2,1]+t*r1[2,1]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    It seems that only one of your equation entries is correct, please check your answer.
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2b</sans>
                <tans>a1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2b</sans>
                <tans>b1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2d</sans>
                <tans>r1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2d</sans>
                <tans>-r1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1815481640</deployedseed>
        <deployedseed>1205749302</deployedseed>
        <deployedseed>1669306292</deployedseed>
        <deployedseed>1112047614</deployedseed>
        <deployedseed>2052486644</deployedseed>
        <deployedseed>336421937</deployedseed>
        <deployedseed>752304319</deployedseed>
        <deployedseed>301237622</deployedseed>
        <deployedseed>2013405524</deployedseed>
        <deployedseed>448995278</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1i</name>
                <value>a1[1,1]+t*r1[1,1]</value>
            </testinput>
            <testinput>
                <name>ans1j</name>
                <value>a1[2,1]+t*r1[2,1]</value>
            </testinput>
            <testinput>
                <name>ans2b</name>
                <value>a1</value>
            </testinput>
            <testinput>
                <name>ans2d</name>
                <value>r1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-3-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1i</name>
                <value>a1[2,1]+t*r1[2,1]</value>
            </testinput>
            <testinput>
                <name>ans1j</name>
                <value>a1[1,1]+t*r1[1,1]</value>
            </testinput>
            <testinput>
                <name>ans2b</name>
                <value>a1</value>
            </testinput>
            <testinput>
                <name>ans2d</name>
                <value>r1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty/>
                <expectedanswernote>prt1-3-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-3-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1i</name>
                <value>a1[2,1]+t*r1[2,1]</value>
            </testinput>
            <testinput>
                <name>ans1j</name>
                <value>a1[2,1]+t*r1[2,1]</value>
            </testinput>
            <testinput>
                <name>ans2b</name>
                <value>b1</value>
            </testinput>
            <testinput>
                <name>ans2d</name>
                <value>negr1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty/>
                <expectedanswernote>prt1-3-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-4-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>ans1i</name>
                <value>a1[1,1]+t*r1[1,1]</value>
            </testinput>
            <testinput>
                <name>ans1j</name>
                <value>a1[1,1]+t*r1[1,1]</value>
            </testinput>
            <testinput>
                <name>ans2b</name>
                <value>a1</value>
            </testinput>
            <testinput>
                <name>ans2d</name>
                <value>r1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty/>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-3-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>