<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.4.1.T1b 2_prod_intro</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Here we are going to find the first derivative of \(y(x) = \displaystyle {@y_x@}.\)
    </span>
    <span class="multilang" lang="no">
        Her skal vi finne den f&#248;rste deriverte av \(y(x) = \displaystyle {@y_x@}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We can factorise \(y(x)\) into the two factors
    </span>
    <span class="multilang" lang="no">
        Vi kan faktorisere \(y(x)\) inn i to faktorer.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \[ f(x) = {@f_x@} \hspace{1cm} g(x) = {@g_x@}. \]
    </span>
    <span class="multilang" lang="no">
        \[ f(x) = {@f_x@} \hspace{1cm} g(x) = {@g_x@}. \]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Now find the derivative of each of these factors.
    </span>
    <span class="multilang" lang="no">
        Finn n&#229; den deriverte av hver av disse faktorene.
    </span>
</p>
<p>
    \( \dfrac{df}{dx} = \)
</p>
[[input:ans1]] [[validation:ans1]][[feedback:prt1]]
<p>
    \( \dfrac{dg}{dx} = \) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
</p>
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Finally, substitute into the product rule to determine
    </span>
    <span class="multilang" lang="no">
        Til slutt, sett inn i produktregelen for &#229; bestemme.
    </span>
</p>
<p>
    \( \dfrac{dy}{dx} = \)&#160;[[input:ans3]] [[validation:ans3]][[feedback:prt3]]
</p>
<br/>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    We have \(y = f(x)g(x)= \displaystyle {@y_x@}\), with \( f(x) = {@f_x@}\) and \(g(x) = \displaystyle {@g_x@} = x^{@-a1@}\).
</p>
<p>
    Then, \( \dfrac{df}{dx} = {@ta1@} \) and \( \dfrac{dg}{dx} = \displaystyle{{@ta2@}}\).
</p>
<p>
    So 
        \[\begin{align*}
  	  	  	  	\dfrac{dy}{dx} &amp;= \dfrac{df}{dx} g(x)+ f(x) \dfrac{dg}{dx}\\
  	  	  	  	&amp;= \displaystyle {@diff(y_x,x)@}\\
  	  	  	  	&amp;= \displaystyle{{@ta3@}}.
        \end{align*}\]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1:rand([2,3,4,5]);
f_x: e^(a1*x);
g_x: x^(-a1);
y_x: f_x*g_x;
ta1: diff(f_x,x);
ta2: diff(g_x,x);
ta3: factor(diff(y_x,x));
</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( y={@y_x@}, \quad  \dfrac{dy}{dx} ={@ta3@}.\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>6</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>6</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1083041560</deployedseed>
        <deployedseed>1329190427</deployedseed>
        <deployedseed>2008513571</deployedseed>
        <deployedseed>139804569</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>