<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.4.1.2</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;velse
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In each case find the derivative of the function
    </span>
    <span class="multilang" lang="no">
        I hvert tilfelle finn den deriverte av funksjonen
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                \( y =\dfrac{x}{\cos (x)} \) (recall that \( \dfrac{1}{\cos(x)}=\sec(x)\)
            </span>
            <span class="multilang" lang="no">
                \( y =\dfrac{x}{\cos (x)} \) (husk at \( \dfrac{1}{\cos(x)}=\sec(x)\))
            </span>
        </p>
        <p>
            \( \dfrac{dy}{dx} =\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            \( y ={@f2(x)@}{@g2(x)@} \)
        </p>
        <p>
            \( \dfrac{dy}{dx} =\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Obtain the derivative of \(y = x e^{x} \tan x \) by using your answers to part (a) and (b)
            </span>
            <span class="multilang" lang="no">
                Finn den deriverte av \(y = x e^{x} \tan x \) ved &#229; bruke svarene fra del (a) og (b).
            </span>
        </p>
        <p>
            \( \dfrac{dy}{dx} =\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We have \(y = f(x)g(x)\), with \( f(x) = {@f1(x)@}\) and \(g(x) = \displaystyle {@g1(x)@} .\)
            </span>
            <span class="multilang" lang="no">
                Vi har \(y = f(x)g(x)\), med \(f(x) = {@f1(x)@}\) og \(g(x) = \displaystyle {@g1(x)@} .\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Then, \( \dfrac{df}{dx} = {@df1@} \) and \( \dfrac{dg}{dx} = \displaystyle{{@dg1@}} .\)
            </span>
            <span class="multilang" lang="no">
                Da blir \( \dfrac{df}{dx} = {@df1@} \) og \( \dfrac{dg}{dx} = \displaystyle{{@dg1@}} .\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                So \[\begin{align*} \dfrac{dy}{dx} &amp;= {@df1@}\left({@g1(x)@}\right) +{@f1(x)@} \left( \displaystyle{{@dg1@}} \right) \\ &amp;= \displaystyle{{@ta1@}}. \end{align*}\]
            </span>
            <span class="multilang" lang="no">
                S&#229; \[\begin{align*} \dfrac{dy}{dx} &amp;= {@df1@}\left({@g1(x)@}\right) +{@f1(x)@} \left( \displaystyle{{@dg1@}} \right) \\ &amp;= \displaystyle{{@ta1@}}. \end{align*}\]
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We have \(y = f(x)g(x)\), with \( f(x) = {@f2(x)@}\) and \(g(x) = {@g2(x)@} .\)
            </span>
            <span class="multilang" lang="no">
                Vi har \(y = f(x)g(x)\), der \( f(x) = {@f2(x)@}\) og \(g(x) = {@g2(x)@}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Then, \( \dfrac{df}{dx} = {@df2@} \) and \( \dfrac{dg}{dx} = \displaystyle{{@dg2@}}.\)
            </span>
            <span class="multilang" lang="no">
                Da blir \( \dfrac{df}{dx} = {@df2@} \) og \( \dfrac{dg}{dx} = \displaystyle{{@dg2@}}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                So \[\begin{align*} \dfrac{dy}{dx} &amp;= {@df2@}(\displaystyle{{@g2(x)@}}) + {@f2(x)@}(\displaystyle{{@dg2@}}) \\ &amp;= \displaystyle{{@ta2@}}. \end{align*}\]
            </span>
            <span class="multilang" lang="no">
                S&#229; \[\begin{align*} \dfrac{dy}{dx} &amp;= {@df2@}(\displaystyle{{@g2(x)@}}) + {@f2(x)@}(\displaystyle{{@dg2@}}) \\ &amp;= \displaystyle{{@ta2@}}. \end{align*}\]
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                The derivative of \(y = x e^{ x} \tan x = \dfrac{x}{\cos x} e^{x} \sin x \) is found by applying the product rule to the results of parts (a) and (b).
            </span>
            <span class="multilang" lang="no">
                Den deriverte av \(y = x e^{x} \tan x = \dfrac{x}{\cos x} e^{x} \sin x\) finnes ved &#229; anvende produktregelen p&#229; resultatene fra del (a) og (b).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                So we have \(y = f(x)g(x)\), with \( f(x) = \displaystyle{{@f3(x)@}}\) and \(g(x) = {@g3(x)@} .\)
            </span>
            <span class="multilang" lang="no">
                Vi har alts&#229; \(y = f(x)g(x)\), der \( f(x) = \displaystyle{{@f3(x)@}}\) og \(g(x) = {@g3(x)@} .\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Then, \( \dfrac{df}{dx} = \displaystyle{{@df3@}} \) and \( \dfrac{dg}{dx} = \displaystyle{{@dg3@}}.\)
            </span>
            <span class="multilang" lang="no">
                Da blir \( \dfrac{df}{dx} = \displaystyle{{@df3@}} \) og \( \dfrac{dg}{dx} = \displaystyle{{@dg3@}}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                So \[\begin{align*} \dfrac{dy}{dx} &amp;= \displaystyle{{@df3@}}(\displaystyle{{@g3(x)@}}) + {@f3(x)@}(\displaystyle{{@dg3@}}) \\ &amp;= \displaystyle{{@ta3@}}. \end{align*}\]
            </span>
            <span class="multilang" lang="no">
                S&#229; \[\begin{align*} \dfrac{dy}{dx} &amp;= \displaystyle{{@df3@}}(\displaystyle{{@g3(x)@}}) + {@f3(x)@}(\displaystyle{{@dg3@}}) \\ &amp;= \displaystyle{{@ta3@}}. \end{align*}\]
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>f1(x):=x;
g1(x):=1/(cos(x));
df1: diff(f1(x),x);
dg1: diff(g1(x),x);
ta1: diff(f1(x)*g1(x),x);


f2(x):=e^(x);
g2(x):=sin(x);
df2: diff(f2(x),x);
dg2: diff(g2(x),x);
ta2: factor(diff(f2(x)*g2(x),x));

f3(x):=x/cos(x);
g3(x):=e^(x)*sin(x);
df3: diff(f3(x),x);
dg3: diff(g3(x),x);
ta3:diff(f3(x)*g3(x),x);
ta3: factor(trigreduce(ta3));


</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text/>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>35</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <br/>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>