<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>3.6.2.3_Find the partial fractions (short)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Find the partial fractions of:
    </span>
    <span class="multilang" lang="no">
        Delbr&#248;ksoppspalt:
    </span>
</p>
<p>
    \(\displaystyle {@WholeExptFact1@}= \) [[input:ans1]][[validation:ans1]][[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given the expression \[\displaystyle {@WholeExptFact1@}.\] The numerator has degree \(1\) and the denominator has degree \(2\), therefore this is a proper fraction. There are \(2\) linear factors in the denominator, therefore
    </span>
    <span class="multilang" lang="no">
        Vi f&#229;r uttrykket \[\displaystyle {@WholeExptFact1@}.\] Telleren har grad \(1\) og nevneren har grad \(2\), derfor er dette en ekte br&#248;k. Det er \(2\) line&#230;re faktorer i nevneren
    </span>
    \[\displaystyle {@WholeExptFact1@}
                = \frac{A}{{@exp11@}} + \frac{B}{{@exp22@}} = \frac{A\cdot ({@exp22@}) + B\cdot ({@exp11@}) }{ ({@exp11@})({@exp22@})} .\]
</p>
<p>
    <span class="multilang" lang="en">
        Equating the top line yields
    </span>
    <span class="multilang" lang="no">
        Likestille &#248;vre linjen gir
    </span>
    \[ {@expNum11@} = A\cdot ({@exp22@}) + B\cdot ({@exp11@}).\]
</p>
<p>
    <span class="multilang" lang="en">
        Then by substituting values of \(x\) we can solve for \(A, \, B\) and obtain \( A = {@expA@}\) and \( B = {@expB@}.\)
    </span>
    <span class="multilang" lang="no">
        Ved &#229; s&#229; erstatte verdiene av \(x\) kan vi l&#248;se for \(A, \, B\) og f&#229; \( A = {@expA@}\) og \( B = {@expB@}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore:
    </span>
    <span class="multilang" lang="no">
        Dermed f&#229;r vi:
    </span>
    \[ {@WholeExptFact1@}= {@ta1@}.\]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>var:x;
var2:A;
var3:B;
rand1:rand_with_step(1,2,1);
rand2:rand_with_step(-5,-1,1);
rand3:rand_with_step(1,3,1);
rand4:rand_with_step(0,2,1);
exp1:rand1*var+(rand1*rand2+1);
exp2:rand3*var+(rand3*rand4-1);
expDenom1:expand(exp1*exp2);
ta11:(exp1)*(exp2);
expA:rand_with_prohib(-9,9,[0]);
expB:rand_with_prohib(-9,9,[0]);
expNum1:expand(expA*exp2+expB*exp1);
wholeExp1: expNum1/expDenom1;
WholeExptFact1: expNum1/ta11;
ta1:expA/exp1+expB/exp2;
exp11: exp1;
exp22: exp2;
expNum11: expNum1;

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[ {@WholeExptFact1@}= {@ta1@}.\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Your answer is incorrect. The correct partial fraction is \({@ta1@}\).
</span>
<span class="multilang" lang="no">
    Svaret ditt er feil. Riktig delbr&#248;k er \({@ta1@}\).
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>268363270</deployedseed>
        <deployedseed>1836852489</deployedseed>
        <deployedseed>291840650</deployedseed>
        <deployedseed>913944909</deployedseed>
        <deployedseed>2060580556</deployedseed>
        <deployedseed>25571352</deployedseed>
        <deployedseed>1120158455</deployedseed>
        <deployedseed>1328382954</deployedseed>
        <deployedseed>1028097588</deployedseed>
        <deployedseed>1589969008</deployedseed>
        <deployedseed>546859691</deployedseed>
        <deployedseed>1753069844</deployedseed>
        <deployedseed>271081459</deployedseed>
        <deployedseed>325235837</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta1+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>