<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>3.6.2.2_Task_Express in partial fractions</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Express \[{@wholeExp1@}\] in partial fractions.
    </span>
    <span class="multilang" lang="no">
        Uttrykk \[{@wholeExp1@}\] som sum av delbr&#248;ker.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        First factorise the denominator (find \({@factor1@}\) and \({@factor2@}\) so that \({@expDenom1@}={@textExp@}\)):
    </span>
    <span class="multilang" lang="no">
        Faktoriser f&#248;rst nevneren (finn \({@faktor1@}\) og \({@faktor2@}\) slik at \({@expDenom1@}={@faktor1faktor2@}\)):
    </span>
</p>
<br/>
<p>
    \({@expDenom1@}=\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]
</p>
<p>
    <span class="multilang" lang="en">
        Because there are two linear factors we write this as:
    </span>
    <span class="multilang" lang="no">
        Siden det er to line&#230;re faktorer, skriver vi dette som:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \[\dfrac{\text{numerator}}{\text{denominator}}=\dfrac{{@var2@}}{{@factor1@}}+\dfrac{{@var3@}}{{@factor2@}}\]
    </span>
    <span class="multilang" lang="no">
        \[\dfrac{\text{teller}}{\text{nevner}}=\dfrac{{@var2@}}{{@faktor1@}}+\dfrac{{@var3@}}{{@faktor2@}}\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Multiply both sides by \({@textExp@}\) to obtain the equation from which to find \({@var2@}\) and \({@var3@}\):
    </span>
    <span class="multilang" lang="no">
        Multipliser begge sider med \({@textExp@}\) for &#229; f&#229; ligningen som kan finne \({@var2@}\) og \({@var3@}\):
    </span>
</p>
<br/>
<p>
    \({@expNum1@}=\)[[input:ans2]][[validation:ans2]][[feedback:prt2]]
</p>
<br/>
<p>
    <span class="multilang" lang="en">
        Substitute an appropriate value for \({@var@}\) to obtain \({@var2@}\):
    </span>
    <span class="multilang" lang="no">
        Sett inn en passende verdi for \({@var@}\) for &#229; f&#229; \({@var2@}\):
    </span>
</p>
<br/>
<p>
    <span class="multilang" lang="en">
        Substitute \({@var@}=\)
    </span>
    <span class="multilang" lang="no">
        Substituer \({@var@}=\)
    </span>
    [[input:ans3]][[validation:ans3]]
    <span class="multilang" lang="en">
        and get
    </span>
    <span class="multilang" lang="no">
        og f&#229;
    </span>
    \({@var2@}=\)[[input:ans4]].[[validation:ans4]]
</p>
[[feedback:prt3]][[feedback:prt4]]
<p>
    <span class="multilang" lang="en">
        Equating coefficients of \({@var@}\) obtain the value of \({@var3@}\):
    </span>
    <span class="multilang" lang="no">
        Ved &#229; likestille koeffisientene til \({@var@}\) f&#229;r du verdien av \({@var3@}\):
    </span>
</p>
<br/>
<p>
    \({@var3@}=\)[[input:ans5]].[[validation:ans5]]
</p>
[[feedback:prt5]]
<br/>
<p>
    <span class="multilang" lang="en">
        Finally, write down the partial fractions:
    </span>
    <span class="multilang" lang="no">
        Skriv til slutt ned delbr&#248;kene:
    </span>
</p>
<br/>
<p>
    \(\displaystyle{{@wholeExp1@}=}\)[[input:ans6]].[[validation:ans6]]
</p>
[[feedback:prt6]]
<br/>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Factorising the denominator we obtain \(({@exp1@})({@exp2@})\).
    </span>
    <span class="multilang" lang="no">
        Ved &#229; faktorisere nevneren f&#229;r vi \(({@exp1@})({@exp2@})\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We can see then that \(\dfrac{{@expNum1@}}{{@expDenom1@}}=\dfrac{{@expNum1@}}{({@exp1@})({@exp2@})}\).
    </span>
    <span class="multilang" lang="no">
        Vi kan da se at \(\dfrac{{@expNum1@}}{{@expDenom1@}}=\dfrac{{@expNum1@}}{({@exp1@})({@exp2@})}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We want to express the right hand side as a sum of two fractions. This means, we need to find expessions \({@var2@}\) and \({@var3@}\) such that
    </span>
    <span class="multilang" lang="no">
        Vi &#248;nsker &#229; uttrykke h&#248;yre side som en sum av to br&#248;ker. Dette betyr at vi m&#229; finne \({@var2@}\) og \({@var3@}\) slik at
    </span>
</p>
\[\dfrac{{@expNum1@}}{({@exp1@})({@exp2@})}=\dfrac{{@var2@}}{{@exp1@}}+\dfrac{{@var3@}}{{@exp2@}}\]
<p>
    <span class="multilang" lang="en">
        Multiplying both sides of the expression by \(({@exp1@})({@exp2@})\) we obtain
    </span>
    <span class="multilang" lang="no">
        Ved &#229; multiplisere begge sider av uttrykket med \(({@exp1@})({@exp2@})\) f&#229;r vi
    </span>
</p>
\[\dfrac{{@expNum1@}}{({@exp1@})({@exp2@})}({@exp1@})({@exp2@})=\dfrac{{@var2@}}{{@exp1@}}({@exp1@})({@exp2@})+\dfrac{{@var3@}}{{@exp2@}}({@exp1@})({@exp2@}).\]
<p>
    <span class="multilang" lang="en">
        Cancelling out, we obtain:
    </span>
    <span class="multilang" lang="no">
        Ved stryking f&#229;r vi:
    </span>
</p>
\[{@expNum1@}={@var2@}({@exp2@})+{@var3@}({@exp1@}).\]
<p>
    <span class="multilang" lang="en">
        In order to obtain \({@var2@}\), we can substitute \({@var@}=\displaystyle {@ta13_1@}\). From there we can see that \({@var2@}={@ta14_1@}\). Similarly we can obtain \({@var3@}={@ta14_1_1@}\). Finally, the correct fraction is:
    </span>
    <span class="multilang" lang="no">
        For &#229; f&#229; \({@var2@}\), kan vi sette inn \({@var@}=\displaystyle {@ta13_1@}\). Derfra kan vi se at \({@var2@}={@ta14_1@}\). P&#229; samme m&#229;te kan vi f&#229; \({@var3@}={@ta14_1_1@}\). Til slutt er den riktige br&#248;ken:
    </span>
</p>
\[{@ta15_1@}\]
]]></text>
        </generalfeedback>
        <defaultgrade>6.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>var:x;
var2:A;
var3:B;
rand1:rand_with_step(1,3,1);
rand2:rand_with_step(-5,-1,1);
rand3:rand_with_step(1,3,1);
rand4:rand_with_step(1,5,1);
exp1:rand1*var+(rand1*rand2+1);
exp2:rand2*var+rand3*rand4;
expDenom1:expand(exp1*exp2);
ta11:(exp1)*(exp2);

expA:rand_with_prohib(-9,9,[0]);
expB:rand_with_prohib(-9,9,[0]);
expNum1:expand(expA*exp2+expB*exp1);
wholeExp1: expNum1/expDenom1;
ta12_1:var2*exp2+var3*exp1;
ta12_2:var2*exp1+var3*exp2;
ta12_3:var2*(-exp2)+var3*(-exp1);
ta12_4:var2*(-exp1)+var3*(-exp2);

ta13_2:-(rand1*rand2+1)/rand1;
ta13_1:-rand3*rand4/rand2;

ta14_1:expA;
ta14_2:expB;

ta14_1_1:expB;
ta14_2_1:expA;

ta15_1:expA/exp1+expB/exp2;
ta15_2:(-expA)/(-exp1)-(expB)/(-exp2);


factor1: factor1;
factor2:factor2;
textExp:factor1*factor2;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[{@wholeExp1@} = {@ta15_1@}\]
\[ {@expDenom1@} = {@ta11@} \]
\[{@expNum1@}= {@ta12_1@}\]
\[ A = {@ta14_1@}, \quad B = {@ ta14_1_1@} \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta11</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta12_1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>A*?+B*?</syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta13_2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta14_1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>ta14_1_1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans6</name>
            <type>algebraic</type>
            <tans>ta15_1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta11</tans>
                <testoptions>var</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>FacForm</answertest>
                <sans>ans1</sans>
                <tans>ta11</tans>
                <testoptions>var</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>5.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta12_1</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0.1000000</falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta12_2</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0.1000000</falsepenalty>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta12_1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.7000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    <span class="multilang" lang="en">
        You answer is algebraically correct but for this problem you should avoid multiplying out your answer as it complicates the rest of the problem.
    </span>
    <span class="multilang" lang="no">
        Svaret ditt er algebraisk riktig, men for denne oppgaven b&#248;r du unng&#229; &#229; multiplisere ut svaret, da det kompliserer resten av oppgaven.
    </span>
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt2-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta12_2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.7000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
</p>
<p dir="ltr">
</p>
<span class="multilang" lang="en">
    You answer is algebraically correct but for this problem you should avoid multiplying out your answer as it complicates the rest of the problem.
</span>
<span class="multilang" lang="no">
    Svaret ditt er algebraisk riktig, men for denne oppgaven b&#248;r du unng&#229; &#229; multiplisere ut svaret, da det kompliserer resten av oppgaven.
</span>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta12_3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-5-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>5</falsenextnode>
                <falseanswernote>prt2-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta12_4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-6-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt2-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta13_1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta13_2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>abs(ans4)</sans>
                <tans>abs(ta14_1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    {@test@}
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>abs(ans4)</sans>
                <tans>abs(ta14_2)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>abs(ans5)</sans>
                <tans>abs(ta14_1_1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>abs(ans5)</sans>
                <tans>abs(ta14_2_1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ans7:-ans6;</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans6</sans>
                <tans>ta15_1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>ta15_2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1484479627</deployedseed>
        <deployedseed>461831087</deployedseed>
        <deployedseed>38331234</deployedseed>
        <deployedseed>229512218</deployedseed>
        <deployedseed>754553697</deployedseed>
        <deployedseed>112035072</deployedseed>
        <deployedseed>594052788</deployedseed>
        <deployedseed>1912727720</deployedseed>
        <deployedseed>374145298</deployedseed>
        <deployedseed>884511265</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta11</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta12_1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta13_1</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta14_1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta14_1_1</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ta15_1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt6-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(expand(ta11)+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(expand(ta12_1)+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ev(ta13_1+10,simp)</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ev(ta14_1+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ev(ta14_1_1+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ev(wholeExp1+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-4-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt6-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>