<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>7.3.4.1 Exercise 1</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;velse
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In this exercise you will use Laplace expansion along the 1st row to determine the determinant of the following matrix:
    </span>
    <span class="multilang" lang="no">
        I denne &#248;velsen skal du bruke Laplace utvidelsen langs 1st r&#229;den til &#229; bestemme determinanten av den f&#248;lgende matrise
    </span>
</p>
<div>
    \(A={@matrixa@}\)
</div>
<p>
    <span class="multilang" lang="en">
        Let the entries in the first row be denoted \(a_{1,1}\), \(a_{1,2}\) and \(a_{1,3}\) in the usual way. First list the cofactors \(A_{i,j}\) of these entries:
    </span>
    <span class="multilang" lang="no">
        La elementene i den f&#248;rste r&#229;den bli \(a_{1,1}\), \(a_{1,2}\) og \(a_{1,3}\) i den vanlige m&#229;ten. Liste de kofaktorene\(A_{i,j}\) av disse elementene:
    </span>
</p>
<p>
    \(A_{1,1}=\)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>
<p>
    \(A_{1,2}=\)[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
</p>
<p>
    \(A_{1,3}=\)[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
</p>
<p>
    <span class="multilang" lang="en">
        Use these values to evaluate the determinant of \(A\):
    </span>
    <span class="multilang" lang="no">
        Bruk de verdiene til &#229; evaluere determinanten til \(A\):
    </span>
</p>
[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        To obtain the cofactor of an element of a \(3 \times 3\) matrix we multiply the minor of that element by the corresponding sign from the \(3 \times 3\) array of signs.
    </span>
    <span class="multilang" lang="no">
        For &#229; f&#229; kofaktoren av en element til en \(3 \times 3\) matrise vi multipliserer den 'minor'-en av elementen ved den korresponderende tegn fra den \(3 \times 3\) array av tegnene.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For \(a_{1,1}\), the sign is \(+\), and the minor is \(\det{@matrix([AA[2,2], AA[2,3]], [AA[3,2], AA[3,3]])@}\). Hence the cofactor is \({@ta1@}.\)
    </span>
    <span class="multilang" lang="no">
        For \(a_{1,1}\), tegnet er \(+\), og 'minor'-en til dette er \(\det{@matrix([AA[2,2], AA[2,3]], [AA[3,2], AA[3,3]])@}\). Herved er kofaktoren: \({@ta1@}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For \(a_{1,2}\), the sign is \(-\), and the minor is \(\det{@matrix([AA[2,1], AA[2,3]], [AA[3,1], AA[3,3]])@}\). Hence the cofactor is \({@ta2@}.\)
    </span>
    <span class="multilang" lang="no">
        For \(a_{1,2}\), tegnet er \(-\), og 'minor'-en til dette er \(\det{@matrix([AA[2,1], AA[2,3]], [AA[3,1], AA[3,3]])@}\). Da er kofaktoren: \({@ta2@}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For \(a_{1,3}\), the sign is \(+\) again, and the minor is \(\det{@matrix([AA[2,1], AA[2,2]], [AA[3,1], AA[3,2]])@}\). Hence the cofactor is \({@ta3@}.\)
    </span>
    <span class="multilang" lang="no">
        For \(a_{1,3}\), tegnet er \(+\), og 'minor'-en til dette er \(\det{@matrix([AA[2,1], AA[2,2]], [AA[3,1], AA[3,2]])@}\). Da er kofaktoren: \({@ta3@}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The determinant of \(A\) is now given by the formula:
\[\det(A) = {@CC@} = {@DD@}\]
    </span>
    <span class="multilang" lang="no">
        Determinanten \(A\) kan beregnes da ved bruk av formelen:
\[\det(A) = {@CC@} = {@DD@}\]
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[matrixa: matrix([rand([-5,-4,-3,-2,-1,1,2,3,4,5]), rand([-5,-4,-3,-2,-1,1,2,3,4,5]), rand([-5,-4,-3,-2,-1,1,2,3,4,5])], [rand([-5,-4,-3,-2,-1,1,2,3,4,5]), rand([-5,-4,-3,-2,-1,1,2,3,4,5]), rand([-5,-4,-3,-2,-1,1,2,3,4,5])], [rand([-5,-4,-3,-2,-1,1,2,3,4,5]), rand([-5,-4,-3,-2,-1,1,2,3,4,5]), rand([-5,-4,-3,-2,-1,1,2,3,4,5])]);
matrixa: ev(matrixa,simp);

AA: matrixa;
MM11: matrix([AA[2,2], AA[2,3]], [AA[3,2], AA[3,3]]);
MM12: matrix([AA[2,1], AA[2,3]], [AA[3,1], AA[3,3]]);
MM13: matrix([AA[2,1], AA[2,2]], [AA[3,1], AA[3,2]]);

ta1: ev(determinant(MM11),simp)
ta2: ev(-determinant(MM12),simp)
ta3: ev(determinant(MM13),simp)
ta4: ev(determinant(matrixa),simp);

minorlist: [determinant(MM11),-determinant(MM12),determinant(MM13)];
simp: false;
CC: apply("+",zip_with("*",AA[1],minorlist));
DD: ev(CC,simp);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( A = {@AA@}, \, A_{1,1} ={@ta1@}, \, A_{1,2} = {@ta2@}, \, A_{1,3}={@ta3@}, \, detA=  {@ta4@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>ans3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans/>
                <tans/>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>815348399</deployedseed>
        <deployedseed>807206414</deployedseed>
        <deployedseed>1359982381</deployedseed>
        <deployedseed>415533922</deployedseed>
        <deployedseed>1666427994</deployedseed>
        <deployedseed>1519906486</deployedseed>
        <deployedseed>1881655594</deployedseed>
        <deployedseed>1203042910</deployedseed>
        <deployedseed>240042115</deployedseed>
        <deployedseed>770303117</deployedseed>
        <deployedseed>736720263</deployedseed>
        <deployedseed>101230784</deployedseed>
        <deployedseed>1543056679</deployedseed>
        <deployedseed>53117613</deployedseed>
        <deployedseed>938344926</deployedseed>
        <deployedseed>1208117384</deployedseed>
        <deployedseed>283144539</deployedseed>
        <deployedseed>392739448</deployedseed>
        <deployedseed>1733957553</deployedseed>
        <deployedseed>250556807</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>-ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>