<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.1.3 differentiate intro</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;velse
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Carry out the procedure above for the function \( y = {@a1@}x^2 \):
            </span>
            <span class="multilang" lang="no">
                Utf&#248;r prosedyren ovenfor for funksjonen \(y = {@a1@}x^2\):
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (1) Let \(A\) be the point \( (a, {@a1@}a^2)\).
            </span>
            <span class="multilang" lang="no">
                (1) La \(A\) v&#230;re punktet \( (a, {@a1@}a^2)\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (2) Let \(B\) be the point \( (a + h, {@a1@}(a + h)^2) \).
            </span>
            <span class="multilang" lang="no">
                (2) La \(B\) v&#230;re punktet \( (a+h, {@a1@}(a+h)^2) \).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (3) Find the gradient of the line \(AB\).
            </span>
            <span class="multilang" lang="no">
                (3) Finn stigningen til linjen \(AB\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (4) Let \(h \to 0\) to find the gradient of the curve at \(A\).
            </span>
            <span class="multilang" lang="no">
                (4) La \(h \to 0\) for &#229; finne stigningen til kurven i \(A\).
            </span>
        </p>
        <p>
            \( \dfrac{dy}{dx} \) = [[input:ans1]]
                                                [[validation:ans1]][[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Carry out the procedure above for the function \( y = x^{@b@} \):
            </span>
            <span class="multilang" lang="no">
                Utf&#248;r prosedyren over for funksjonen \( y = x^{@b@} \):
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (1) Let \(A\) be the point \( (a, a^{@b@})\).
            </span>
            <span class="multilang" lang="no">
                (1) La \(A\) v&#230;re punktet \((a, a^{@b@})\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (2) Let \(B\) be the point \( (a + h, (a + h)^{@b@}) \).
            </span>
            <span class="multilang" lang="no">
                (2) La \(B\) v&#230;re punktet \( (a + h, (a + h)^{@b@}) \).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (3) Find the gradient of the line \(AB\).
            </span>
            <span class="multilang" lang="no">
                (3) Finn stigningen til linjen \(AB\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (4) Let \(h \to 0\) to find the gradient of the curve at \(A\).
            </span>
            <span class="multilang" lang="no">
                (4) La \(h \to 0\) for &#229; finne stigningen til kurven i \(A\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
            </span>
        </p>
        <p>
            \( \dfrac{dy}{dx} \) = [[input:ans2]]
                                                [[validation:ans2]][[feedback:prt2]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                The gradient of the line \(AB\) is
            </span>
            <span class="multilang" lang="no">
                Stigningen til linjen \(AB\) er
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                \[ \dfrac{&#160;{@a1@}(a + h)^2 - {@a1@}a^2 }{h} = \dfrac{ {@a1@}(a^2 + 2ah + h^2) - {@a1@}a^2}{h} = \dfrac{ {@a1@}(2ah + h^2)}{h} = {@2*a1@}a + h \]
            </span>
            <span class="multilang" lang="no">
                \[ \dfrac{&#160;{@a1@}(a + h)^2 - {@a1@}a^2 }{h} = \dfrac{ {@a1@}(a^2 + 2ah + h^2) - {@a1@}a^2}{h} = \dfrac{ {@a1@}(2ah + h^2)}{h} = {@2*a1@}a + h \]
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Letting \( h \to 0 \) gives the gradient at \(A\) as \({@2*a1@}a\).
            </span>
            <span class="multilang" lang="no">
                N&#229;r \( h \to 0 \), blir stigningen ved \(A\) \({@2*a1@}a\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Thus \( \dfrac{dy}{dx} = {@ta1@} \)
            </span>
            <span class="multilang" lang="no">
                Derfor blir \( \dfrac{dy}{dx} = {@ta1@} \)
            </span>
        </p>
        <br/>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                The gradient of the line \(AB\) is
            </span>
            <span class="multilang" lang="no">
                Stigningen til linjen \(AB\) er
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                \[ \dfrac{&#160;&#160;(a + h)^{@b@} - a^{@b@} }{h} = \dfrac{ {@ex@} - a^{@b@}}{h} = \dfrac{ {@ex1@} }{h} = {@ex2@} \]
            </span>
            <span class="multilang" lang="no">
                \[ \dfrac{&#160;&#160;(a + h)^{b} - a^{b} }{h} = \dfrac{ eks - a^{b}}{h} = \dfrac{ eks1 }{h} = eks2 \]
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Letting \( h \to 0 \) gives the gradient at \(A\) as \({@b@}a^{@b-1@} \).
            </span>
            <span class="multilang" lang="no">
                N&#229;r vi lar \( h \to 0 \), f&#229;r vi stigningstallet ved \(A\) som \({@b@}a^{@b-1@} \).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Thus \( \dfrac{dy}{dx} = {@ta2@} \)
            </span>
            <span class="multilang" lang="no">
                Dermed blir  \( \dfrac{dy}{dx} = {@ta2@} \) oversatt til "Dermed blir $\dfrac{dy}{dx} = {@ta2@}$".
            </span>
        </p>
        <br/>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: rand([2,3,4,5,6]);
ta1:diff(a1*x^2,x);

b:rand([3]);
ex:expand((a+h)^3);
ex1:ex-a^b;
ex2:simplify(ex1/h);
ta2:diff(x^b,x);
</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( y = {@a1@}x^2 \Rightarrow y'={@ta1@}, \quad y = x^{@b@} \Rightarrow y'={@ta2@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <br/>
</p>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>867377300</deployedseed>
        <deployedseed>1225449357</deployedseed>
        <deployedseed>499458877</deployedseed>
        <deployedseed>206670610</deployedseed>
        <deployedseed>1583780896</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>