<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.1.4.2 differentiate and evaluate</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    Exercise
</p>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            Find the gradient of the following graphs at the given points:
        </span>
        <span class="multilang" lang="no">
            Finne gradienten av de f&#248;lgende grafene p&#229; de f&#248;lgende punkter:
        </span>
    </p>
    <li>
        <p>
        </p>
        <p>
        </p>
        \( f(x) = {@f1(x)@} \) at \(x = {@n1@}\)
        <p>
        </p>
        <p>
        </p>
        <p>
            gradient = [[input:ans1]]
                                                [[validation:ans1]][[feedback:prt1]]
        </p>
    </li>
    <li>
        \( f(x) = {@f2(x)@} \)&#160;at \(x = {@n2@}\)
        <p>
        </p>
        <p>
            gradient =&#160; [[input:ans2]]
                                                [[validation:ans2]][[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
        </p>
        <p>
        </p>
        \( f(x) = {@f3(x)@} \)&#160;at \(x = {@n3@}\)
        <p>
        </p>
        <p>
        </p>
        <p>
            gradient = [[input:ans3]]
                                                [[validation:ans3]][[feedback:prt3]]
        </p>
    </li>
    <li>
        <p>
        </p>
        <p>
        </p>
        \( f(x) = \displaystyle {@f4(x)@} \)&#160;at \(x = {@n4@}\)
        <p>
        </p>
        <p>
        </p>
        <p>
            gradient = [[input:ans4]]
                                                [[validation:ans4]][[feedback:prt4]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We have
                                \[ \dfrac{df}{dx}= {@df1@} \] which, when evaluated at \( x = {@n1@} \) is \( {@ta1@}\)
            </span>
            <span class="multilang" lang="no">
                Vi har
                                \[ \dfrac{df}{dx}= {@df1@} \] hvis, n&#229;r den er evaluert p&#229; \( x = {@n1@} \) er \( {@ta1@}\)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We have

    
    
    
    \[ \dfrac{df}{dx} = {@df2@} \] which, when evaluated at \( x = {@n2@} \) is \( {@ta2@}\)
            </span>
            <span class="multilang" lang="no">
                Vi har

    
    
    
    \[ \dfrac{df}{dx} = {@df2@} \] som, n&#229;r evaluert ved \( x = {@n2@} \), er \( {@ta2@}\)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We have

    
    
    
    \[ \dfrac{df}{dx} = {@df3@} \] which, when evaluated at \( x = {@n3@} \) is \( {@ta3@}\)
            </span>
            <span class="multilang" lang="no">
                Vi har 




\[ \dfrac{df}{dx} = {@df3@} \]




n&#229;r vi evaluerer denne ved \( x = {@n3@} \), f&#229;r vi \( {@ta3@}\) .
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We have

    
    
    
    \[ \dfrac{df}{dx} = {@df4@} \] which, when evaluated at \( x = {@n4@} \) is \( {@ta4@}\)
            </span>
            <span class="multilang" lang="no">
                Vi har




\[ \dfrac{df}{dx} = {@df4@} \],




som n&#229;r evaluert ved \( x = {@n4@} \) gir \( {@ta4@}\).
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: rand([2,3,4]);
n1:-2;
f1(x):=a1*x^2;
a2: expand(f1(x+h)) - f1(x);
a3: simplify(a2/h);
df1:diff(f1(x),x);
ta1:2*a1*n1


b:rand([2,3,4]);
n2:2;
f2(x):=b*x^3;
b2: expand(f2(x+h)) - f2(x);
b3: simplify(b2/h);
df2:diff(f2(x),x);
ta2:3*b*n2^2

ca:rand([6,7,8]);
cb:rand([3,4]);
n3:-5;
f3(x):=ca*x + cb;
c2: expand(f3(x+h)) - f3(x);
c3: simplify(c2/h);
df3:diff(f3(x),x);
ta3:ca

d1:rand([1]);
n4:0.5;
f4(x):=d1/x;
d2: expand(f4(x+h)) - f4(x);
d3: simplify(d2/h);
df4:diff(f4(x),x);
ta4:-n4^(-2)

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[ f(x) = {@f1(x)@}, \, x = {@n1@}, \, f'= {@ta1@} \]
\[ f(x) = {@f2(x)@}, \, x = {@n2@}, \, f'= {@ta2@} \]
\[ f(x) = {@f3(x)@}, \, x = {@n3@}, \, f'= {@ta3@} \]
\[ f(x) = {@f4(x)@}, \, x = {@n4@}, \, f'= {@ta4@} \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>867377300</deployedseed>
        <deployedseed>1225449357</deployedseed>
        <deployedseed>499458877</deployedseed>
        <deployedseed>2134512923</deployedseed>
        <deployedseed>1204641410</deployedseed>
        <deployedseed>1356534604</deployedseed>
        <deployedseed>865722290</deployedseed>
        <deployedseed>1701587251</deployedseed>
        <deployedseed>1896902844</deployedseed>
        <deployedseed>1249906673</deployedseed>
        <deployedseed>708033166</deployedseed>
        <deployedseed>1389447944</deployedseed>
        <deployedseed>1401325268</deployedseed>
        <deployedseed>2011594621</deployedseed>
        <deployedseed>1892016619</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3+1</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>