<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2.1.3.4 Composite function 3</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If \(f\left(x\right) = {@fx@}\), find \(f\left(f\left(x\right)\right)\).
    </span>
    <span class="multilang" lang="no">
        Hvis \(f\left(x\right) = {@fx@}\), finn \(f\left(f\left(x\right)\right)\).
    </span>
</p>
<p>
    [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>
<p>
    <br/>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We substitute \(f\left( x \right)= {@fx@}\) back into it self, so \[f \left( f(x) \right) = f \left ({@fx@} \right)= {@ffx@}={@expffx@}.\]
    </span>
    <span class="multilang" lang="no">
        Vi setter \(f\left( x \right)= {@fx@}\) inn i seg selv, s&#229; \[f \left( f(x) \right) = f \left ({@fx@} \right)= {@ffx@}={@expffx@}.\]
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022060100</text>
        </stackversion>
        <questionvariables>
            <text>amag:rand(8)+2;
 asign:2*rand(2) -1;
 a:amag * asign;
 bmag:rand(8)+2;
 bsign:2*rand(2) -1;
 b:bmag * bsign;
 c:rand(11)-5;
 mb:-b;
 mc:-c;

fx:b*x + a;

ffx:b*fx+a;
expffx: expand(ffx);
ta1: expffx;

 sc:sconcat(c);
 smb:sconcat(mb);
 smc:sconcat(mc);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>Given \(f\left(x\right) = {@fx@}\)  then \(f\left(f\left(x\right)\right)={@expffx@}\) .</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1504293482</deployedseed>
        <deployedseed>1998130634</deployedseed>
        <deployedseed>1465004566</deployedseed>
        <deployedseed>1123752268</deployedseed>
        <deployedseed>87106695</deployedseed>
        <deployedseed>2031804246</deployedseed>
        <deployedseed>1892774009</deployedseed>
        <deployedseed>720934101</deployedseed>
        <deployedseed>769044697</deployedseed>
        <deployedseed>1247711650</deployedseed>
        <deployedseed>1225206925</deployedseed>
        <deployedseed>1059976966</deployedseed>
        <deployedseed>157846395</deployedseed>
        <deployedseed>1198171971</deployedseed>
        <deployedseed>104843417</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1-2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>fx*fx</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>