<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2.1.2.3 Evaluate value of function and domain</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Suppose \( \displaystyle G\left({@v@}\right) = {@fv@}\).
    </span>
    <span class="multilang" lang="no">
        Gitt \( \displaystyle G\left({@v@}\right) = {@fv@}\).
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Evaluate: \( G\left({@v_val@}\right)=\)
            </span>
            <span class="multilang" lang="no">
                Evaluer: \( G\left({@v_val@}\right)=\)
            </span>
            [[input:ans1]]. [[validation:ans1]] [[feedback:Result]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Find a value of \( {@v@} \) such that&#160;\(G\left({@v@}\right) \) does not exist. \({@v@}=\)
            </span>
            <span class="multilang" lang="no">
                Finn en verdi for \( {@v@} \) slik at&#160;\(G\left({@v@}\right) \) ikke eksisterer. \({@v@}=\)
            </span>
            [[input:ans2]]. [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
    <p>
        <br/>
    </p>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We must substitute {@v_val@} in place of {@v@} in the definition of the function:
            </span>
            <span class="multilang" lang="no">
                Vi m&#229; substituere {@v@} med {@v_val@} i definisjonen for funksjonen:
            </span>
        </p>
        \[\begin{array}{rcl} G\left({@v_val@}\right) = \displaystyle \frac{{@a@}\cdot({@v_val@})+({@b@})}{{@c@}\cdot({@v_val@})+({@d@})} = \displaystyle \frac{{@num_val@}}{{@den_val@}}
                                = \displaystyle {@ta1@} \end{array}\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We need to find the value of {@v@} for which the denominator is \(0\). So,
            </span>
            <span class="multilang" lang="no">
                Vi m&#229; finne verdien for {@v@} s&#229; nevneren er \(0\). Da er,
            </span>
        </p>
        \[ {@den@}=0 \\ {@v@} = {@ta2@} \]
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>/* Function G(x) = (a*x+b)/(c*x+d) */
v: rand([x,y,z]);
a: rand_with_prohib(-4,4,[0]);
b_magnitude :rand(4)+2;
b_sign: rand_with_prohib(-1,1,[0]);
b :b_magnitude*b_sign;
c: rand(2)+1;
d: rand_with_prohib(-5,5,[0]);
num: a*v+b;
den: c*v+d;
fv: num/den;
v_val: rand_with_prohib(-3,3,[-d/c]);   /*Exclude the value that makes denominator 0 */
ta1: (a*v_val+b)/(c*v_val+d);
num_val: a*v_val+b;
den_val: c*v_val+d;
ta2: -d/c;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>Given \(G\left({@v@}\right) = {@fv@}\) then \(G\left({@v_val@}\right) = {@ta1@} \), and \(G\left({@v@}\right) \) does not exit when  {@v@} = {@ta2@}.</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>*,+,</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>*</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>Result</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>FNC-EVL-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>FNC-EVL-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1300386662</deployedseed>
        <deployedseed>1373118239</deployedseed>
        <deployedseed>472023574</deployedseed>
        <deployedseed>304421122</deployedseed>
        <deployedseed>2080816889</deployedseed>
        <deployedseed>1218903653</deployedseed>
        <deployedseed>1634564423</deployedseed>
        <deployedseed>276193739</deployedseed>
        <deployedseed>156345992</deployedseed>
        <deployedseed>1540114888</deployedseed>
        <deployedseed>990512950</deployedseed>
        <deployedseed>703133214</deployedseed>
        <deployedseed>537356561</deployedseed>
        <deployedseed>1220751083</deployedseed>
        <deployedseed>375361921</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>Result</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>FNC-EVL-TRUE</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1-1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2-1</value>
            </testinput>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>Result</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>FNC-EVL-FALSE</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>