<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>3.2.3.4 Completing the square in practice</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        A chemical manufacturer found that the sales figures for a certain chemical X\(_2\)O depended on its selling price. At present, the company can sell all of its weekly production of \(300\text{kg}\) at a price of
        <b>
            &#163;
        </b>
        \(600\) per \(\text{kg}\). The company's market research department advised that the amount sold would decrease by only \(1\)&#160;\(\text{kg}\) per week for every
        <b>
            &#163;
        </b>
        \(2\) per \(\text{kg}\) increase in the price of X\(_2\)O. If the total production costs are made up of a fixed cost of
        <b>
            &#163;
        </b>
        \(30000\) per week, plus
        <b>
            &#163;
        </b>
        \(400\) per&#160;\(\text{kg}\) of product.
    </span>
    <span class="multilang" lang="no">
        En kjemikalieprodusent fant at salgstallene for et bestemt kjemikalie X\(_2\)O var avhengig av salgsprisen. For tiden kan selskapet selge hele sin ukentlige produksjon av \(300\text{kg}\) til en pris p&#229; \(600\) kr per \(\text{kg}\). Selskapets markedsunders&#248;kelsesavdeling informerte om at det solgte bel&#248;pet ville reduseres med bare \(1\)&#160;\(\text{kg}\) per uke for hver \(2\) kr per \( \text{kg}\) &#248;kning i prisen p&#229; X\(_2\)O. Hvis de totale produksjonskostnadene best&#229;r av en fast kostnad p&#229; \(30000\) kr per uke, pluss
        <b>
        </b>
        \(400\) kr per&#160;\(\text{ kg}\) av produktet.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Let
        <b>
            &#163;
        </b>
        \(x\) be the new price per \(\text{kg}\) of X\(_2\)O.&#160;Find a quadratic polynomial \(P\) in \(x\) which gives the weekly profit.
    </span>
    <span class="multilang" lang="no">
        La
        <b>
        </b>
        \(x\) kr v&#230;re den nye prisen per \(\text{kg}\) av X\(_2\)O.&#160;Finn et kvadratisk polynom \(P\) i \(x\) som gir den ukentlige fortjenesten.
    </span>
</p>
<p style="text-align: center;">
    \(P =\) [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
</p>
<p style="text-align: left;">
    <span class="multilang" lang="en">
        Complete the square for the above&#160; expression and hence find
    </span>
    <span class="multilang" lang="no">
        Fullf&#248;r kvadratet over&#160; uttrykk og finn dermed
    </span>
</p>
<p>
</p>
<ol>
    <li>
        <span class="multilang" lang="en">
            the price which maximises the weekly profit on sales of X\(_2\)O;
            <b>
                &#163;
            </b>
        </span>
        <span class="multilang" lang="no">
            prisen som maksimerer den ukentlige fortjenesten ved salg av X\(_2\)O;
            <b>
            </b>
        </span>
        [[input:ans1]] [[validation:ans1]][[feedback:prt1]] kr
    </li>
    <li>
        <span class="multilang" lang="en">
            the maximum weekly profit;
        </span>
        <span class="multilang" lang="no">
            maksimal ukentlig fortjeneste;
        </span>
        <b>
        </b>
        [[input:ans2]] [[validation:ans2]][[feedback:prt2]] kr
    </li>
    <li>
        <span class="multilang" lang="en">
            the weekly production rate.
        </span>
        <span class="multilang" lang="no">
            den ukentlige produksjonsraten.
        </span>
        [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Firstly, we must derive the equation for weekly profit. If \(x\) is the new price per kg of the chemical, then the new amount produced is
    </span>
    <span class="multilang" lang="no">
        For det f&#248;rste m&#229; vi utlede ligningen for ukentlig fortjeneste. Hvis \(x\) er den nye prisen per kg av kjemikaliet, s&#229; er den nye produserte mengden
    </span>
</p>
<p>
    \[300-\frac{(x-600)}{2} = 600-\frac{x}{2}.\]
</p>
<p>
    <span class="multilang" lang="en">
        Now we must work out the profit. This will be what we make from selling the product minus any costs.
    </span>
    <span class="multilang" lang="no">
        N&#229; m&#229; vi regne ut overskuddet. Dette vil v&#230;re hva vi tjener p&#229; &#229; selge produktet minus eventuelle kostnader.
    </span>
</p>
<p>
    \[P = x\left(600-\frac{x}{2}\right)-30000-400\left(600-\frac{x}{2}\right) =\]
</p>
<p>
    \[-\dfrac{x^2}{2}+800x-270000.\]
</p>
<p>
    <span class="multilang" lang="en">
        Completing the square we obtain
    </span>
    <span class="multilang" lang="no">
        Ved &#229; fullf&#248;re kvadratet f&#229;r vi
    </span>
</p>
<p>
    \[P = -\frac{1}{2}(x^2-1600x+540000) = -\frac{1}{2}((x-800)^2+540000-800^2 =\]
</p>
<p>
    \[ -\frac{1}{2}((x-800)^2+540000-800^2) =-\frac{1}{2}((x-800)^2-100000).\]
</p>
<p>
</p>
<ol>
    <li>
        <span class="multilang" lang="en">
            We want to maximise \(P\). This will occur when \((x-800)^2=0\), so \(x=800\).
        </span>
        <span class="multilang" lang="no">
            Vi &#248;nsker &#229; maksimere \(P\). Dette vil skje n&#229;r \((x-800)^2=0\), s&#229; \(x=800\).
        </span>
    </li>
    <li>
        At \(x=800\), \(P = 50000\).
    </li>
    <li>
        <span class="multilang" lang="en">
            The production rate is \(600-\frac{x}{2}\). Subbing in our answer from part 1 we obtain
        </span>
        <span class="multilang" lang="no">
            Produksjonsraten er \(600-\frac{x}{2}\). Ved &#229; sende inn v&#229;rt svar fra del 1 f&#229;r vi
        </span>
        \[600-\frac{800}{2} = 200.\]
    </li>
</ol>
<p>
</p>
<p>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ta1:800;
ta2:50000;
ta3:200;
tpoly:-x^2/2+800*x-270000;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text/>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>tpoly</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>diff(diff(diff(ans4,x),x),x)</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans><![CDATA[is(coeff(ans4,x,2)<0)]]></sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>-coeff(ans4,x)/(2*coeff(ans4,x,2))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>tpoly</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>tpoly</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta1+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(ta2+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ev(ta3+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ev(tpoly+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-4-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3+3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-3-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>