<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>3.2.2.2 Solve the quadratic equation (integer roots)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the quadratic equation \({@exp@}=0\). If there is a repeated root, enter this twice.
    </span>
    <span class="multilang" lang="no">
        L&#248;s andregradsligningen \({@exp@}=0\). Hvis det er en dobbel rot, skriv inn denne to ganger.
    </span>
</p>
<p>
    <nowrap>
        \(x = \) [[input:ans1]] [[validation:ans1]] or \(x = \) [[input:ans2]][[validation:ans2]]
        <nowrap>
        </nowrap>
    </nowrap>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are asked to factorise the quadratic equation \({@exp@}=0.\)
    </span>
    <span class="multilang" lang="no">
        Vi blir bedt om &#229; l&#248;se andregradsligningen \({@exp@}=0.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        First factorise the left-hand side:
    </span>
    <span class="multilang" lang="no">
        Vi faktoriser f&#248;rst venstre side:
    </span>
</p>
<p>
    \[{@exp@}={@exp1@}.\]
</p>
<p>
    <span class="multilang" lang="en">
        Equate each factor to zero to obtain the two solutions:
    </span>
    <span class="multilang" lang="no">
        Sett hver faktor til null for &#229; f&#229; de to l&#248;sningene:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solution 1:
    </span>
    <span class="multilang" lang="no">
        L&#248;sning 1:
    </span>
    \(x= {@aa@}\)
</p>
<p>
    <span class="multilang" lang="en">
        Solution 2:
    </span>
    <span class="multilang" lang="no">
        L&#248;sning 2:
    </span>
    \(x= {@dd@}\).
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022082900</text>
        </stackversion>
        <questionvariables>
            <text>aa:1+rand(5);
bb:rand_with_prohib(1,5,[aa]);
cc : (-1)^(bb);
dd : cc*bb;
exp1: (x-aa)*(x-dd);
exp : expand(exp1);
ta1:aa;
ta2:dd;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\[{@exp@}=0 \text{ so } x_{1,2} = {@ta1@}, {@ta2@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>{ans1,ans2}</sans>
                <tans>{ta1,ta2}</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    You answered correctly!
</span>
<span class="multilang" lang="no">
    Du svarte rett!
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    You answered wrong!
</span>
<span class="multilang" lang="no">
    Du svarte feil!
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1441005928</deployedseed>
        <deployedseed>646728586</deployedseed>
        <deployedseed>2129230970</deployedseed>
        <deployedseed>1274810135</deployedseed>
        <deployedseed>699206350</deployedseed>
        <deployedseed>208893348</deployedseed>
        <deployedseed>21086862</deployedseed>
        <deployedseed>596421197</deployedseed>
        <deployedseed>1469182629</deployedseed>
        <deployedseed>1746074607</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>