<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>3.2.5.3 Solve the quadratic equation (rationals) (copy for end)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the quadratic equation \({@exp@}=0\). If there is a repeated root, enter this twice.
    </span>
    <span class="multilang" lang="no">
        L&#248;s andregradsligningen \({@exp@}=0\). Hvis det er en dobbel rot, skriv inn denne to ganger.
    </span>
</p>
<p>
    <nowrap>
        \(x = \)[[input:ans1]] [[validation:ans1]]
        <span class="multilang" lang="en">
            or
        </span>
        <span class="multilang" lang="no">
            eller
        </span>
        \(x = \)[[input:ans2]][[validation:ans2]]
        <nowrap>
        </nowrap>
    </nowrap>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given the quadratic equation \({@exp@}=0.\)
    </span>
    <span class="multilang" lang="no">
        Vi har andregradsligningen \({@exp@}=0.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        First factorise the left-hand side:
    </span>
    <span class="multilang" lang="no">
        Faktoriser f&#248;rst venstre side:
    </span>
</p>
<p>
    \[{@exp@}={@exp1@}.\]
</p>
<p>
    <span class="multilang" lang="en">
        Equate each factor is then equated to zero to obtain the two solutions:
    </span>
    <span class="multilang" lang="no">
        Sett hver faktor lik 0 for &#229; f&#229; de to l&#248;sningene:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solution
    </span>
    <span class="multilang" lang="no">
        L&#248;sning
    </span>
    1: \(x=&#160;{@ta1@}\)
</p>
<p>
    <span class="multilang" lang="en">
        Solution
    </span>
    <span class="multilang" lang="no">
        L&#248;sning
    </span>
    2: \(x= {@ta2@}\).
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa:1+rand(5)
bb:1+rand(5)
exp1: (2*x-aa)*(3*x-bb)
exp : expand(exp1)
ta1:aa/2
ta2:bb/3</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\({@exp@}=0\) roots \( x= {@ta1@}, x = {@ta2@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>{ratsimp(ans1),ratsimp(ans2)}</sans>
                <tans>{ratsimp(ta1),ratsimp(ta2)}</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>296953638</deployedseed>
        <deployedseed>200420641</deployedseed>
        <deployedseed>1741041881</deployedseed>
        <deployedseed>2140628164</deployedseed>
        <deployedseed>225101550</deployedseed>
        <deployedseed>1431006916</deployedseed>
        <deployedseed>419462610</deployedseed>
        <deployedseed>1605014443</deployedseed>
        <deployedseed>519137390</deployedseed>
        <deployedseed>888782538</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>