<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>3.5.2.2_Task_Solve Inequality</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the inequality
    </span>
    <span class="multilang" lang="no">
        L&#248;s ulikheten
    </span>
    \({@leftvar@}{@var@}+{@leftcoef@}&lt;{@rightvar@}{@var@}+{@rightcoef@}\).
</p>
<p>
    <span class="multilang" lang="en">
        This is done by making \({@var@}\) the subject and obtain it on its own on the left-hand side.
    </span>
    <span class="multilang" lang="no">
        Dette gj&#248;res ved &#229; isolere \({@var@}\) p&#229; venstre side av ulikheten.
    </span>
</p>
<br/>
<p>
    <span class="multilang" lang="en">
        Start by subtracting \({@rightvar@}{@var@}\) from both sides to remove quantities involving \({@var@}\) from the right:
    </span>
    <span class="multilang" lang="no">
        Start med &#229; trekke \({@rightvar@}{@var@}\) fra begge sider av ulikheten for &#229; fjerne mengder som involverer \({@var@}\) fra h&#248;yre side:
    </span>
</p>
<br/>
<p>
    [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
</p>
<br/>
<p>
    <span class="multilang" lang="en">
        Now subtract \({@leftcoef@}\) from both sides to remove the \({@leftcoef@}\) on the left:
    </span>
    <span class="multilang" lang="no">
        Trekk n&#229; \({@leftcoef@}\) fra begge sider for &#229; fjerne \({@leftcoef@}\) p&#229; venstre side:
    </span>
</p>
<br/>
<p>
    [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
</p>
<p>
    <span class="multilang" lang="en">
        Finally, the range of values of \({@var@}\) is:
    </span>
    <span class="multilang" lang="no">
        Til slutt er det bestemt at verdiomr&#229;det for \({@var@}\) er:
    </span>
</p>
<br/>
<p>
    [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    We have the inequality
</span>
<span class="multilang" lang="no">
    Vi har ulikheten
</span>
\[{@leftvar@}{@var@}+{@leftcoef@}&lt;{@rightvar@}{@var@}+{@rightcoef@}\]
<p>
</p>
<p>
    <span class="multilang" lang="en">
        We start by subtracting \({@rightvar@}{@var@}\) from both sides to remove quantities involving \({@var@}\)&#160;from the right. Hence, we get:
    </span>
    <span class="multilang" lang="no">
        Vi starter ved &#229; trekke \({@rightvar@}{@var@}\) fra begge sider for &#229; fjerne mengder som involverer \({@var@}\)&#160;fra h&#248;yre. Dermed f&#229;r vi:
    </span>
    \[ {@correctanswer1@}\]
</p>
<p>
    <span class="multilang" lang="en">
        Now subtract \({@leftcoef@}\) from both sides to remove the \({@leftcoef@}\) on the left:
    </span>
    <span class="multilang" lang="no">
        Trekk n&#229; \({@leftcoef@}\) fra begge sider for &#229; fjerne \({@leftcoef@}\) p&#229; venstre side:
    </span>
    \[ {@correctanswer2@} \]
</p>
<p>
    <span class="multilang" lang="en">
        Finally, the range of values of \({@var@}\) is:
    </span>
    <span class="multilang" lang="no">
        Da f&#229;r vi endelig at verdiomr&#229;det for \({@var@}\) er:
    </span>
    \[ {@correctanswer3@} \]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022060100</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[var:x;
leftvar: rand_with_step(10,25,1);
rightvar: rand_with_step(1,9,1);
leftcoef: rand(9)+1;
rightcoef: rand(9)+1;
newcoef: rightcoef-leftcoef;
newvar: leftvar-rightvar
newrange: newcoef/newvar;
correctanswer1: newvar*var+leftcoef<rightcoef;
correctanswer2: newvar*var<newcoef;
correctanswer3: var<newrange;


]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text><![CDATA[\({@leftvar@}{@var@}+{@leftcoef@}&lt;{@rightvar@}{@var@}+{@rightcoef@}\), \mbox{ the first answer is: } {@correctanswer1@} \mbox{ the second part is: }{@correctanswer2@}  \mbox{ the third part is: }{@correctanswer3@}]]></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>correctanswer1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>(,)</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>correctanswer2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>(,)</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>correctanswer3</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>PropLogic</answertest>
                <sans>ans1</sans>
                <tans>correctanswer1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text>The correct answer is \({@correctanswer1@}\).</text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>PropLogic</answertest>
                <sans>ans2</sans>
                <tans>correctanswer2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ans1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text>You answer here is correct, but your subtraction above is wrong.</text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text>The correct answer is \({@correctanswer2@}\).</text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>PropLogic</answertest>
                <sans>ans3</sans>
                <tans>correctanswer3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ans2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text>This is correct, but you calculations above are incorrect.</text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text>The correct answer is \({@correctanswer3@}\).</text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1389453098</deployedseed>
        <deployedseed>40422752</deployedseed>
        <deployedseed>1791475607</deployedseed>
        <deployedseed>924068491</deployedseed>
        <deployedseed>1238709344</deployedseed>
        <deployedseed>98624745</deployedseed>
        <deployedseed>1474340704</deployedseed>
        <deployedseed>2106013424</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>correctanswer1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>correctanswer2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>correctanswer3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA[x<10000000]]></value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value><![CDATA[x<1000000000]]></value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA[x<100000000]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>