<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>13.3.3.2c Finding areas</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            Find the area bounded by \(y= e^x\), the \(y\)-axis and the line \(x=2\).
        </span>
        <span class="multilang" lang="no">
            Finn arealet avgrenset av \(y= e^x\), \(y\)-aksen og linja \(x=2\).
        </span>
        [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]
    </li>
    <li>
        <span class="multilang" lang="en">
            Find the area enclosed between \(y=x(x-{@aa@})(x-{@bb@})\) and the \(x\)-axis.
        </span>
        <span class="multilang" lang="no">
            Finn arealet avgrenset av \(y=x(x-{@aa@})(x-{@bb@})\) og \(x\)-aksen.
        </span>
        [[input:ans6]] [[validation:ans6]] [[feedback:prt6]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            The curve \(y= e^x\), the \(y\)-axis and the line \(x=2\) intersect at \(x=0,2\).This area is above the \(x\) axis hence the area is \[\int_0^2 e^x \, dx = {@ta5@}.\]
        </span>
        <span class="multilang" lang="no">
            Kurva \(y= e^x\), \(y\)-aksen og linja \(x=2\) skj&#230;rer hverandre i \(x=0,2\). Dette omr&#229;det ligger over \(x\)-aksen og dermed blir arealet \[\int_0^2 e^x \, dx = {@ta5@}.\]
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            The curve \(y={@exp6@}\) and the \(x\) axis intersect at \(x = 0,{@aa@},{@bb@}\). Between \(0\) and {@aa@}, the area under the curve is above the \(x\) axis and between {@aa@} and {@bb@} , it is below the \(x\) axis. Hence the area is: \[\int_{@l2@}^{@aa@}
                                x(x-{@aa@})(x-{@bb@}) \, dx - \int_{@aa@}^{@bb@} x(x-{@aa@})(x-{@bb@}) \, dx =\] \[\int_{@l2@}^{@aa@} ({@expand(exp6)@}) \, dx - \int_{@aa@}^{@bb@}({@expand(exp6)@}) \, dx =\] \[ \left[{@int(exp6,x)@} \right]_{@l2@}^{@aa@} -\left[{@int(exp6,x)@}\right]_{@aa@}^{@bb@}
                                = {@ta6@} .\]
        </span>
        <span class="multilang" lang="no">
            Kurva \(y={@exp6@}\) og \(x\)-aksen skj&#230;rer hverandre i \(x = 0,{@aa@},{@bb@}\). Mellom \(0\) og {@aa@} ligger ligger arealet under grafen over \(x\)-aksen, mens mellom {@aa@} og {@bb@} , ligger det under \(x\)-aksen. Dermed blir arealet: \[\int_{@l2@}^{@aa@}
                                x(x-{@aa@})(x-{@bb@}) \, dx - \int_{@aa@}^{@bb@} x(x-{@aa@})(x-{@bb@}) \, dx =\] \[\int_{@l2@}^{@aa@} ({@expand(exp6)@}) \, dx - \int_{@aa@}^{@bb@}({@expand(exp6)@}) \, dx =\] \[ \left[{@int(exp6,x)@} \right]_{@l2@}^{@aa@} -\left[{@int(exp6,x)@}\right]_{@aa@}^{@bb@}
                                = {@ta6@} .\]
        </span>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>exp5: e^x;
ta5: int(exp5, x, 0, 2); 

/* polynomial function*/
aa: rand(3)+1;
bb: aa+1+rand(2);
exp6: x*(x-aa)*(x-bb);
l2:0;
u2:aa;
ta6: int(exp6, x, l2, aa)+ abs(int(exp6, x, aa, bb));</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[\int_0^2 e^x \, dx = {@ta5@}, \quad \int_{@l2@}^{@aa@} {@exp6@} \, dx - \int_{@aa@}^{@bb@} {@exp6@} \, dx = {@ta6@} \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans5</name>
            <type>numerical</type>
            <tans>ta5</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans6</name>
            <type>numerical</type>
            <tans>ta6</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>ta5</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>ta6</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>