<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>13.3.3.2a Finding areas</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            Find the total area enclosed between the \(x\)-axis and the curve \(y={@exp1@}\) between \(x={@l1@}\) and \(x={@u1@}\).
        </span>
        <span class="multilang" lang="no">
            Finn det totale arealet mellom \(x\)-aksen og kurven \(y={@exp1@}\) mellom \(x={@l1@}\) og \(x={@u1@}\).
        </span>
        [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        <span class="multilang" lang="en">
            Find the area under \(y={@exp2@}\) from \(t={@l2@}\) to \(t={@u2@}\). Give your answer to \(4\) significant figures.
        </span>
        <span class="multilang" lang="no">
            Finn arealet under \(y={@exp2@}\) fra \(t={@l2@}\) til \(t={@u2@}\). Angi svaret ditt med \(4\) gjeldende siffer.
        </span>
        [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
    <li>
        <span class="multilang" lang="en">
            Find the area enclosed by \(y={@exp3@}\) and the \(x\)-axis from \(x={@l3@}\) to \(x={@u3@}\).
        </span>
        <span class="multilang" lang="no">
            Finn arealet mellom \(y={@exp3@}\) og \(x\)-aksen fra \(x={@l3@}\) til \(x={@u3@}\).
        </span>
        [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            Since the curve \(y={@exp1@}\) is above the \(x\) axis between 0 and {@u1@}, and below the \(x\) axis between {@l1@} and 0. The area under the curve is: \[-\int_{@l1@}^{0} {@exp1@} \, dx + \int_{0}^{@u1@} {@exp1@} \, dx = {@ta1@}.\]
        </span>
        <span class="multilang" lang="no">
            Siden kurven \(y={@exp1@}\) ligger over \(x\)-aksen mellom 0 og {@u1@}, og under \(x\)-aksen mellom {@l1@} og 0. Arealet blir derfor: \[-\int_{@l1@}^{0} {@exp1@} \, dx + \int_{0}^{@u1@} {@exp1@} \, dx = {@ta1@}.\]
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            The curve \(y={@exp2@}\) is above the \(t\) axis from \(t={@l2@}\) to \(t={@u2@}\). The area is: \[\int_{@l2@}^{@u2@} {@exp2@} \, dt = \left[ {@int(exp2,t)@} \right]_{@l2@}^{@u2@} = {@ta2@}.\]
        </span>
        <span class="multilang" lang="no">
            Kurven \(y={@exp2@}\) ligger over \(t\)-aksen fra \(t={@l2@}\) til \(t={@u2@}\). Arealet blir: \[\int_{@l2@}^{@u2@} {@exp2@} \, dt = \left[ {@int(exp2,t)@} \right]_{@l2@}^{@u2@} = {@ta2@}.\]
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            The area under the curve \(y={@exp3@}\) is above the \(x\) axis from \(x={@l3@}\) to \(x={@u3@}\), hence the area is: \[ \int_{@l3@}^{@u3@} ({@exp3@}) \, dx =\left[ {@int(exp3,x)@} \right]_{@l3@}^{@u3@} = {@ta3@}.\]
        </span>
        <span class="multilang" lang="no">
            Arealet under kurven \(y={@exp3@}\) ligger over \(x\)-aksen fra \(x={@l3@}\) til \(x={@u3@}\), dermed blir arealet: \[ \int_{@l3@}^{@u3@} ({@exp3@}) \, dx =\left[ {@int(exp3,x)@} \right]_{@l3@}^{@u3@} = {@ta3@}.\]
        </span>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand([1,1/2, 1/3]);
exp1: aa*x^3;
l1: rand([-1,-2]);
u1: -l1;
ta1:  2*int(exp1, x, 0, l1); 

exp2: cos(2*t);
l2: 0;
u2: 0.5;
ta2: significantfigures(float(int(exp2, t, l2, u2)),4); 

bb: rand([4,9,16]);
exp3: bb-x^2;
l3: rand(2);
u3: abs(sqrt(bb));
ta3: int(exp3, x, l3, u3); </text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[-\int_{@l1@}^{0} {@exp1@} \, dx + \int_{0}^{@u1@} {@exp1@} \, dx = {@ta1@}, \quad  \int_{@l2@}^{@u2@} {@exp2@} \, dt = {@ta2@}, \quad \int_{@l3@}^{@u3@} {@exp3@} \, dx = {@ta3@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>dispsf(ta2,4)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>numerical</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumSigFigs</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>4</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>