<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>13.4.2.2a Bestemt integrasjon</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    <p class="HELM_exercise">
        Exercise
    </p>
    <ol class="HELM_parts">
        <p>
            Evaluate the following. Answers should be given as a decimal rounded to \(4\) d.p.
        </p>
        <li>
            \({\displaystyle\int_{@l1@}^{@u1@} {@exp1@}\mathrm{d}x}\).[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </li>
        <li>
            \({\displaystyle\int_{@l2@}^{@u2@} {@exp2@}\mathrm{d}t}\).[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </li>
        <li>
            \({\displaystyle\int_{@l3@}^{@u3@} {@exp3@}\mathrm{d}x}\).[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </li>
    </ol>
</span>
<span class="multilang" lang="no">
    <p class="HELM_exercise">
        Oppgave
    </p>
    <ol class="HELM_parts">
        <p>
            Regn ut f&#248;lgende integral. Rund av svarene til \(4\) gjeldende siffer.
        </p>
        <li>
            \({\displaystyle\int_{@l1@}^{@u1@} {@exp1@}\mathrm{d}x}\).[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </li>
        <li>
            \({\displaystyle\int_{@l2@}^{@u2@} {@exp2@}\mathrm{d}t}\).[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </li>
        <li>
            \({\displaystyle\int_{@l3@}^{@u3@} {@exp3@}\mathrm{d}x}\).[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </li>
    </ol>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    [[facts:calc_int_methods_parts]]
    <br/>
</p>
<p>
    <br/>
</p>
<ol class="HELM_parts">
    <li>
        \(\displaystyle{\int_{@l1@}^{@u1@} {@exp1@}\mathrm{d}x = \left[ {@int(exp1,x)@}\right]_{@l1@}^{@u1@} = {@ta1r@}={@ta1@}}.\)
    </li>
    <li>
        \(\displaystyle{\int_{@l2@}^{@u2@} {@exp2@}\mathrm{d}t = \left[ {@int(exp2,x)@}\right]_{@l2@}^{@u2@} ={@ta2r@}= {@ta2@}}.\)
    </li>
    <li>
        \(\displaystyle{\int_{@l3@}^{@u3@} {@exp3@}\mathrm{d}x=&#160;\left[ {@int(exp3,t)@}\right]_{@l3@}^{@u3@} = {@ta3r@}={@ta3@}}.\)
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand([2, 4, 6, 8]); 
bb: rand([3, 5, 7, 9]); 
cc: rand(4)+1;

exp1: x*cos(aa*x); 
exp2: x*sin(aa*x); 
exp3: t*e^(aa*t); 

l1: 0;
u1: 1;
ta1r: int(exp1, x, l1, u1);
ta1: decimalplaces(abs(float(int(exp1, x, l1, u1))),4);

l2: 0;
u2: pi/2; 
ta2r: int(exp2, x, l2, u2);
ta2: decimalplaces(float(int(exp2, x, l2, u2)),4); 

l3: -1;
u3: 1;
ta3r: int(exp3, t, l3, u3);
ta3: decimalplaces(float(int(exp3, t, l3, u3)),4); 
</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\(\int_{@l1@}^{@u1@} {@exp1@}\mathrm{d}x = {@ta1@}, \quad \int_{@l2@}^{@u2@} {@exp2@}\mathrm{d}t = {@ta2@}, \quad \int_{@l3@}^{@u3@} {@exp3@}\mathrm{d}x= {@ta3@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta1</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>ta2</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>/</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>numerical</type>
            <tans>ta3</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>4</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>4</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions>4</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>