<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>13.4.1.2f Delvis integrasjon</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    <p class="HELM_exercise">
        Exercise
    </p>
    <p>
        In some questions below it will be necessary to apply integration by parts more than once.
    </p>
    <p>
        Find
    </p>
    <ol class="HELM_parts">
        <li>
            \({\displaystyle\int t e^{-st}\mathrm{d}t}\) where \(s\) is a constant,[[input:ans11]] [[validation:ans11]] [[feedback:prt11]]
        </li>
        <li>
            \({\displaystyle\int t^2 e^{-st}\mathrm{d}t}\) where \(s\) is a constant.[[input:ans12]] [[validation:ans12]] [[feedback:prt12]]
        </li>
    </ol>
</span>
<span class="multilang" lang="no">
    <p class="HELM_exercise">
        Oppgave
    </p>
    <p>
        Regn ut
    </p>
    <ol class="HELM_parts">
        <li>
            \({\displaystyle\int t e^{-st}\mathrm{d}t}\) hvor \(s\) er en konsant.[[input:ans11]] [[validation:ans11]] [[feedback:prt11]]
        </li>
        <li>
            \({\displaystyle\int t^2 e^{-st}\mathrm{d}t}\) hvor \(s\) er en konsant.[[input:ans12]] [[validation:ans12]] [[feedback:prt12]]
        </li>
    </ol>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    <p>
    </p>
    <p>
        We apply integraton by parts: \[ \int u{{\rm d}v\over {\rm d}x}{\rm d}x=uv- \int{{\rm d}u\over {\rm d}x}v\,{\rm d}x\]
    </p>
    <p>
        or alternatively: \[\int f(x)g(x)\,{\rm d}x=f(x)\,\int g(x){\rm d}x -\int {{\rm d}f\over {\rm d}x}\left\{\int g(x){\rm d}x\right\}{\rm d}x.\]
    </p>
    <ol class="HELM_parts">
        <li>
            <p>
                We are given the integral \({\displaystyle\int {@exp11@} \mathrm{d}t}.\)
            </p>
            <p>
                We have \( f(x) = {@f11@}\) and \(g(x) = {@g11@}.\)
            </p>
            <p>
                Hence:
            </p>
            <p>
                \(\begin{align} \int {@exp11@} \mathrm{d}t &amp;={@f11@} \int {@g11@} \mathrm{d}t-\int \left({@df11@} \cdot \int {@g11@}\mathrm{d}t \right)\mathrm{d}t \\ &amp;= {@f11* intg11@}- \int {@df11*intg11@}\mathrm{d}t \\ &amp;= {@ta11@}. \end{align}\)
            </p>
        </li>
        <li>
            <p>
                We are given the integral \({\displaystyle\int {@exp12@} \mathrm{d}t}.\)
            </p>
            <p>
                We have \( f(x) = {@f12@}\) and \(g(x) = {@g12@}.\)
            </p>
            <p>
                Hence:
            </p>
            <p>
                \(\begin{align} \int {@exp12@} \mathrm{d}t &amp;={@f12@} \int {@g12@} \mathrm{d}t-\int \left({@df12@} \cdot \int {@g12@}\mathrm{d}t \right)\mathrm{d}t \\ &amp;= {@f12* intg12@}- \int {@df12*intg12@}\mathrm{d}t. \end{align}\)
            </p>
            <p>
                We now need to find \(\displaystyle \int {@df12*intg12@} \mathrm{d}t \) using intefration by parts again. We get:
            </p>
            <p>
                \(\displaystyle {\int {@df12*intg12@} \mathrm{d}t= {@int(df12*intg12,t)@}}.\)
            </p>
            <p>
                Altogether we have:
            </p>
            <p>
                \(\displaystyle{\int {@exp12@} \mathrm{d}t= {@ta12@}}.\)
            </p>
        </li>
    </ol>
</span>
<span class="multilang" lang="no">
    <p>
    </p>
    <p>
        Vi bruker delvis integrasjon: \[ \int u{{\rm d}v\over {\rm d}x}{\rm d}x=uv- \int{{\rm d}u\over {\rm d}x}v\,{\rm d}x\]
    </p>
    <p>
        alternativt: \[\int f(x)g(x)\,{\rm d}x=f(x)\,\int g(x){\rm d}x -\int {{\rm d}f\over {\rm d}x}\left\{\int g(x){\rm d}x\right\}{\rm d}x.\]
    </p>
    <ol class="HELM_parts">
        <li>
            <p>
                Gitt integralet \({\displaystyle\int {@exp11@} \mathrm{d}t}.\)
            </p>
            <p>
                Velger \( f(x) = {@f11@}\) og \(g(x) = {@g11@}.\)
            </p>
            <p>
                Dermed:
            </p>
            <p>
                \(\begin{align} \int {@exp11@} \mathrm{d}t &amp;={@f11@} \int {@g11@} \mathrm{d}t-\int \left({@df11@} \cdot \int {@g11@}\mathrm{d}t \right)\mathrm{d}t \\ &amp;= {@f11* intg11@}- \int {@df11*intg11@}\mathrm{d}t \\ &amp;= {@ta11@}. \end{align}\)
            </p>
        </li>
        <li>
            <p>
                Gitt integralet \({\displaystyle\int {@exp12@} \mathrm{d}t}.\)
            </p>
            <p>
                Velger \( f(x) = {@f12@}\) and \(g(x) = {@g12@}.\)
            </p>
            <p>
                Dermed:
            </p>
            <p>
                \(\begin{align} \int {@exp12@} \mathrm{d}t &amp;={@f12@} \int {@g12@} \mathrm{d}t-\int \left({@df12@} \cdot \int {@g12@}\mathrm{d}t \right)\mathrm{d}t \\ &amp;= {@f12* intg12@}- \int {@df12*intg12@}\mathrm{d}t. \end{align}\)
            </p>
            <p>
                N&#229; m&#229; vi finne \(\displaystyle \int {@df12*intg12@} \mathrm{d}t \) med en ny runde delvis integrasjon. Vi f&#229;r at:
            </p>
            <p>
                \(\displaystyle {\int {@df12*intg12@} \mathrm{d}t= {@int(df12*intg12,t)@}}.\)
            </p>
            <p>
                Tilsammen f&#229;r vi:
            </p>
            <p>
                \(\displaystyle{\int {@exp12@} \mathrm{d}t= {@ta12@}}.\)
            </p>
        </li>
    </ol>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>exp11: t*e^(-s*t); 
f11: t;                         /* needed for worked solution */
g11: e^(-s*t);
intg11: int(g11,t);
df11: diff(f11,t);

exp12: t^2*e^(-s*t);
f12: t^2;                         /* needed for worked solution */
g12: e^(-s*t);
intg12: int(g12,t);
df12: diff(f12,t);

ta11: int(exp11, t) + c; 
ta12: int(exp12, t) + c;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text/>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans11</name>
            <type>algebraic</type>
            <tans>ta11</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans12</name>
            <type>algebraic</type>
            <tans>ta12</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt11</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans11</sans>
                <tans>ta11</tans>
                <testoptions>t</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt11-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt11-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt12</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans12</sans>
                <tans>ta12</tans>
                <testoptions>t</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt12-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt12-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>