<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>13.4.1.2a Delvis integrasjon</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In some questions below it will be necessary to apply integration by parts more than once.
    </span>
    <span class="multilang" lang="en">
        I noen av oppgavene under vil du m&#229;tte bruke delvis integrasjon flere ganger.
    </span>
</p>
<br/>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            Find
        </span>
        <span class="multilang" lang="no">
            Regn ut
        </span>
    </p>
    <li>
        \({\displaystyle\int {@exp1@} \mathrm{d}x}= \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        \({\displaystyle\int {@exp2@} \mathrm{d}t} \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
    <li>
        \({\displaystyle\int x\cos x \mathrm{d}x} \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand([2, 4, 6]); 
bb: rand(5)+3;

exp1: x*sin(aa*x);
f1: x;                         /* needed for worked solution */
g1: sin(aa*x);
intg1: int(g1,x);
df1: diff(f1,x);

exp2: t*e^(bb*t);
f2: t;                         /* needed for worked solution */
g2: e^(bb*t);
intg2: int(g2,t);
df2: diff(f2,t);

exp3: x*cos(x); 
f3: x;                         /* needed for worked solution */
g3: cos(x);
intg3: int(g3,x);
df3: diff(f3,x);

ta1: int(exp1, x) + c; 
ta2: int(exp2, t) + c; 
ta3: int(exp3, x) + c; </text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We apply integraton by parts:
    </span>
    <span class="multilang" lang="no">
        Vi anvender delvis integrasjon:
    </span>
    \[
    <br/>
    \int u{{\rm d}v\over {\rm d}x}{\rm d}x=uv- \int{{\rm d}u\over {\rm d}x}v\,{\rm d}x\]
    <br/>
    or alternatively: \[\int f(x)g(x)\,{\rm d}x=f(x)\,\int
    <br/>
    g(x){\rm d}x -\int {{\rm d}f\over {\rm d}x}\left\{\int g(x){\rm d}x\right\}{\rm
                d}x.\]
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We are given the integral \({\displaystyle\int {@exp1@} \mathrm{d}x}.\)
            </span>
            <span class="multilang" lang="no">
                Vi har gitt integralet\({\displaystyle\int {@exp1@} \mathrm{d}x}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                We have \( f(x) = {@f1@}\) and \(g(x) = {@g1@}.\)
            </span>
            <span class="multilang" lang="no">
                Vi har at \( f(x) = {@f1@}\) og \(g(x) = {@g1@}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Hence:
            </span>
            <span class="multilang" lang="no">
                Dermed:
            </span>
        </p>
        <p>
            \(\begin{align} \int {@exp1@} \mathrm{d}x &amp;={@f1@} \int {@g1@} \mathrm{d}x -\int \left({@df1@} \cdot \int {@g1@}\mathrm{d}x \right)\mathrm{d}x \\ &amp;= {@f1* intg1@}- \int {@df1*intg1@}\mathrm{d}x \\ &amp;= {@ta1@}. \end{align}\)
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We are given the integral \({\displaystyle\int {@exp2@} \mathrm{d}x}.\)
            </span>
            <span class="multilang" lang="no">
                Vi har gitt integralet \({\displaystyle\int {@exp2@} \mathrm{d}x}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                We have \( f(x) = {@f2@}\) and \(g(x) = {@g2@}.\)
            </span>
            <span class="multilang" lang="no">
                Vi har at \( f(x) = {@f2@}\) and \(g(x) = {@g2@}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Hence:
            </span>
            <span class="multilang" lang="no">
                Dermed:
            </span>
        </p>
        <p>
            \(\begin{align} \int {@exp2@} \mathrm{d}x &amp;={@f2@} \int {@g2@} \mathrm{d}x -\int \left({@df2@} \cdot \int {@g2@}\mathrm{d}x \right)\mathrm{d}x \\ &amp;= {@f2* intg2@}- \int {@df2*intg2@}\mathrm{d}x \\ &amp;= {@ta2@}. \end{align}\)
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We are given the integral \({\displaystyle\int {@exp3@} \mathrm{d}x}.\)
            </span>
            <span class="multilang" lang="no">
                Vi har gitt integralet \({\displaystyle\int {@exp3@} \mathrm{d}x}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                We have \( f(x) = {@f3@}\) and \(g(x) = {@g3@}.\)
            </span>
            <span class="multilang" lang="no">
                Vi har at \( f(x) = {@f3@}\) and \(g(x) = {@g3@}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Hence:
            </span>
            <span class="multilang" lang="no">
                Dermed:
            </span>
        </p>
        <p>
            \(\begin{align} \int {@exp3@} \mathrm{d}x &amp;={@f3@} \int {@g3@} \mathrm{d}x -\int \left({@df3@} \cdot \int {@g3@}\mathrm{d}x \right)\mathrm{d}x \\ &amp;= {@f3* intg3@}- \int {@df3*intg3@}\mathrm{d}x \\ &amp;= {@ta3@}. \end{align}\)
        </p>
    </li>
</ol>
]]></text>
        </specificfeedback>
        <questionnote>
            <text>\[ \int {@exp1@}\mathrm{d}x = {@ta1@}, \quad \int {@exp2@} \mathrm{d}t = {@ta2@}\]
\[ \int x\cos x \mathrm{d}x = {@ta3@} \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>t</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>