<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>13.4.1.2c Delvis integrasjon</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    <p class="HELM_exercise">
        Exercise
    </p>
    <br/>
    <p>
        By writing \(\ln x\) as \(1\cdot \ln x\) find
    </p>
    <p>
        \({\displaystyle\int \ln x\, dx}= \) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]
    </p>
</span>
<span class="multilang" lang="no">
    <p class="HELM_exercise">
        Oppave
    </p>
    <br/>
    <p>
        Ved &#229; skrive \(\ln x\) som \(1\cdot \ln x\), regn ut
    </p>
    <p>
        \({\displaystyle\int \ln x\, dx}= \) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    <p>
    </p>
    <p>
        We apply integraton by parts: \[
        <br/>
        \int u{{\rm d}v\over {\rm d}x}{\rm d}x=uv- \int{{\rm d}u\over {\rm d}x}v\,{\rm d}x\]
        <br/>
        or alternatively: \[\int f(x)g(x)\,{\rm d}x=f(x)\,\int
        <br/>
        g(x){\rm d}x -\int {{\rm d}f\over {\rm d}x}\left\{\int g(x){\rm d}x\right\}{\rm
                d}x.\]
    </p>
    <p>
        We are given the integral \({\displaystyle\int {@exp5@} \mathrm{d}x}.\)
    </p>
    <p>
        We have \( f(x) = {@f5@}\) and \(g(x) = {@g5@}.\)
    </p>
    <p>
        Hence:
    </p>
    <p>
        \(\begin{align} \int {@exp5@} \mathrm{d}x &amp;={@f5@} \int {@g5@} \mathrm{d}x -\int \left({@df5@} \cdot \int {@g5@}\mathrm{d}x \right)\mathrm{d}x \\ &amp;= {@f5* intg5@}- \int {@df5*intg5@}\mathrm{d}x \\ &amp;= {@ta5@}. \end{align}\)
    </p>
</span>
<span class="multilang" lang="no">
    <p>
    </p>
    <p>
        Vi bruker delvis integrasjon: \[
        <br/>
        \int u{{\rm d}v\over {\rm d}x}{\rm d}x=uv- \int{{\rm d}u\over {\rm d}x}v\,{\rm d}x\]
        <br/>
        alternativt: \[\int f(x)g(x)\,{\rm d}x=f(x)\,\int
        <br/>
        g(x){\rm d}x -\int {{\rm d}f\over {\rm d}x}\left\{\int g(x){\rm d}x\right\}{\rm
                d}x.\]
    </p>
    <p>
        Vi har integralet \({\displaystyle\int {@exp5@} \mathrm{d}x}.\)
    </p>
    <p>
        Velger s&#229; \( f(x) = {@f5@}\) and \(g(x) = {@g5@}.\)
    </p>
    <p>
        Dermed blir:
    </p>
    <p>
        \(\begin{align} \int {@exp5@} \mathrm{d}x &amp;={@f5@} \int {@g5@} \mathrm{d}x -\int \left({@df5@} \cdot \int {@g5@}\mathrm{d}x \right)\mathrm{d}x \\ &amp;= {@f5* intg5@}- \int {@df5*intg5@}\mathrm{d}x \\ &amp;= {@ta5@}. \end{align}\)
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>exp5: ln(x);
f5: ln(x);                         /* needed for worked solution */
g5: 1;
intg5: int(g5,x);
df5: diff(f5,x);
ta5: int(exp5, x) + c;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text/>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>ta5</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans5</sans>
                <tans>ta5</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>