<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>13.5.1.2 Substitution</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave
    </span>
</p>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            Use suitable substitutions to find
        </span>
        <span class="multilang" lang="no">
            Bruk subsitusjon for &#229; finne
        </span>
    </p>
    <li>
        \({\displaystyle\int {@exp1@}\mathrm{d}x}=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        \({\displaystyle\int {@exp2@}\mathrm{d}t}=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    <ol class="HELM_parts">
        <li>
            <p>
                We need to calculate the integral \( I =\int {@exp1@} \mathrm{d}x .\)
            </p>
            <p>
                Using \( u = {@u1@}\) we get \(\dfrac{du}{dx} = {@diff(u1,x)@} \Rightarrow \mathrm{d}x = \frac{1}{{@diff(u1,x)@}} \mathrm{d}u.\)
            </p>
            <p>
                \(I = \int {@exp1@}\mathrm{d}x = \int u^{{@bb@}} \mathrm{d}x = \int \dfrac{u^{{@bb@}}}{{@aa@}}\mathrm{d}u = {@ta1@}.\)
            </p>
        </li>
        <li>
            <p>
                We need to calculate the integral \( I =\int {@exp2@} \mathrm{d}t .\)
            </p>
            <p>
                Using \( u = {@v1@}\) we get \(\dfrac{du}{dx} = {@diff(v1,t)@} \Rightarrow \mathrm{d}t = \frac{1}{{@diff(v1,t)@}} \mathrm{d}u.\)
            </p>
            <p>
                \(I = \int {@exp2@}\mathrm{d}t = \int {@p1@} \sin u \, \mathrm{d}t =&#160;&#160;\int {@p1@} \sin u \, \frac{1}{{@diff(v1,t)@}} \mathrm{d}u = \int \dfrac{\sin (u)}{3} ={@ta2@}.\)
            </p>
        </li>
    </ol>
</span>
<span class="multilang" lang="no">
    <ol class="HELM_parts">
        <li>
            <p>
                Vi m&#229; regne ut integralet \( I =\int {@exp1@} \mathrm{d}x .\)
            </p>
            <p>
                Ved &#229; bruke \( u = {@u1@}\) f&#229;r vi\(\dfrac{du}{dx} = {@diff(u1,x)@} \Rightarrow \mathrm{d}x = \frac{1}{{@diff(u1,x)@}} \mathrm{d}u.\)
            </p>
            <p>
                \(I = \int {@exp1@}\mathrm{d}x = \int u^{{@bb@}} \mathrm{d}x = \int \dfrac{u^{{@bb@}}}{{@aa@}}\mathrm{d}u = {@ta1@}.\)
            </p>
        </li>
        <li>
            <p>
                Vi m&#229; regne ut integrlet \( I =\int {@exp2@} \mathrm{d}t .\)
            </p>
            <p>
                Ved &#229; bruke \( u = {@v1@}\) f&#229;r vi \(\dfrac{du}{dx} = {@diff(v1,t)@} \Rightarrow \mathrm{d}t = \frac{1}{{@diff(v1,t)@}} \mathrm{d}u.\)
            </p>
            <p>
                \(I = \int {@exp2@}\mathrm{d}t = \int {@p1@} \sin u \, \mathrm{d}t =&#160;&#160;\int {@p1@} \sin u \, \frac{1}{{@diff(v1,t)@}} \mathrm{d}u = \int \dfrac{\sin (u)}{3} ={@ta2@}.\)
            </p>
        </li>
    </ol>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand([2, 4, 6, 8]); 
bb: rand([6, 7, 8, 9]);

u1: aa*x + 1;
v1: t^3 + 1;
exp1: (aa*x + 1)^bb; 
p1: t^2;
exp2: (t^2)*sin(t^3 + 1);

ta1: int(exp1, x) + c; 
ta2: int(exp2, t) + c;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( \int {@exp1@}\mathrm{d}x= {@ta1@}, \quad  \int {@exp2@}\mathrm{d}t = {@ta2@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>t</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>