<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>13.5.3.2 Integralser og logaritmer</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave
    </span>
</p>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            Write down the result of finding the following integrals..
        </span>
        <span class="multilang" lang="no">
            Skriv ned resultatet av de f&#248;lgende integralene
        </span>
    </p>
    <li>
        \({\displaystyle\int {@exp1@}\mathrm{d}x}=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        \({\displaystyle\int {@exp2@}\mathrm{d}t}=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
    <li>
        \({\displaystyle\int {@exp3@}\mathrm{d}x}=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </li>
    <li>
        \({\displaystyle\int {@exp4@}\mathrm{d}x}=\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    <ol class="HELM_parts">
        <li>
            \({\displaystyle\int {@exp1@}\mathrm{d}x}={@aa@} \int \dfrac{1}{x}\mathrm{d}x = {@ta1@}\)
        </li>
        <li>
            <p>
                Using \( z= t^{{@nn@}} +1\), we observe that \({\displaystyle\int {@exp2@}\mathrm{d}t= \int \frac{z'}{z} \mathrm{d}t}.\)
            </p>
            <p>
                Hence, \({\displaystyle\int {@exp2@}\mathrm{d}t= \ln|z|+c = {@ta2@}}.\)
            </p>
        </li>
        <li>
            <p>
                Using \( z= {@den3@}\), we observe that the derivative of the denominator is {@diff(den3,x)@} whereas the numerator is just \(1.\) We adjust the numerator as follows:
            </p>
            <p>
                \({\displaystyle{ \int {@exp3@}\mathrm{d}x}=\frac{1}{@diff(den3,x)@} \int {@diff(den3,x)*exp3@} = {@ta3@}}.\)
            </p>
        </li>
        <li>
            <p>
                Using \( z= {@den4@}\), we observe that the derivative of the denominator is {@diff(den4,x)@} whereas the numerator is just \({@num4@}.\) We adjust the numerator as follows:
            </p>
            <p>
                \({\displaystyle{ \int {@exp4@}\mathrm{d}x}=\frac{{@num4@}}{@diff(den4,x)@} \int {@diff(den4,x)* (1/den4)@} = {@num4@} {@int(1/den4,x)@}={@ta4@}}.\)
            </p>
        </li>
    </ol>
</span>
<span class="multilang" lang="no">
    <ol class="HELM_parts">
        <li>
            \({\displaystyle\int {@exp1@}\mathrm{d}x}={@aa@} \int \dfrac{1}{x}\mathrm{d}x = {@ta1@}\)
        </li>
        <li>
            <p>
                Ved &#229; bruke \( z= t^{{@nn@}} +1\), ser vi at \({\displaystyle\int {@exp2@}\mathrm{d}t= \int \frac{z'}{z} \mathrm{d}t}.\)
            </p>
            <p>
                Derfor er, \({\displaystyle\int {@exp2@}\mathrm{d}t= \ln|z|+c = {@ta2@}}.\)
            </p>
        </li>
        <li>
            <p>
                Ved &#229; bruke \( z= {@den3@}\), ser vi at nevneren derivert er {@diff(den3,x)@} mens telleren er bare \(1.\) Vi justerer telleren:
            </p>
            <p>
                \({\displaystyle{ \int {@exp3@}\mathrm{d}x}=\frac{1}{@diff(den3,x)@} \int {@diff(den3,x)*exp3@} = {@ta3@}}.\)
            </p>
        </li>
        <li>
            <p>
                Ved &#229; bruke \( z= {@den4@}\),  ser vi at nevneren derivert er {@diff(den4,x)@} mens telleren er bare \({@num4@}.\) vi justerer telleren:
            </p>
            <p>
                \({\displaystyle{ \int {@exp4@}\mathrm{d}x}=\frac{{@num4@}}{@diff(den4,x)@} \int {@diff(den4,x)* (1/den4)@} = {@num4@} {@int(1/den4,x)@}={@ta4@}}.\)
            </p>
        </li>
    </ol>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>/* Integration gives a logarithm. Complicated expresion may not always marked corectlly even if SAns in mathematically correct. Part c,d try to catch various possible answers but may not catch all of them*/ 

aa: rand(7)+2; 
nn: rand([2, 4, 6, 8])
bb: rand(6)+2; 
cc: rand(7)+2;

exp1: aa/x; 
exp2:(nn*t^(nn - 1))/(t^(nn) + 1) ; 
den3: bb*x + cc;
exp3: 1/(bb*x + cc);
num4: 2; 
den4: cc*x - bb;
exp4: 2/(cc*x - bb);

ta1: int(exp1, x) + c; 
ta2: int(exp2, t) + c; 
ta3: int(exp3, x) + c; 
ta3_n:int(1/(factor(den3)), x) + c; 
ta4: int(exp4, x) + c;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[ \int {@exp1@}\mathrm{d}x={@ta1@}, \quad \int {@exp2@}\mathrm{d}t= {@ta2@} \]
\[\int {@exp3@}\mathrm{d}x= {@ta3@}, \quad \int {@exp4@}\mathrm{d}x={@ta4@} \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>t</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions>[x,exp3]</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>diff(ans3,x)</sans>
                <tans>diff(ta3,x)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>length(lv1)</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>diff(ans4,x)</sans>
                <tans>diff(ta4,x)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt4-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>length(lv2)</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>