<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2.362</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
</p>
<p>
    I Ohio har NASA bygd et langt fallr&#248;r. Inne i fallr&#248;ret er det et minilaboratorium som kan falle fritt \(132m\). I dette fallende minilaboratoriet kan NASA utf&#248;re eksperimenter i vektl&#248;s tilstand. Massen av minilaboratoriet er \(225kg\). (F&#248;r eksperimentet begynner, pumpes r&#248;ret tomt for luft for &#229; unng&#229; luftmotstand.) Velg positiv retning oppover gjennom hele oppgaven.
</p>
<p>
    <br/>
</p>
<p>
    <strong>
        a)
    </strong>
    Regn ut hvor lang tid det frie fallet varer.
    <span style="font-size: 0.9375rem;">
        [[input:ans1]] [[validation:ans1]]
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        <br/>
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        <br/>
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        Etter det frie fallet bremses&#160;minilaboratoriet av sm&#229; plastkuler i bunnen av fallr&#248;ret. Oppbremsingen skjer over en strekning p&#229;
        <strong>
            4.5m
        </strong>
        .
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        <br/>
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        <strong>
            b)
        </strong>
        Finn verdien av den gjennomsnittlige akselerasjonen under oppbremsingen.
    </span>
    <span style="font-size: 0.9375rem;">
        [[input:ans2]] [[validation:ans2]]
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        <br/>
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        <strong>
            c)
        </strong>
        Hvor stor er den gjennomsnittlige kraften p&#229;&#160;minilaboratoriet fra plastkulene under oppbremsingen?&#160;[[input:ans3]] [[validation:ans3]]
        <br/>
        <br/>
    </span>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020061000</text>
        </stackversion>
        <questionvariables>
            <text>stack_unit_si_declare(true);
ta1:5.19*s;
ta2:290*m/s^2;
ta3:67000*N;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[feedback:prt1]]
<br/>
[[feedback:prt2]]
<br/>
[[feedback:prt3]]
<br/>
]]></text>
        </specificfeedback>
        <questionnote>
            <text/>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>units</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>1</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options>minsf:2,maxsf:2</options>
        </input>
        <input>
            <name>ans2</name>
            <type>units</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>1</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>units</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>1</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>UnitsAbsolute</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>0.05</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Hint: Her kan vi bruke energibetraktning: \(E_p=E_k \Rightarrow mgh=\frac{1}{2}mv^2\) for &#229; finne farten.
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>UnitsAbsolute</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>2</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>UnitsAbsolute</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions>2</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>