<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>4.122</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
Mari skal ta sertifikat for bil. Utgiftene til de obligatoriske kursene og kj&#248;ringene kommer til &#229; koste \({@b@}\) kr. I tillegg trenger hun kj&#248;retimer som koster \({@a@}\) kr.
<br/>
<br/>
Finn en formel for utgiftene \(y\) i kroner n&#229;r hun bruker \(x\) kj&#248;retimer.
<br/>
\(y=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>b:15000+rand_with_step(100,4000,50)
a:rand([4,5,6,7,8])*100
ta:a*x+b

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\(y={@a@}*x+{@b@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>fx:ans1
s_b:ev(ans1,x=0)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text>[[jsxgraph]] var board = JXG.JSXGraph.initBoard(divid, {boundingbox: [-5, 40000, 40, -2000], axis: true, showCopyright: false}); var f = board.jc.snippet('{#fx#}', true, 'x', true); board.create('functiongraph', [f,0,40]); [[/jsxgraph]]</text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text>[[jsxgraph]] var board = JXG.JSXGraph.initBoard(divid, {boundingbox: [-5, 40000, 40, -2000], axis: true, showCopyright: false}); var f = board.jc.snippet('{#fx#}', true, 'x', true); board.create('functiongraph', [f,0,40]); 
var p = board.create('point', [0,{#b#}], {color:'#003399',name:'A', showInfobox:true, snapToGrid:true});
[[/jsxgraph]]</text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>s_b</sans>
                <tans>b</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Du har riktig konstantledd, men stigningstallet er ikke riktig.
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Konstantleddet sier her hva kostnaden er da Mari har tatt 0 kj&#248;retimer, alts&#229; \(x=0\). Du har funnet \({@s_b@}\), men det skulle v&#230;rt \({@b@}\).
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>2019801395</deployedseed>
        <deployedseed>1765973357</deployedseed>
        <deployedseed>186286825</deployedseed>
        <deployedseed>1297312947</deployedseed>
        <deployedseed>488424536</deployedseed>
        <deployedseed>2028830252</deployedseed>
        <deployedseed>927507610</deployedseed>
        <deployedseed>919069598</deployedseed>
        <deployedseed>1314690122</deployedseed>
        <deployedseed>1089560965</deployedseed>
    </question>
    

</quiz>