<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>7.112 d,e</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
    Nedenfor er grafen til funksjonsuttrykket \(f(x) = {@f(x)@}\) fra forrige oppgave. Tegn tangenten til \(f\) i punktet \((1,f(1))\) ved &#229; plassere \(a\) der tangenten krysser y-aksen og \(b\) i \((1,f(1))\)
    <br/>
</p>
<p>
    [[jsxgraph input-ref-stateStore="stateRef1"]]
        // Create a board like normal.
        var board = JXG.JSXGraph.initBoard(divid, {boundingbox: [-15, 15, 15, -15], axis: true, showCopyright: false});
        // Create a point, its initial position will be the default position if no state is present.
        var a = board.create('point', [4, 3], {snapToGrid:true, name:'a'});
        var b = board.create('point', [4, 4], {snapToGrid:true, name:'b'});

  
  
  
  var l1 = board.create('line', [a, b]);

  
  
  
  var f = board.jc.snippet('{#f(x)#}', true, 'x', true);
        board.create('functiongraph', [f,-10,10]);

  
  
  
  // Bind it to the input and state stored in it.
        stack_jxg.bind_point_dual(stateRef1, a, b);

  
  
  
  // As a side note, you typically do not want the state storing input to be directly visible to the user
        // although it may be handy during development to see what happens in it. You might hide it like this:
        var stateInput = document.getElementById(stateRef1);

  
  
  
  stateInput.style.display = 'none';
[[/jsxgraph]]
    <br/>
</p>
<p>
    Hva er likningen til tangenten?
    <br/>
</p>
<p>
    \(y=\)[[input:ans2]] [[validation:ans2]]
    <br/>
</p>
<p style="display:none">
    [[input:stateStore]] [[validation:stateStore]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>a: rand([-6, -4, 2, 4, 6])
f(x):=x^2 + a *x

tangent: f(1)

ta1:  [[0,-1], [1, tangent]]
ta2: (f(1)+1)*x - 1</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<p>
    [[feedback:prt1]]
</p>
<p>
    [[feedback:prt2]]
    <br/>
</p>
]]></text>
        </specificfeedback>
        <questionnote>
            <text>f(x)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed>seed</variantsselectionseed>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>1</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>stateStore</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>1</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>stateStore</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>