<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>7.201</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
    Funksjonen \(f\) er gitt ved \[f(x) = x^2 + bx + c\]
Bestem \(b\) og \(c\) n&#229;r vi vet at \(f({@pX@}) = {@pY@}\) og at \(f\) har et bunnpunkt for \(x = {@bunnX@}\)
</p>
<p>
    [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>b: rand([-10, -8, -6, 6, 8, 10])
g(x):= x^2 + b*x
bunnX: -b/2
bunnY:  f(bunnX)
c: rand(7) +1 - g(bunnX)
f(x):=x^2 + b*x + c

pX: bunnX + rand([-1, 1])
pY: f(pX)

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\({@f(x)@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>f(x)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>1</insertstars>
            <syntaxhint>x^2 + b*x + c</syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>f(x)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    Riktig! Her er grafen til svaret ditt, sammen med bunnpunktet og  punktet for \(f({@pX@})\)
</p>
[[jsxgraph]]
        // boundingbox:[left, top, right, bottom]
        var board = JXG.JSXGraph.initBoard(divid, {boundingbox: [-11, 20, 11, -2], axis: true, showCopyright: false});
        var f = board.jc.snippet('{#ans1#}', true, 'x', true);
        board.create('functiongraph', [f,-10,10]);
board.create('point', [{#pX#}, {#pY#}], {fixed:true, name:''});
board.create('point', [{#bunnX#}, {#bunnY#}], {fixed:true, name:''});




[[/jsxgraph]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    Feil. Her er grafen til svaret ditt sammen med punktet for \(f({@pX@})\). Som du kan se oppfyller ikke svaret ditt kravet om at bunnpunktet skal v&#230;re i \(x= {@bunnX@}\) og/eller kravet om at  \(f({@pX@} = {@pY@})\)
</p>
[[jsxgraph]]
        // boundingbox:[left, top, right, bottom]
        var board = JXG.JSXGraph.initBoard(divid, {boundingbox: [-20, 20, 20, -20], axis: true, showCopyright: false});
        var f = board.jc.snippet('{#ans1#}', true, 'x', true);
        board.create('functiongraph', [f,-30, 30]);
        board.create('point', [{#pX#}, {#pY#}], {fixed:true, name:''});
[[/jsxgraph]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>