<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Oppgave 1_f</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
</p>
<p>
    Deriver funksjonen:
</p>
<p>
    \(f(x)={@exp@}\)
</p>
<p>
    \(f'(x)=\) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020061000</text>
        </stackversion>
        <questionvariables>
            <text>a1:1+rand(6)
a2:3+rand(3)
a3:1+rand(9)
exp:a1*x/(x^a2+a3);
ta:diff(exp,x);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\(({@exp@})'\)   =   \({@ta@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Diff</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>FacForm</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Riktig!
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Riktig svar, men kunne du faktorisert ytterligere?
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1145420222</deployedseed>
        <deployedseed>140278454</deployedseed>
        <deployedseed>594314008</deployedseed>
        <deployedseed>232147148</deployedseed>
        <deployedseed>1791185275</deployedseed>
        <deployedseed>1736142258</deployedseed>
        <deployedseed>771922320</deployedseed>
        <deployedseed>1794039863</deployedseed>
        <deployedseed>565568389</deployedseed>
        <deployedseed>355872741</deployedseed>
        <deployedseed>175158838</deployedseed>
        <deployedseed>2068067422</deployedseed>
        <deployedseed>1002632584</deployedseed>
        <deployedseed>1888174200</deployedseed>
        <deployedseed>1236588927</deployedseed>
        <deployedseed>1125435139</deployedseed>
        <deployedseed>522696219</deployedseed>
        <deployedseed>1231113438</deployedseed>
        <deployedseed>1850995216</deployedseed>
        <deployedseed>198843228</deployedseed>
    </question>
    

</quiz>