<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Oppgave 1_b</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
</p>
<p>
    Gj&#248;r uttrykket s&#229; enkelt som mulig
</p>
<p>
    \({@exp@}=\) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020061000</text>
        </stackversion>
        <questionvariables>
            <text>r1:1+rand(8);
r2:1+rand(6);
r3:1+rand(3);
simp:false;
exp:(a/r3^r3)^(1/r3)*((r1*a^r3)/a^r2)*(a^(2/r3)/2);
exp1:(a/27)^(1/3)*((6*a^3)/a^2)*(a^(2/3)/2);
ta:simplify(exp);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\({@exp@}={@ta@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Uttrykkene er ikke ekvivalente
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>FacForm</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>a</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Riktig
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Uttrykkene er ekvivalente, men kan du forenkle ytterligere?
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>100073347</deployedseed>
        <deployedseed>39157649</deployedseed>
        <deployedseed>1887882157</deployedseed>
        <deployedseed>1014902267</deployedseed>
        <deployedseed>866981674</deployedseed>
        <deployedseed>87367268</deployedseed>
        <deployedseed>2120051066</deployedseed>
        <deployedseed>1484680651</deployedseed>
        <deployedseed>2140963871</deployedseed>
        <deployedseed>1691176354</deployedseed>
        <deployedseed>1545405358</deployedseed>
        <deployedseed>930524634</deployedseed>
        <deployedseed>108265565</deployedseed>
        <deployedseed>360289575</deployedseed>
        <deployedseed>1827144114</deployedseed>
        <deployedseed>393811663</deployedseed>
        <deployedseed>2025048250</deployedseed>
        <deployedseed>1920690827</deployedseed>
        <deployedseed>1039531807</deployedseed>
        <deployedseed>114632948</deployedseed>
        <deployedseed>1458219842</deployedseed>
        <deployedseed>893602147</deployedseed>
        <deployedseed>2101437364</deployedseed>
        <deployedseed>1465195497</deployedseed>
        <deployedseed>434510102</deployedseed>
        <deployedseed>1960671377</deployedseed>
        <deployedseed>648463857</deployedseed>
        <deployedseed>1058156368</deployedseed>
        <deployedseed>2074252516</deployedseed>
        <deployedseed>424952299</deployedseed>
        <deployedseed>152276457</deployedseed>
        <deployedseed>1182058942</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>