<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Oppgave 1_b</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
</p>
<p>
    L&#248;s likningen ved regning.&#160;Laveste verdi f&#248;rst.
</p>
<p>
    \({@exp@}\)&#160; der&#160; \(x\in [0&#176;, 360&#176;&gt;\)
</p>
<p>
    <br/>
</p>
<p>
    \(x=\) [[input:ans1]] [[validation:ans1]] eller \(x=\)&#160;[[input:ans2]] [[validation:ans2]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020061000</text>
        </stackversion>
        <questionvariables>
            <text>a1:sqrt(3+rand_with_step(2,5,1))
a2:6+rand(6);
a3:sqrt(1+rand_with_prohib(1,5,[3]))

exp:a2*sin(x)-a1=a3;
losn:solve(exp,x);
ta1:ev(x,losn[1]);
ta2:pi-ta1;
ta1:ta1*180/pi;
x1:significantfigures(ta1,10);
x2:180-x1;

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[feedback:prt1]]
<br/>
[[feedback:prt2]]
<br/>
]]></text>
        </specificfeedback>
        <questionnote>
            <text><![CDATA[\({@exp@}=>\)   \({@x1@}\) eller \({@x2@}\)]]></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>x1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>x2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumAbsolute</answertest>
                <sans>ans1</sans>
                <tans>x1</tans>
                <testoptions>2</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumAbsolute</answertest>
                <sans>ans2</sans>
                <tans>x2</tans>
                <testoptions>2</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>459027656</deployedseed>
        <deployedseed>1935959782</deployedseed>
        <deployedseed>1146663720</deployedseed>
        <deployedseed>1014267950</deployedseed>
        <deployedseed>815655660</deployedseed>
        <deployedseed>2065686452</deployedseed>
        <deployedseed>1917518773</deployedseed>
        <deployedseed>152593565</deployedseed>
        <deployedseed>1126793176</deployedseed>
        <deployedseed>965595340</deployedseed>
        <deployedseed>2065063560</deployedseed>
        <deployedseed>2105069548</deployedseed>
        <deployedseed>617013580</deployedseed>
        <deployedseed>1062490076</deployedseed>
        <deployedseed>945465516</deployedseed>
        <deployedseed>332146186</deployedseed>
        <deployedseed>1931298633</deployedseed>
        <deployedseed>206477662</deployedseed>
    </question>
    

</quiz>