<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Oppgave 1_c_1</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
</p>
<p>
    Regn ut integralet
</p>
<p>
    \(\int ({@exp1@})\)dx
</p>
<p>
    [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022060100</text>
        </stackversion>
        <questionvariables>
            <text>a1:1+rand(6);
a2:1+rand(6)
a3:1+rand(4)
exp1: a1*x+a2/x-a3/(x^2);
ta1:int(exp1,x)+c;
</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\[\int{@exp1@} dx={@ta1@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>int</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Int</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    Riktig!
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    Pr&#248;v igjen
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>9724664</deployedseed>
        <deployedseed>2112230639</deployedseed>
        <deployedseed>911521766</deployedseed>
        <deployedseed>1159432126</deployedseed>
        <deployedseed>49825798</deployedseed>
        <deployedseed>1844934118</deployedseed>
        <deployedseed>490587231</deployedseed>
        <deployedseed>463974068</deployedseed>
        <deployedseed>1989972324</deployedseed>
        <deployedseed>531064040</deployedseed>
        <deployedseed>1627740416</deployedseed>
        <deployedseed>864654198</deployedseed>
        <deployedseed>1187184848</deployedseed>
        <deployedseed>209253542</deployedseed>
        <deployedseed>403436602</deployedseed>
        <deployedseed>862832187</deployedseed>
        <deployedseed>1289906793</deployedseed>
        <deployedseed>1130986688</deployedseed>
        <deployedseed>1554492532</deployedseed>
        <deployedseed>1919230254</deployedseed>
        <deployedseed>600242322</deployedseed>
        <deployedseed>999791752</deployedseed>
        <deployedseed>908965523</deployedseed>
        <deployedseed>646359938</deployedseed>
        <deployedseed>1230091471</deployedseed>
        <deployedseed>1175335415</deployedseed>
        <deployedseed>644177306</deployedseed>
        <deployedseed>493847647</deployedseed>
        <deployedseed>1496589617</deployedseed>
        <deployedseed>1599281370</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>