<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Oppgave 1_e</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
</p>
<p>
    Deriver funksjonen:
</p>
<p>
    \(f(x)={@exp@}\)
    <br/>
</p>
<p>
    <br/>
</p>
<p>
    \(\frac{df}{dx}=\) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020061000</text>
        </stackversion>
        <questionvariables>
            <text>a1:1+rand(6);
a2:2+rand(6);
a3:2+rand(6);

exp:(a1*x^a2)*cos(a3*x);
losn:diff(exp,x);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text><![CDATA[\({@exp@}=>\)   \({@losn@}\)]]></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>losn</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Diff</answertest>
                <sans>ans1</sans>
                <tans>losn</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Hint: Bruk produktregelen
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1779851263</deployedseed>
        <deployedseed>1922921016</deployedseed>
        <deployedseed>1716335976</deployedseed>
        <deployedseed>459093881</deployedseed>
        <deployedseed>989535978</deployedseed>
        <deployedseed>1179959311</deployedseed>
        <deployedseed>799583723</deployedseed>
        <deployedseed>1128727104</deployedseed>
        <deployedseed>31831139</deployedseed>
        <deployedseed>1509665009</deployedseed>
        <deployedseed>113628104</deployedseed>
        <deployedseed>1978672271</deployedseed>
        <deployedseed>691532880</deployedseed>
        <deployedseed>1568624224</deployedseed>
        <deployedseed>1297672763</deployedseed>
        <deployedseed>496793256</deployedseed>
        <deployedseed>2115727703</deployedseed>
        <deployedseed>747460132</deployedseed>
        <deployedseed>92572862</deployedseed>
        <deployedseed>289561355</deployedseed>
        <deployedseed>766470409</deployedseed>
        <deployedseed>318635606</deployedseed>
        <deployedseed>1863267197</deployedseed>
        <deployedseed>1510840096</deployedseed>
        <deployedseed>1930423320</deployedseed>
        <deployedseed>152168462</deployedseed>
        <deployedseed>1367285305</deployedseed>
        <deployedseed>1311650787</deployedseed>
        <deployedseed>1280349709</deployedseed>
    </question>
    

</quiz>