<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Oppgave 1_a</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p>
</p>
<p>
    Gj&#248;r uttrykket s&#229; enkelt som mulig
</p>
<p>
    \({@teller@}:{@nevner@}=\)
    <span style="font-size: 0.9375rem;">
        [[input:ans1]] [[validation:ans1]]
    </span>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020081700</text>
        </stackversion>
        <questionvariables>
            <text>a1:rand_with_step(2,10,2);
a2:a1*4;
a3:rand_with_step(2,10,2);
a4:a3*2;
a5:rand_with_step(2,10,2);
a6:4*a5;
simp:false;
teller:((a1*x^2-a2)/(a1*x+a2));
nevner:((a3*x+a4)/(a5*x+a6));
exp:teller/nevner;
ta:simplify(exp);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\({@exp@}={@ta@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>exp</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Uttrykkene er ikke ekvivalente
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>FacForm</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>x</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Riktig
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    Utrykkene er ekvivalente, men kan du forenkle ytterligere?
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>944881384</deployedseed>
        <deployedseed>337745953</deployedseed>
        <deployedseed>1420578876</deployedseed>
        <deployedseed>882459832</deployedseed>
        <deployedseed>1374118570</deployedseed>
        <deployedseed>1511259471</deployedseed>
        <deployedseed>997510855</deployedseed>
        <deployedseed>1158294386</deployedseed>
        <deployedseed>634238426</deployedseed>
        <deployedseed>1005857887</deployedseed>
        <deployedseed>1453226388</deployedseed>
        <deployedseed>2066242861</deployedseed>
        <deployedseed>1686454850</deployedseed>
        <deployedseed>129446924</deployedseed>
        <deployedseed>39511040</deployedseed>
        <deployedseed>1926732322</deployedseed>
        <deployedseed>430128195</deployedseed>
        <deployedseed>311856451</deployedseed>
        <deployedseed>2007709701</deployedseed>
        <deployedseed>1121977579</deployedseed>
        <deployedseed>956759673</deployedseed>
        <deployedseed>2128115596</deployedseed>
        <deployedseed>1663014159</deployedseed>
        <deployedseed>1921184896</deployedseed>
        <deployedseed>15490851</deployedseed>
        <deployedseed>1586479952</deployedseed>
        <deployedseed>2065261101</deployedseed>
        <deployedseed>1325248040</deployedseed>
        <deployedseed>949339201</deployedseed>
        <deployedseed>109374370</deployedseed>
        <deployedseed>659619069</deployedseed>
        <deployedseed>240388201</deployedseed>
        <deployedseed>1100602086</deployedseed>
        <deployedseed>1247317869</deployedseed>
        <deployedseed>2088896618</deployedseed>
        <deployedseed>489264897</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
        </qtest>
        <tags>
            <tag>
                <text>Matematikk</text>
            </tag>
        </tags>
    </question>
    

</quiz>