<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>6.4.2.6 Logs and Temperature</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The temperature \(T\), in degrees C, of a chemical reaction is given by the formula
                \[\displaystyle{T={@aa@}e^{{@b2@}t}}\]
        where \(t\geq 0\) is the time, in seconds.
    </span>
    <span class="multilang" lang="no">
        Temperaturen \(T\), i grader C, in en kjemisk reaksjon er gitt av formelen
                    \[\displaystyle{T={@aa@}e^{{@b2@}t}}\]
                    der \(t\geq 0\) er tiden, i sekunder.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Calculate the time taken for the temperature to reach \({@cc@}^\circ \text{C}.\) Please give your answer to 3 significant figures.
    </span>
    <span class="multilang" lang="no">
        Beregn tiden det tar f&#248;r temperaturen n&#229;r \({@cc@}^\circ \text{C}.\) Vennligst gi svaret med 3 gjeldenede gjeldende sifre.
    </span>
</p>
<p>
    \( t = \) [[input:ans1]] [[validation:ans1]] \( \textrm{s}\)
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    We solve
\[{@cc@}={@aa@}e^{{@b1@}t}.\]
Rearranging gives
\[e^{{@b1@}t} = \frac{{@cc@}}{{@aa@}}\]
and taking logarithms
\[t = \frac{1}{{@b1@}} \ln\left(\frac{{@cc@}}{{@aa@}}\right) \approx {@fl@} .\]
        So the time is {@dispsf(ta,3)@} seconds (to 3 s.f.).
</span>
<span class="multilang" lang="no">
    Vi l&#248;ser
\[{@cc@}={@aa@}e^{{@b1@}t}.\]
Omorganisering gir
\[e^{{@b1@}t} = \frac{{@cc@}}{{@aa@}}\]
N&#229;r vi tar logaritmen f&#229;r vi at
\[t = \frac{1}{{@b1@}} \ln\left(\frac{{@cc@}}{{@aa@}}\right) \ca. {@fl@} .\]
S&#229; tiden er {@dispsf(ta,3)@} sekunder (til 3 s.f.).
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2022071300</text>
        </stackversion>
        <questionvariables>
            <text>aa : 70 + rand(20);
b1 : 2 + rand(5);
b2 : b1/100;
cc : 140 + rand(20);
exp : aa*e^(b2*t);
fl: float(log(cc/aa)/b2);
ta : significantfigures(fl,3);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\[ {@exp@} = {@cc@} \Rightarrow t = {@dispsf(ta,3)@} sec\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>dispsf(ta,3)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumSigFigs</answertest>
                <sans>ans1</sans>
                <tans>dispsf(ta,3)</tans>
                <testoptions>3</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>318278591</deployedseed>
        <deployedseed>359978827</deployedseed>
        <deployedseed>2104700052</deployedseed>
        <deployedseed>7324307</deployedseed>
        <deployedseed>50548011</deployedseed>
        <deployedseed>1826302372</deployedseed>
        <deployedseed>192397823</deployedseed>
        <deployedseed>381344467</deployedseed>
        <deployedseed>1758554717</deployedseed>
        <deployedseed>922416136</deployedseed>
        <deployedseed>153242994</deployedseed>
        <deployedseed>1138128281</deployedseed>
        <deployedseed>1206952831</deployedseed>
        <deployedseed>746705668</deployedseed>
        <deployedseed>2112379713</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>dispsf(ta,3)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>dispsf(ev(ta+1,simp),3)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>