<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>10.1.1.7 Equal complex numbers</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If the two complex numbers \( z_1 = {@z1_r@}+{@z1_i*%i@} \) and \( z_2 = {@z2_r@}-{@z2_i*%i@}\) are equal then find the values of the real numbers \(k\) and \(l\).
    </span>
    <span class="multilang" lang="no">
        Hvis de to komplekse tallene \( z_1 = {@z1_r@}+{@z1_i*%i@} \) og \( z_2 = {@z2_r@}-{@z2_i*%i@}\) er like, finn verdiene til de reelle tallene \(k\) og \(l\).
    </span>
</p>
<p>
    \( k = \) [[input:ans1]] [[validation:ans1]]&#160;[[feedback:prt1]]
</p>
<p>
    \( l = \) [[input:ans2]] [[validation:ans2]]&#160;[[feedback:prt2]]
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Give your answers as rational numbers.
    </span>
    <span class="multilang" lang="no">
        Gi svarene dine som rasjonelle tall.
    </span>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Since \( z_1 = z_2 \) we can equate their real and imaginary parts. So: \[ {@a@}k+{@b@} = k+ {@d@} \quad&#160; \text{and} \quad {@c@}l-{@d@} =&#160;-(k-l) \]
    </span>
    <span class="multilang" lang="no">
        Siden \( z_1 = z_2 \) kan vi sette deres reelle og imagin&#230;re deler lik hverandre. S&#229;: \[ {@a@}k+{@b@} = k+ {@d@} \quad&#160; \text{og} \quad {@c@}l-{@d@} =&#160;-(k-l) \]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        By solving these equations we find \( k = \displaystyle {@ta1@}\) and \(l = \displaystyle {@ta2@}\).
    </span>
    <span class="multilang" lang="no">
        Ved &#229; l&#248;se disse ligningene finner vi \( k = \displaystyle {@ta1@}\) og \(l = \displaystyle {@ta2@}\).
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>orderless(i)
a: rand(6)+2;
b: a+2;
c: rand(5)+2;
d: a+1;

z1: (a*k+b)+(c*l-d)*%i;
z1_r: (a*k+b);     /* Maxima doesn't display real part first and then imaginary so I display seperately */
z1_i: (c*l-d);

z2: (k+d)-(k-l)*%i;
z2_r: (k+d);
z2_i: (k-l);

k: (d-b)/(a-1);
l: (d-k)/(c-1)
ta1: k
ta2: l</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( z_1 = {@z1_r@}+{@z1_i*%i@} \), \( z_2 = {@z2_r@}-{@z2_i*%i@} \) so \(k = {@ta1@}\), \(l = {@ta2@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1277516676</deployedseed>
        <deployedseed>1815257382</deployedseed>
        <deployedseed>1590268928</deployedseed>
        <deployedseed>1028942880</deployedseed>
        <deployedseed>252066610</deployedseed>
        <deployedseed>1113191377</deployedseed>
        <deployedseed>1320301875</deployedseed>
        <deployedseed>2092538225</deployedseed>
        <deployedseed>1066331828</deployedseed>
        <deployedseed>30274902</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>