<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>10.1.2.7 Multiplication of Complex numbers</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If \( z = {@z@} \) determine \( z^2. \)&#160; Input your answer, in the form p + q*i, in the box below:
    </span>
    <span class="multilang" lang="no">
        Hvis \( z = {@z@} \) finn \( z^2. \)&#160; Skriv inn svaret ditt, i formen p + q*i, i boksen nedenfor:
    </span>
</p>
<p>
    \(z^2 =\) [[input:ans1]] [[validation:ans1]]&#160;[[feedback:Result]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given \(z = {@z@}\).
    </span>
    <span class="multilang" lang="no">
        Vi f&#229;r \(z = {@z@}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore \( z^2 =({@z@})^2 = ({@z@})({@z@}) = {@ta@} \)
    </span>
    <span class="multilang" lang="no">
        Derfor \( z^2 =({@z@})^2 = ({@z@})({@z@}) = {@ta@} \)
    </span>
    <br/>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>n1:rand(8)+1;
n2:rand(6)+2;
z: n1+%i*n2;
ta:ev(expand((n1+%i*n2)^2),simp);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\(z = {@z@}\),    so \(z^2 = {@ta@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>1</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>Result</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>sansr:ev(realpart(ans1),simp);
sansi:ev(imagpart(ans1),simp);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>CMP-MUL-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>CMP-MUL-FALSE-0</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>sansr</sans>
                <tans>pval</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>REAL_PART_CORRECT</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The real part of your answer is correct.
    </span>
    <span class="multilang" lang="no">
        Den reelle delen av svaret ditt er riktig.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>CMP-QUOT-FALSE-2</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>sansi</sans>
                <tans>qval</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>IM_PART_CORRECT</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The imaginary part of your answer is correct.
    </span>
    <span class="multilang" lang="no">
        Den imagin&#230;re delen av svaret ditt er riktig.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>CMP-QUOT-FALSE-3</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>wrong1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>NOT_SIMPLIFIED_ENOUGH</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is equivalent to the correct answer, but not in the required form \(p + q\,i\).
    </span>
    <span class="multilang" lang="no">
        Svaret ditt er ekvivalent til det riktige svaret, men ikke i den riktige formen \(p + q\,i\).
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>CMP-QUOT-FALSE-1</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1181331979</deployedseed>
        <deployedseed>1061939536</deployedseed>
        <deployedseed>943677024</deployedseed>
        <deployedseed>390567766</deployedseed>
        <deployedseed>1751471833</deployedseed>
        <deployedseed>1748398506</deployedseed>
        <deployedseed>1377512507</deployedseed>
        <deployedseed>2070986315</deployedseed>
        <deployedseed>357697941</deployedseed>
        <deployedseed>2128708882</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>CMP-MUL-TRUE</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta+1,simp)</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>CMP-QUOT-FALSE-3</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>0</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>CMP-QUOT-FALSE-3</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>ans1</name>
                <value>z^2</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>CMP-QUOT-FALSE-3</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>