<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>10.2.3.5 Modulus and argument</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            If \( \displaystyle{z_1 = {@z1@}}\) and&#160;\( \displaystyle{z_2 = {@z2@}}\) find:
        </span>
        <span class="multilang" lang="no">
            Hvis \( \displaystyle{z_1 = {@z1@}}\) and&#160;\( \displaystyle{z_2 = {@z2@}}\) finn:
        </span>
    </p>
    <li>
        <p>
            \( |z_1 z_2| = \) [[input:ans1]] [[validation:ans1]]&#160;[[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            \( \text{Arg}(z_1 z_2) = \)&#160;[[input:ans2]] [[validation:ans2]]&#160;[[feedback:prt2]]
        </p>
    </li>
    <p>
        <span class="multilang" lang="en">
            Give your answers to \(3\) d.p.
        </span>
        <span class="multilang" lang="no">
            Gi svarene dine til \(3\) d.p.
        </span>
    </p>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            \( z_1 = {@z1@} \) so \( |z_1| = {@r1@} \) and \( \text{arg}(z_1) = {@theta1@}\)
        </span>
        <span class="multilang" lang="no">
            \( z_1 = {@z1@} \) s&#229; \( |z_1| = {@r1@} \) og \( \text{arg}(z_1) = {@theta1@}\)
        </span>
    </p>
    <p>
        <span class="multilang" lang="en">
            \( z_2 = {@z2@} \) so \( |z_2| = {@r2@} \) and \( \text{arg}(z_2) = {@theta2@}\)
        </span>
        <span class="multilang" lang="no">
            \( z_2 = {@z2@} \) s&#229; \( |z_2| = {@r2@} \) og \( \text{arg}(z_2) = {@theta2@}\)
        </span>
        <br/>
    </p>
    <li>
        <p>
            <b>
            </b>
            \( |z_1 z_2| = |z_1||z_2| = {@ta1@} \)
            <br/>
        </p>
    </li>
    <li>
        <p>
            <b>
            </b>
            \( \text{arg}(z_1 z_2) = \text{arg}(z_1) + \text{arg}(z_2) = {@ta2@} \)
        </p>
        <p>
            <span class="multilang" lang="en">
                Have in mind that \( -\pi &lt; \text{arg}(z) \leq \pi \) so you may need to add or subtract a \( \pi\) from your \( \text{arg} \) results.
            </span>
            <span class="multilang" lang="no">
                Husk at \( -\pi &lt; \text{arg}(z) \leq \pi \) s&#229; du m&#229; kanskje legge til eller trekke fra \( \pi\) fra dine \( \text{arg} \) resultater.
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: rand_with_prohib(-15,15,[0]);
b1: rand_with_prohib(-12,12,[0]);
a2: rand_with_prohib(-10,10,[0]);
b2: rand_with_prohib(-8,8,[0]);
z1: a1+(b1)*%i;
z2: a2+(b2)*%i;
r1: sqrt((a1)^2+(b1)^2);    /*mod z1*/
r2: sqrt((a2)^2+(b2)^2);    /*mod z2*/
r: decimalplaces(float(r1*r2),3);              /*mod z1*z2 */
theta1: decimalplaces(float(atan2(b1, a1)),4);
theta2: decimalplaces(float(atan2(b2, a2)),4);

x: a1*a2-b1*b2             /*Real part of z */
y: a1*b2+a2*b1            /*Imaginary part of z*/
theta: decimalplaces(float(atan2(y,x)),3);  /* Arg of z */

ta1: dispdp(r,3);
ta2: dispdp(theta,3);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( z_1 = {@z1@}\), \( z_2 = {@z2@}\) so \( |z_1 z_2| = {@ta1@}\) and \( \text{arg}(z_1) = {@theta1@}\), \( \text{arg}(z_2) = {@theta2@}\)  so \( \text{arg}(z_1 z_2) = {@ta2@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>units</type>
            <tans>ta1</tans>
            <boxsize>6</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>units</type>
            <tans>ta2</tans>
            <boxsize>6</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans1</sans>
                <tans>r</tans>
                <testoptions>3</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>r</tans>
                <testoptions>0.005</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans2</sans>
                <tans>theta</tans>
                <testoptions>3</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>NumAbsolute</answertest>
                <sans>ans2</sans>
                <tans>theta</tans>
                <testoptions>0.005</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>248085827</deployedseed>
        <deployedseed>1319085549</deployedseed>
        <deployedseed>1195234449</deployedseed>
        <deployedseed>797627641</deployedseed>
        <deployedseed>274530023</deployedseed>
        <deployedseed>12073410</deployedseed>
        <deployedseed>1106947110</deployedseed>
        <deployedseed>1933979670</deployedseed>
        <deployedseed>1464661752</deployedseed>
        <deployedseed>585504294</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(r+0.01,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(theta+0.007,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>