<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>8.3.2.1 General Gauss Elimination + Final exercise on back substitution</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<h5>
    <span class="multilang" lang="en">
        The general system of three simultaneous linear equations
    </span>
    <span class="multilang" lang="no">
        Det generelle systemet av tre simultane line&#230;re likninger
    </span>
</h5>
<p>
    <span class="multilang" lang="en">
        In the previous section we met systems of equations which could be solved by back-substitution alone. In this Section we meet systems which are not so amenable and where preliminary work must be done before back-substitution can be used.
    </span>
    <span class="multilang" lang="no">
        I forrige avsnitt m&#248;tte vi ligningssystemer som kunne l&#248;ses ved hjelp av kun bakover-substitusjon. I dette avsnittet m&#248;ter vi systemer som ikke er s&#229; lette &#229; h&#229;ndtere, og hvor det m&#229; gj&#248;res forarbeid f&#248;r bakover-substitusjon kan benyttes.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Consider the system \[\begin{matrix}
                                x_1 &amp; + &amp; 3x_2 &amp; + &amp; 5x_3 &amp; = &amp; 14 \\
                                2x_1 &amp; - &amp; x_2 &amp; - &amp; 3x_3 &amp; = &amp; 3 \\
                                4x_1 &amp; + &amp; 5x_2 &amp; - &amp; x_3 &amp; = &amp; 7 \\
                                \end{matrix}\] We will use the solution method known as Gauss elimination, which has three stages. In the first stage the equations are written in matrix form. In the second stage the matrix equations are replaced by a system of equations having the same solution but which are in triangular form. In the final stage the new system is solved by back-substitution.
    </span>
    <span class="multilang" lang="no">
        Vi tar for oss systemet \[\begin{matrix}
                                x_1 &amp; + &amp; 3x_2 &amp; + &amp; 5x_3 &amp; = &amp; 14 \\
                                2x_1 &amp; - &amp; x_2 &amp; - &amp; 3x_3 &amp; = &amp; 3 \\
                                4x_1 &amp; + &amp; 5x_2 &amp; - &amp; x_3 &amp; = &amp; 7 \\
                                \end{matrix}\] Vi vil bruke l&#248;sningsmetoden kjent som Gauss eliminering, som har tre stadier. I det f&#248;rste stadiet blir ligningene skrevet i matriseform. I det andre stadiet blir matriseligningene erstattet av et system av ligninger som har samme l&#248;sning, men som er i triangul&#230;r form. I det siste stadiet blir det nye systemet l&#248;st ved bakover-substitusjon.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Stage 1: Matrix Formulation
    </span>
    <span class="multilang" lang="no">
        Trinn 1: Matriseframstilling
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The first step is to write the equations in matrix form:
    </span>
    <span class="multilang" lang="no">
        Det f&#248;rste steget er &#229; skrive likningene i matriseform:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \[\left[\begin{matrix}
                                1&amp;3&amp;5\\
                                2&amp;-1&amp;-3\\
                                4&amp;5&amp;-1 \\
                                \end{matrix}\right]
                                \left[\begin{matrix}
                                x_1\\
                                x_2\\
                                x_3 \\
                                \end{matrix}\right] =
                                \left[\begin{matrix}
                                14\\
                                3\\
                                7\\
                                \end{matrix}\right].\] Then, for conciseness, we combine the matrix of coefficients with the column vector of right-hand sides to produce the augmented matrix:
    </span>
    <span class="multilang" lang="no">
        \[\left[\begin{matrix}
                                1&amp;3&amp;5&amp;14\\
                                2&amp;-1&amp;-3&amp;3\\
                                4&amp;5&amp;-1&amp;7 \\
                                \end{matrix}\right].\] Videre, for korthets skyld, kombinerer vi matrisen av koeffisientene med kolonnevektoren p&#229; h&#248;yre side for &#229; lage den utvidete matrisen.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \[\left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                2 &amp; -1 &amp; -3 &amp; 3 \\
                                4 &amp; 5 &amp; -1 &amp; 7 \\
                                \end{array}\right].\]
    </span>
    <span class="multilang" lang="no">
        \[\left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                2 &amp; -1 &amp; -3 &amp; 3 \\
                                4 &amp; 5 &amp; -1 &amp; 7 \\
                                \end{array}\right].\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If the general system of equations is written \(A X =B\) then the augmented matrix is written \([A|B]\). Hence the first equation \[x_1+3x_2+5x_3=14\] is replaced by the first row of the augmented matrix, \[1 \qquad 3\qquad 5\qquad |\qquad 14\] and so on. Stage 1 has now been completed. We will next triangularise the matrix of coefficients by means of row operations. There are three possible row operations:
    </span>
    <span class="multilang" lang="no">
        Hvis det generelle ligningssystemet er skrevet \(A X =B\), da er den utvidede matrisen skrevet som \([A|B]\). Dermed blir den f&#248;rste ligningen \[x_1+3x_2+5x_3=14\] erstattet av den f&#248;rste raden i den utvidede matrisen, \[1 \qquad 3\qquad 5\qquad |\qquad 14\] og s&#229; videre. F&#248;rste steg er n&#229; fullf&#248;rt. Vi vil neste triangulere koeffisientmatrisen ved hjelp av radoperasjoner. Det er tre mulige radoperasjoner:
    </span>
</p>
<ul>
    <li>
        <p>
            <span class="multilang" lang="en">
                interchange two rows;
            </span>
            <span class="multilang" lang="no">
                Bytt om to rader;
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                multiply or divide a row by a non-zero constant factor;
            </span>
            <span class="multilang" lang="no">
                multipliser eller divider en rad med en ikke-null konstant faktor;
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                add to, or subtract from, one row a multiple of another row.
            </span>
            <span class="multilang" lang="no">
                legg til eller trekk fra en multiplum av en annen rad til en rad.
            </span>
        </p>
    </li>
</ul>
<p>
    <span class="multilang" lang="en">
        Note that interchanging two rows of the augmented matrix is equivalent to interchanging the two corresponding equations. The shorthand notation we use is introduced by example. To interchange row 1 and row 3 we write \(R1\leftrightarrow R3\). To divide row 2 by 5 we write \(R2\div 5\). To add three times row 1 to row 2, we write \(R2+3R1\). In the Task which follows you will see where these annotations are placed.
    </span>
    <span class="multilang" lang="no">
        Merk at &#229; bytte om to rader i den utvidede matrisen er ekvivalent med &#229; bytte om de to korresponderende likningene. Den forkortede notasjonen vi bruker blir introdusert med et eksempel. For &#229; bytte om rad 1 og rad 3 skriver vi \(R1\leftrightarrow R3\). For &#229; dele rad 2 med 5 skriver vi \(R2\div 5\). For &#229; legge 'tre ganger rad 1' til rad 2 skriver vi \(R2+3R1\). I oppgaven som f&#248;lger vil du se hvor disse merknadene plasseres.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Note that these operations neither create nor destroy solutions so that at every step the system of equations has the same solution as the original system.
    </span>
    <span class="multilang" lang="no">
        Merk at disse operasjonene verken skaper eller &#248;delegger l&#248;sninger slik at systemet av ligninger p&#229; hvert trinn har samme l&#248;sning som det opprinnelige systemet.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Stage 2: Triangularisation
    </span>
    <span class="multilang" lang="no">
        Trinn 2: Triangulering
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The second stage proceeds by first eliminating \(x_1\) from the second and third equations using row operations.
    </span>
    <span class="multilang" lang="no">
        Den andre fasen fortsetter ved &#229; f&#248;rst eliminere \(x_1\) fra den andre og tredje ligningen ved &#229; bruke radoperasjoner.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \[\left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                2 &amp; -1 &amp; -3 &amp; 3 \\
                                4 &amp; 5 &amp; -1 &amp; 7
                                \end{array}\right]
                                \begin{array}{c}
                                \\ R2-2 \times R1 \\
                                R3-4 \times R1 \\
                                \end{array}
                                \Rightarrow
                                \left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                0 &amp; -7 &amp; -13 &amp; -25 \\
                                0 &amp; -7 &amp; -21 &amp; -49 \\
                                \end{array}\right]\]
    </span>
    <span class="multilang" lang="no">
        \[\left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                2 &amp; -1 &amp; -3 &amp; 3 \\
                                4 &amp; 5 &amp; -1 &amp; 7
                                \end{array}\right]
                                \begin{array}{c}
                                \\ R2-2 \times R1 \\
                                R3-4 \times R1 \\
                                \end{array}
                                \Rightarrow
                                \left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                0 &amp; -7 &amp; -13 &amp; -25 \\
                                0 &amp; -7 &amp; -21 &amp; -49 \\
                                \end{array}\right]\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In the above we have subtracted twice row (equation) 1 from row (equation) 2.
    </span>
    <span class="multilang" lang="no">
        I den ovennevnte matrisen har vi trukket fra 'to ganger rad (likning) 1' fra rad (likning) 2.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In full these operations would be written, respectively, as \[(2x_1-x_2-3x_3)-2(x_1+3x_2+5x_3)=3-2\times 14 \qquad \text{or} \qquad -7x_2-13x_3=-25\] and \[(4x_1+5x_2-x_3)-4(x_1+3x_2+5x_3)=7-4\times 14 \qquad \text{or} \qquad -7x_2-21x_3=-49.\]
    </span>
    <span class="multilang" lang="no">
        Skrevet ut fullt ville disse operasjonene v&#230;re, henholdsvis, som f&#248;lger: \[(2x_1-x_2-3x_3)-2(x_1+3x_2+5x_3)=3-2\times 14 \qquad \text{eller} \qquad -7x_2-13x_3=-25\] og \[(4x_1+5x_2-x_3)-4(x_1+3x_2+5x_3)=7-4\times 14 \qquad \text{eller} \qquad -7x_2-21x_3=-49.\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Now since all the elements in rows 2 and 3 are negative we multiply throughout by \(-1\), this is not really necessary but it will clean up the matrix a bit: \[\left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                0 &amp; -7 &amp; -13 &amp; -25 \\
                                0 &amp; -7 &amp; -21 &amp; -49
                                \end{array}\right]
                                \begin{array}{c}
                                \\
                                R2\times (-1) \\
                                R3\times (-1) \\
                                \end{array}
                                \Rightarrow
                                \left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                0 &amp; 7 &amp; 13 &amp; 25 \\
                                0 &amp; 7 &amp; 21 &amp; 49
                                \end{array}\right]\]
    </span>
    <span class="multilang" lang="no">
        Siden alle elementene i rad 2 og 3 er negative, multipliserer vi med \(-1\). Dette er egentlig ikke n&#248;dvendig, men det vil rydde opp matrisen litt: \[\left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                0 &amp; -7 &amp; -13 &amp; -25 \\
                                0 &amp; -7 &amp; -21 &amp; -49
                                \end{array}\right]
                                \begin{array}{c}
                                \\
                                R2\times (-1) \\
                                R3\times (-1) \\
                                \end{array}
                                \Rightarrow
                                \left[\begin{array}{rrr|r}
                                1 &amp; 3 &amp; 5 &amp; 14 \\
                                0 &amp; 7 &amp; 13 &amp; 25 \\
                                0 &amp; 7 &amp; 21 &amp; 49
                                \end{array}\right]\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Finally, we eliminate \(x_2\) from the third equation by subtracting equation 2 from equation 3, i.e., \(R3-R2\): \[\left[\begin{array}{rrr|r}
                                1&amp;3&amp;5&amp;14\\
                                0&amp;7&amp;13&amp;25\\
                                0&amp;7&amp;21&amp;49 \\
                                \end{array}\right]
                                \begin{array}{c}
                                \\
                                \\
                                R3-R2 \\
                                \end{array}
                                \Rightarrow
                                \left[\begin{array}{rrr|r}
                                1&amp;3&amp;5&amp;14\\
                                0&amp;7&amp;13&amp;25\\
                                0&amp;0&amp;8&amp;24 \\
                                \end{array}\right]\]
    </span>
    <span class="multilang" lang="no">
        Til slutt eliminerer vi \(x_2\) fra tredje ligning ved &#229; trekke ligning 2 fra ligning 3, det vil si \(R3-R2\): \[\left[\begin{array}{rrr|r}
                                1&amp;3&amp;5&amp;14\\
                                0&amp;7&amp;13&amp;25\\
                                0&amp;7&amp;21&amp;49 \\
                                \end{array}\right]
                                \begin{array}{c}
                                \\
                                \\
                                R3-R2 \\
                                \end{array}
                                \Rightarrow
                                \left[\begin{array}{rrr|r}
                                1&amp;3&amp;5&amp;14\\
                                0&amp;7&amp;13&amp;25\\
                                0&amp;0&amp;8&amp;24 \\
                                \end{array}\right]\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The system is now in triangular form.
    </span>
    <span class="multilang" lang="no">
        Systemet er n&#229; i triangel form.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Stage 3: Back Substitution
    </span>
    <span class="multilang" lang="no">
        Trinn 3: Bakover-substitusjon
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Here we solve the equations from bottom to top. At each step of the back substitution process we encounter equations which only have a single unknown and so can be easily solved.
    </span>
    <span class="multilang" lang="no">
        Her l&#248;ser vi likningene fra bunnen og opp. Ved hvert trinn i bakover-substitusjonsprosessen st&#248;ter vi p&#229; likninger som bare har &#233;n ukjent, og kan dermed enkelt l&#248;ses.
    </span>
</p>
<hr/>
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;ving
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Complete the solution to the above system by back-substitution.
    </span>
    <span class="multilang" lang="no">
        Fullf&#248;r l&#248;sningen p&#229; systemet ovenfor ved bakover-substitusjon.
    </span>
</p>
<div>
    \(x=\) [[input:ans1]] [[validation:ans1]]
</div>
<div>
    <br/>
</div>
<p>
    <span class="multilang" lang="en">
        Check that these values satisfy the original system of equations.
    </span>
    <span class="multilang" lang="no">
        Sjekk om disse verdiene tilfredsstiller det opprinnelige ligningssystemet.
    </span>
</p>
<p>
    [[feedback:prt1]]
    <br/>
</p>
<hr/>
<p>
    <span class="multilang" lang="en">
        You should always check your solution in the original equations to ensure that the equations balance.
    </span>
    <span class="multilang" lang="no">
        Du b&#248;r alltid sjekke l&#248;sningen din mot de opprinnelige likningene for &#229; forsikre deg om at likningene balanserer.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Again we emphasise that we chose a particular set of procedures in Stage 1. This was chosen mainly to keep the arithmetic simple by delaying the introduction of fractions. Sometimes we are courageous and take fewer, harder steps.
    </span>
    <span class="multilang" lang="no">
        Igjen legger vi vekt p&#229; at vi valgte en bestemt rekkef&#248;lge av prosedyrer i Trinn 1. Dette ble hovedsakelig valgt for &#229; gj&#248;re aritmetikken enkel ved &#229; utsette introduksjonen av br&#248;k. Noen ganger er vi modige og tar f&#230;rre, vanskeligere steg.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        An important point to note is that when in Stage 2 we wrote \(R2-2R1\) against row 2; what we meant is that row 2 is replaced by the combination (row \(2)-2\times\)(row 1).
    </span>
    <span class="multilang" lang="no">
        En viktig ting &#229; merke seg er at n&#229;r vi i Trinn 2 skrev \(R2-2R1\) ved rad 2; hva vi mente er at rad 2 erstattes av kombinasjonen (rad \(2)-2\times\)(rad 1).
    </span>
</p>
<div class="HELM_keypoint">
    <h4>
        <span class="multilang" lang="en">
            Key Point: Elementary Row Operations
        </span>
        <span class="multilang" lang="no">
            N&#248;kkelpunkt: Grunnleggende radoperasjoner
        </span>
    </h4>
    <p>
        <span class="multilang" lang="en">
            The operation \[\text{ row } i-\alpha \times \text{ row }\ j\] means replace row \(i\) by the combination \[\text{ row } i-\alpha \times \text{ row } j.\]
        </span>
        <span class="multilang" lang="no">
            Operasjonen \[\text{ rad } i-\alpha \times \text{ rad }\ j\] betyr &#229; erstatte rad \(i\) med kombinasjonen \[\text{ rad } i-\alpha \times \text{ rad } j.\]
        </span>
    </p>
</div>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        In full the equations are \[\begin{matrix}
                x_1 &amp; + &amp; 3x_2 &amp; + &amp; 5x_3 &amp; = &amp; 14 \\
                &amp; &amp; 7x_2 &amp; + &amp; 13x_3 &amp; = &amp; 25 \\
                &amp; &amp; &amp; &amp; 8x_3 &amp; = &amp; 24 \\
                \end{matrix}\] From the last equation we see that \(x_3=3\). Substituting this value into the second equation gives \[7x_2+39=25 \qquad \text{or} \qquad 7x_2=-14 \qquad \text{so that} \qquad x_2=-2\] Finally, using these values for \(x_2\) and \(x_3\) in equation 1 gives \[x_1-6+15=14.\] Hence \(x_1=5\). The solution
is therefore \([x_1,x_2,x_3]^T=[5,-2,3]^T\). You can check that these values satisfy the original system of equations.
    </span>
    <span class="multilang" lang="no">
        I sin helhet er likningene \[\begin{matrix}
                x_1 &amp; + &amp; 3x_2 &amp; + &amp; 5x_3 &amp; = &amp; 14 \\
                &amp; &amp; 7x_2 &amp; + &amp; 13x_3 &amp; = &amp; 25 \\
                &amp; &amp; &amp; &amp; 8x_3 &amp; = &amp; 24 \\
                \end{matrix}\] Fra den siste likningen ser vi at \(x_3=3\). Substitusjon av denne verdien i den andre likningen gir \[7x_2+39=25 \qquad \text{eller} \qquad 7x_2=-14 \qquad \text{slik at} \qquad x_2=-2\] Til slutt bruker vi disse verdiene for \(x_2\) og \(x_3\) i likning 1 som gir \[x_1-6+15=14.\] Dermed er \(x_1=5\). L&#248;sningen er derfor \([x_1,x_2,x_3]^T=[5,-2,3]^T\). Du kan sjekke at disse verdiene tilfredsstiller det opprinnelige ligningssystemet.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ordergreat(x, y, z);
/* Input the system as in the example*/
a11 : 1;
a12 : 3;
a13 : 5;
a22 : 7;
a23 : 13;
a33 : 8;
matA : matrix([a11, a12, a13], [0, a22, a23], [0, 0, a33]);
vecB : matrix([14], [25], [24]);
vecX : matrix([x], [y], [z]);

/* Construct the solution */
aux_LHS : matA.vecX;
eq1 : aux_LHS[1][1] = vecB[1][1];
eq2 : aux_LHS[2][1] = vecB[2][1];
eq3 : aux_LHS[3][1] = vecB[3][1];
ta_aux : algsys([eq1, eq2, eq3], [x, y, z]);
ta : matrix([rhs(ta_aux[1][1])], [rhs(ta_aux[1][2])], [rhs(ta_aux[1][3])])
</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text/>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Well done!
</span>
<span class="multilang" lang="no">
    Bra jobba!
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Be careful, your answer doesn't satisfy the system \[ {@matA@} {@ans1@} = {@matA . ans1@} \neq {@vecB@} \].
</span>
<span class="multilang" lang="no">
    V&#230;r forsiktig, svaret ditt tilfredsstiller ikke systemet \[ {@matA@} {@ans1@} = {@matA . ans1@} \neq {@vecB@} \].
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(2*ta,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>