<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.2.2.T1</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;velse
    </span>
</p>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            Find \(\dfrac{dy}{dx}\) where
                                                \( y = {@ff@} \)
        </span>
        <span class="multilang" lang="no">
            Finn \(\dfrac{dy}{dx}\) hvor
                                                \( y = {@ff@} \)
        </span>
    </p>
    <li>
        <p>
            <span class="multilang" lang="en">
                First find the derivative of \( {@f1@} \)
            </span>
            <span class="multilang" lang="no">
                F&#248;rst finn den deriverte til \( {@f1@} \)
            </span>
        </p>
        [[input:ans1]]
                                [[validation:ans1]][[feedback:prt1]]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Next find the derivative of \( {@f2@} \)
            </span>
            <span class="multilang" lang="no">
                Finn derivatet til \( {@f2@} \) neste.
            </span>
        </p>
        [[input:ans2]]
                                [[validation:ans2]][[feedback:prt2]]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Finally, combine your results to find the derivative of
                                                                \( {@ff@}\)
            </span>
            <span class="multilang" lang="no">
                Til slutt kombinerer du resultatene for &#229; finne den deriverte av \({@ff@}\).
            </span>
        </p>
    </li>
</ol>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            \( \dfrac{dy}{dx} =\)
        </span>
        <span class="multilang" lang="no">
            \(\dfrac{dy}{dx}\) betyr "differensialkvotienten til y med hensyn p&#229; x".
        </span>
    </p>
    [[input:ans3]]
                                                [[validation:ans3]][[feedback:prt3]]
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We have \( y= {@ff@}.\)
    </span>
    <span class="multilang" lang="no">
        Vi har \( y= {@ff@}.\)
    </span>
</p>
<br/>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                The derivarive of {@f1@} is {@ta1@}.
            </span>
            <span class="multilang" lang="no">
                Derivert av {@f1@} er {@ta1@}.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                The derivative of {@f2@} is {@ta2@}.
            </span>
            <span class="multilang" lang="no">
                Den deriverte av {@f2@} er {@ta2@}.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                So the derivative of \(y\) is \(\dfrac{dy}{dx} = {@ta3@}.\)
            </span>
            <span class="multilang" lang="no">
                S&#229; den deriverte av \(y\) er \(\dfrac{dy}{dx} = {@ta3@}.\)
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: rand(5)+5;
a2: rand_with_prohib(5,9,[a1]);
a3: rand_with_prohib(3,9,[a1,a2]);
f1: a1*x^a2;
f2: a3*e^(a2*x)
ff: f1-f2;

ta1:diff( f1,x);
ta2:diff( f2,x);
ta3:diff( ff,x);

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( y= {@ff@} \rightarrow y'={@ta3@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>12</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1109815843</deployedseed>
        <deployedseed>222692975</deployedseed>
        <deployedseed>211394853</deployedseed>
        <deployedseed>1022446471</deployedseed>
        <deployedseed>1520726962</deployedseed>
        <deployedseed>1633632809</deployedseed>
        <deployedseed>1226487974</deployedseed>
        <deployedseed>421503170</deployedseed>
        <deployedseed>1051421359</deployedseed>
        <deployedseed>1863212791</deployedseed>
        <deployedseed>440312242</deployedseed>
        <deployedseed>1840153085</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>