<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.2.1.T1 differentiate</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave.
    </span>
</p>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            Use the table of derivatives to find \( \dfrac{dy}{dx}\) when
        </span>
        <span class="multilang" lang="no">
            Bruk tabellen til &#229; finne \( \dfrac{dy}{dx}\) n&#229;r
        </span>
    </p>
    <li>
        <p>
            <span class="multilang" lang="en">
                \(y = {@f1@} \)
            </span>
            <span class="multilang" lang="no">
                \(y = {@f1@} \)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Hint: Write \(\sqrt{x}\) as \(x^{1/2}\) , and use the result for differentiating \(x^n\) with \(n = \dfrac{1}{2}\).
            </span>
            <span class="multilang" lang="no">
                Tips: Skriv \(\sqrt{x}\) som \(x^{1/2}\), og bruk resultatet for &#229; derivere \(x^n\) med \(n = \dfrac{1}{2}\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                \( \dfrac{dy}{dx} \) =
            </span>
            <span class="multilang" lang="no">
                \( \dfrac{dy}{dx} \) =
            </span>
        </p>
        [[input:ans1]][[validation:ans1]][[feedback:prt1]]
    </li>
    <li>
        \(y =&#160; \displaystyle {@f2@} \)
        <p>
            <span class="multilang" lang="en">
                Hint: Write \( \dfrac{ {@b@}}{x^3} \) as \( {@b@}x^{-3}\) and use the result for differentiating \(kx^n\) with \(k = {@b@}\) and \(n = -3\).
            </span>
            <span class="multilang" lang="no">
                Hint: Skriv  \( \dfrac{ {@b@}}{x^3} \) som \( {@b@}x^{-3}\) og bruk resultatet for &#229; derivere \(kx^n\) med \(k = {@b@}\) og \(n = -3\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                \( \dfrac{dy}{dx} \) =
            </span>
            <span class="multilang" lang="no">
                \( \dfrac{dy}{dx} \) = \( \dfrac{d y}{d x} \) =
            </span>
        </p>
        [[input:ans2]][[validation:ans2]][[feedback:prt2]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            We have \(y = {@f1@} = x^{1/2}, \) so:
        </p>
        <p>
            \[ \dfrac{dy}{dx} = nx^{n-1} = \dfrac{1}{2}x^{1/2 - 1} = \dfrac{1}{2}x^{-1/2} \]
        </p>
        <p>
            This may be written as \( \dfrac{1}{2\sqrt{x}} \).
        </p>
    </li>
    <li>
        <p>
            We have \(y = \displaystyle {@f2@} = {@b@}x^{-3}, \) so:
        </p>
        <p>
            \[ \dfrac{dy}{dx} = {@b@}(-3)x^{-3-1} = - {@b*3@}x^{-4} \]
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ta1:diff(x^(1/2),x);

b:rand([3,4,5]);
f1: sqrt(x)
f2: b/(x^3);
ta2: diff(f2,x);

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( y = {@f1@}, \, y' = {@ta1@}, \quad y = {@f2@}, \, y' = {@ta2@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>555319474</deployedseed>
        <deployedseed>64752438</deployedseed>
        <deployedseed>1351530464</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta1+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(ta2+2,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>