<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.2.1.T2 differentiate</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Oppgave
    </span>
</p>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            Although the derivative table is written using \(x\) as the independent variable, the table can be used for any variable. Use the table of derivatives to find
        </span>
        <span class="multilang" lang="no">
            Selv om den derivasjonstabellen er skrevet ved &#229; bruke \(x\) som den uavhengige variabelen, kan tabellen brukes for enhver variabel. Bruk derivasjonstabellen til &#229; finne
        </span>
    </p>
    <li>
        <p>
            <span class="multilang" lang="en">
                Find \(\dfrac{dz}{dt}\) given \( z = e^t \)
            </span>
            <span class="multilang" lang="no">
                Finn \(\dfrac{dz}{dt}\) gitt \(z = e^t\)
            </span>
        </p>
        <p>
            ( \dfrac{dz}{dt} \) = [[input:ans1]]
        [[validation:ans1]][[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Find&#160;\( \dfrac{dp}{dt}\) given \( p = {@f2@} \)
            </span>
            <span class="multilang" lang="no">
                Finn&#160;\( \dfrac{dp}{dt}\) gitt at \( p = {@f2@} \)
            </span>
        </p>
        <p>
            \( \dfrac{dp}{dt} \) = [[input:ans2]]
                                                [[validation:ans2]][[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Find&#160;\( \dfrac{dz}{dy}\) given \( z = {@f3@} \)
            </span>
            <span class="multilang" lang="no">
                Finn \( \dfrac{dz}{dy}\) gitt \( z = {@f3@} \)
            </span>
        </p>
        <p>
            \( \dfrac{dz}{dy} \) = [[input:ans3]]
                                                [[validation:ans3]][[feedback:prt3]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                From the Table, if \(y = e^x\) then \( \dfrac{dy}{dx} = e^x \). Hence if

        
        
        
        \(z = e^t\) then \( \dfrac{dz}{dt} = e^t \).
            </span>
            <span class="multilang" lang="no">
                Fra tabellen, hvis \(y = e^x\), s&#229; er \( \dfrac{dy}{dx} = e^x \). Derfor, hvis 




\( z = e^t \), s&#229; er \( \dfrac{dz}{dt} = e^t \).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                From the Table, if \(y = e^{ {@b@}x}\) then \( \dfrac{dy}{dx} = {@b@}e^{ {@b@}x} \). Hence if

        
        
        
        \(p = e^{ {@b@}t}\) then \( \dfrac{dp}{dt} = {@b@}e^{ {@b@}t} \).
            </span>
            <span class="multilang" lang="no">
                Fra tabellen, hvis \(y = e^{ {@b@}x}\), s&#229; er \( \dfrac{dy}{dx} = {@b@}e^{ {@b@}x} \). Derfor, hvis 




\(p = e^{ {@b@}t}\), s&#229; er \( \dfrac{dp}{dt} = {@b@}e^{ {@b@}t} \).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                From the Table, if \(y = e^{- {@c@}x}\) then \( \dfrac{dy}{dx} = -{@c@}e^{- {@c@}x} \). Hence if

        
        
        
        \(z = e^{ -{@c@}y}\) then \( \dfrac{dz}{dy} = - {@c@}e^{ -{@c@}y} \).
            </span>
            <span class="multilang" lang="no">
                Fra tabellen, hvis \(y = e^{-{@c@}x}\), da er \( \dfrac{dy}{dx} = -{@c@}e^{-{@c@}x} \). Derfor, hvis 




\(z = e^{ -{@c@}y}\) s&#229; er \( \dfrac{dz}{dy} = -{@c@}e^{ -{@c@}y} \).
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ta1:diff(e^t,t);

b:rand([6,7,8,9]);
f2: e^(b*t);
ta2:diff(f2,t);

c:rand([3,4,5]);
f3: e^(-c*y);
ta3: diff(f3,y);

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\( z = e^t, z'= {@ta1@}, \quad p = {@f2@}, p'= {@ta2@}, \quad z = {@f3@}, z'= {@ta3@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
</p>
<p>
    <br/>
</p>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>555319474</deployedseed>
        <deployedseed>64752438</deployedseed>
        <deployedseed>1351530464</deployedseed>
        <deployedseed>352442436</deployedseed>
        <deployedseed>777996274</deployedseed>
        <deployedseed>1016710591</deployedseed>
        <deployedseed>678941322</deployedseed>
        <deployedseed>23878499</deployedseed>
        <deployedseed>172101399</deployedseed>
        <deployedseed>2006881662</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>