<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.2.1.2 Differentiation</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;velse
    </span>
</p>
<ol class="HELM_parts">
    <p>
        <span class="multilang" lang="en">
            Find \(\dfrac{dz}{dt} \)&#160; when \(z\) is given by:
        </span>
        <span class="multilang" lang="no">
            Finn \(\frac{dz}{dt}\) n&#229;r \(z\) er gitt ved:
        </span>
    </p>
    <li>
        \(\displaystyle {@f1@}\)
                            [[input:ans1]]
                                                [[validation:ans1]][[feedback:prt1]]
    </li>
    <li>
        \( {@f2@} \)
        
                
                
                
        [[input:ans2]]
                                                [[validation:ans2]][[feedback:prt2]]
    </li>
    <li>
        \({@c1@} t^{-2}\)
        
                
                
                
        [[input:ans3]]
                                                [[validation:ans3]][[feedback:prt3]]
    </li>
    <li>
        \(-\dfrac{{@d1@}}{2} t^{ \frac{ {@d1@} }{2} }\)
        
                
                
                
        [[input:ans4]]
                                                [[validation:ans4]][[feedback:prt4]]
    </li>
    <li>
        \( {@f5@} \)
        
                
                
                
        [[input:ans5]]
                                                [[validation:ans5]][[feedback:prt5]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We can rewrite \( y= \displaystyle{@f1@}\) as \({@a1@} t^{-3}.\)
            </span>
            <span class="multilang" lang="no">
                Vi kan omskrive \( y= \displaystyle{@f1@}\) som \({@a1@} t^{-3}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Then, using the \(kx^n\) rule, we have that the derivative is: \( y' = {@a1@} \cdot (-3) \cdot t^{-4} = \displaystyle {@ta1@}. \)
            </span>
            <span class="multilang" lang="no">
                Deretter, ved &#229; bruke regelen \(kx^n\), f&#229;r vi at den deriverte er: \( y' = {@a1@} \cdot (-3) \cdot t^{-4} = \displaystyle {@ta1@}.\)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \(y = {@f2@}\), we use the \(x^n\) rule to get \( y' = \dfrac{3}{2} t^{\frac{3}{2}-1} = \displaystyle {@ta2@} .\)
            </span>
            <span class="multilang" lang="no">
                For \(y = {@f2@}\), we use the \(x^n\) regelen for &#229; f&#229; \( y' = \dfrac{3}{2} t^{\frac{3}{2}-1} = \displaystyle {@ta2@} .\)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \( y= {@c1@} t^{-2}\), we use the \(kx^n\) rule.
            </span>
            <span class="multilang" lang="no">
                For \( y= {@c1@} t^{-2}\), we use the \(kx^n\) regelen.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                The derivative is \( y' = {@c1@} \cdot (-2) \cdot t^{-2-1} = \displaystyle {@ta3@} .\)
            </span>
            <span class="multilang" lang="no">
                Den deriverte er \( y' = {@c1@} \cdot (-2) \cdot t^{-2-1} = \displaystyle {@ta3@} .\)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \( y= {@f4@} \), we also use the \(kx^n\) rule.
            </span>
            <span class="multilang" lang="no">
                For \( y= {@f4@} \), vi bruker ogs&#229; regelen \(kx^n\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                The derivative is \( y' = -\frac{{@d1@}}{2} \cdot \frac{{@d1@}}{2} \cdot t^{\frac{{@d1@}}{2} - 1} = \displaystyle{@ta4@} .\)
            </span>
            <span class="multilang" lang="no">
                Den deriverte er \(y' = -\frac{{@d1@}}{2} \cdot \frac{{@d1@}}{2} \cdot t^{\frac{{@d1@}}{2} - 1} = \displaystyle{@ta4@}\).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \( y= {@f5@}\), we use the \(\ln kx\) rule, which gives \( y' =\displaystyle {@ta5@} .\)
            </span>
            <span class="multilang" lang="no">
                For \( y= {@f5@}\), we use the \(\ln kx\) rule, which gives \( y' =\displaystyle \frac{1}{x}.\)
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1:rand([5,7,9]);
f1: a1/t^3;
ta1:diff(f1,t);

b:rand([3,4,5]);
f2: sqrt(t^3);
ta2:diff(f2,t);

c1:rand([4,6,8]);
f3: c1*t^(-2);
ta3:diff(f3,t);

d1:rand([3,5]);
f4: -d1/2*t^(d1/2);
ta4:diff(f4,t);

e1:rand([3,5]);
f5: ln(e1*t);
ta5:diff(f5,t);

</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[ z={@f1@}, z'={@ta1@}, \quad z= {@f2@}, z'={@ta2@}, \quad z= {@f3@}, z'={@ta3@}\] 
\[ z= {@f4@}, z'= {@ta4@}, \quad z={@f5@}, z'= {@ta5@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>ta5</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>diff</forbidwords>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>ta5</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>555319474</deployedseed>
        <deployedseed>64752438</deployedseed>
        <deployedseed>1351530464</deployedseed>
        <deployedseed>1160406212</deployedseed>
        <deployedseed>1748445226</deployedseed>
        <deployedseed>1990717505</deployedseed>
        <deployedseed>964121452</deployedseed>
        <deployedseed>805516359</deployedseed>
        <deployedseed>1182947299</deployedseed>
        <deployedseed>15656004</deployedseed>
        <deployedseed>893917385</deployedseed>
        <deployedseed>481786772</deployedseed>
        <deployedseed>1007656668</deployedseed>
        <deployedseed>1122411965</deployedseed>
        <deployedseed>136131629</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta5</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3+1</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4+1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta5+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>