<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>11.2.3.1 evaluate trig derivatives</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        &#216;velse
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Calculate the derivative of \(y(x) = {@f1@}\) when \(x = {@a1@}\)
                <b>
                    radians
                </b>
                to \(3\) decimal places.
            </span>
            <span class="multilang" lang="no">
                Beregn den deriverte av \(y(x) = {@f1@}\) n&#229;r \(x = {@a1@}\) radianer, til \(3\) desimaler.
            </span>
        </p>
        <p>
            \( \dfrac{dy}{dx}({@a1@}) =\) [[input:ans1]]
                                                [[validation:ans1]][[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Calculate the rate of change of \(i(t) = {@f2@}\) when \(t =\displaystyle {@b1@} \)
                <b>
                    radians
                </b>
            </span>
            <span class="multilang" lang="no">
                Regn ut raten av endring av \(i(t) = {@f2@}\) n&#229;r \(t =\displaystyle {@b1@} \)
                <b>
                    radianer
                </b>
                .
            </span>
        </p>
        <p>
            \( \dfrac{di}{dt}\left(\displaystyle {@b1@}\right) =\) [[input:ans2]]
                                                                [[validation:ans2]][[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Calculate the rate of change of \(F(t) = {@f3@}\) when \(t = {@c1@}\)
                <b>
                    radians
                </b>
                to \(3\) decimal places.
            </span>
            <span class="multilang" lang="no">
                Beregn endringsraten til \(F(t) = {@f3@}\) n&#229;r \(t = {@c1@}\)
                <b>
                    radianer
                </b>
                til 3 desimaler.
            </span>
        </p>
        <p>
            \( \dfrac{dF}{dt}({@c1@}) =\) [[input:ans3]]
                                                [[validation:ans3]][[feedback:prt3]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We have \( y(x) = {@f1@}\) so \[ \dfrac{dy}{dx} = {@df1@},\] which when evaluated at \(x = {@a1@}\) is \({@ta1@}. \)
            </span>
            <span class="multilang" lang="no">
                Vi har \(y(x) = {@f1@}\), s&#229; \[\dfrac{dy}{dx} = {@df1@}, \] og n&#229;r vi evaluerer dette ved \(x = {@a1@}\), blir svaret \({@ta1@}. \)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We have \(i(t) = {@f2@}\) so \[ \dfrac{di}{dt} = {@df2@},\] which when evaluated at \( t = \displaystyle {@b1@} \) is {@ta2@}.
            </span>
            <span class="multilang" lang="no">
                Vi har \(i(t) = {@f2@}\), s&#229; \[\dfrac{di}{dt} = {@df2@},\] som n&#229;r evaluert ved \( t = \displaystyle {@b1@} \) er {@ta2@}.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We have \(F(t) = {@f3@}\) so \[ \dfrac{dF}{dt}= {@df3@},\] which when evaluated at \(t = {@c1@}\) is {@ta3@}.
            </span>
            <span class="multilang" lang="no">
                Vi har \(F(t) = {@f3@}\) s&#229; \[\dfrac{dF}{dt}= {@df3@},\] som n&#229;r evaluert ved \(t = {@c1@}\) er {@ta3@}.
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1:rand([0.1, 0.3, 0.4]);
f1: x^2 + sin(x);
df1: diff(f1,x);
ta1:decimalplaces(ev(df1,x:a1),3);

b1:rand([pi/3]);
f2: 4*sin(2*t) + 3*t;
df2: diff(f2,t);
ta2:ev(df2, t:b1);

c1:rand([1.3,1.5,1.7]);
f3: 5*sin(t) - 3*cos(2*t);
df3: diff(f3,t);
ta3:decimalplaces(ev(df3,t:c1),3);


</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\[ y(x) = {@f1@}, \quad y'({@a1@})={@ta1@}\] 
\[ i(t) = {@f2@}, \quad  i'({@b1@})={@ta2@} \]
\[ F(t) = {@f3@}, \quad F'({@c1@})={@ta3@} \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options>mindp:3, maxdp:3</options>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>numerical</type>
            <tans>ta3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options>mindp:3, maxdp:3</options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1856783899</deployedseed>
        <deployedseed>1228400075</deployedseed>
        <deployedseed>1760920483</deployedseed>
        <deployedseed>1843501837</deployedseed>
        <deployedseed>1263490896</deployedseed>
        <deployedseed>1749226452</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>


</quiz>